% Refutation found. Thanks to Tanya!
% SZS status Theorem for E075-29
% SZS output start Proof for E075-29
fof(f31842,plain,(
  $false),
  inference(avatar_sat_refutation,[],[f1247,f2218,f2406,f2884,f3300,f3951,f8700,f8864,f8943,f9030,f10511,f10554,f31826])).
fof(f31826,plain,(
  ~spl29_226 | ~spl29_232 | ~spl29_274 | ~spl29_328),
  inference(avatar_contradiction_clause,[],[f31825])).
fof(f31825,plain,(
  $false | (~spl29_226 | ~spl29_232 | ~spl29_274 | ~spl29_328)),
  inference(subsumption_resolution,[],[f31824,f24441])).
fof(f24441,plain,(
  ( ! [X6,X8,X7,X9] : (para(X6,X7,X8,X9)) ) | (~spl29_226 | ~spl29_232 | ~spl29_274 | ~spl29_328)),
  inference(subsumption_resolution,[],[f24419,f24411])).
fof(f24411,plain,(
  ( ! [X10,X9] : (midp(sK20,X9,X10)) ) | (~spl29_226 | ~spl29_232 | ~spl29_274 | ~spl29_328)),
  inference(subsumption_resolution,[],[f24410,f12824])).
fof(f12824,plain,(
  ( ! [X1] : (midp(sK25,X1,sK20)) ) | (~spl29_226 | ~spl29_232 | ~spl29_274)),
  inference(subsumption_resolution,[],[f12823,f12790])).
fof(f12790,plain,(
  ( ! [X3] : (midp(sK22,X3,sK23)) ) | (~spl29_226 | ~spl29_232 | ~spl29_274)),
  inference(subsumption_resolution,[],[f9153,f10174])).
fof(f10174,plain,(
  ( ! [X33,X34] : (midp(X33,X34,X33)) ) | (~spl29_226 | ~spl29_274)),
  inference(resolution,[],[f9954,f7763])).
fof(f7763,plain,(
  ( ! [X28,X29,X27] : (~midp(X27,X28,X29) | midp(X28,X27,X29)) ) | (~spl29_226 | ~spl29_274)),
  inference(subsumption_resolution,[],[f7725,f4153])).
fof(f4153,plain,(
  ( ! [X2,X0,X1] : (coll(X0,X1,X2)) ) | ~spl29_226),
  inference(subsumption_resolution,[],[f4150,f4138])).
fof(f4138,plain,(
  ( ! [X6,X5] : (coll(X5,X5,X6)) ) | ~spl29_226),
  inference(resolution,[],[f4121,f359])).
fof(f359,plain,(
  ( ! [X2,X0,X1] : (~coll(X0,X1,X2) | coll(X0,X2,X1)) )),
  inference(cnf_transformation,[],[f168])).
fof(f168,plain,(
  ! [X0,X1,X2] : (coll(X0,X2,X1) | ~coll(X0,X1,X2))),
  inference(ennf_transformation,[],[f1])).
fof(f1,axiom,(
  ! [X0,X1,X2] : (coll(X0,X1,X2) => coll(X0,X2,X1))),
  file('./7_Book/00EE/11/E075-29.p',ruleD1)).
fof(f4121,plain,(
  ( ! [X4,X3] : (coll(X3,X4,X3)) ) | ~spl29_226),
  inference(resolution,[],[f4103,f360])).
fof(f360,plain,(
  ( ! [X2,X0,X1] : (~coll(X0,X1,X2) | coll(X1,X0,X2)) )),
  inference(cnf_transformation,[],[f169])).
fof(f169,plain,(
  ! [X0,X1,X2] : (coll(X1,X0,X2) | ~coll(X0,X1,X2))),
  inference(ennf_transformation,[],[f2])).
fof(f2,axiom,(
  ! [X0,X1,X2] : (coll(X0,X1,X2) => coll(X1,X0,X2))),
  file('./7_Book/00EE/11/E075-29.p',ruleD2)).
fof(f4103,plain,(
  ( ! [X4,X3] : (coll(X3,X4,X4)) ) | ~spl29_226),
  inference(resolution,[],[f2214,f425])).
fof(f425,plain,(
  ( ! [X2,X0,X1] : (~para(X0,X1,X0,X2) | coll(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f267])).
fof(f267,plain,(
  ! [X0,X1,X2] : (coll(X0,X1,X2) | ~para(X0,X1,X0,X2))),
  inference(ennf_transformation,[],[f67])).
fof(f67,axiom,(
  ! [X0,X1,X2] : (para(X0,X1,X0,X2) => coll(X0,X1,X2))),
  file('./7_Book/00EE/11/E075-29.p',ruleD66)).
fof(f2214,plain,(
  ( ! [X2,X3] : (para(X2,X3,X2,X3)) ) | ~spl29_226),
  inference(avatar_component_clause,[],[f2213])).
fof(f2213,plain,(
  spl29_226 <=> ! [X3,X2] : para(X2,X3,X2,X3)),
  introduced(avatar_definition,[new_symbols(naming,[spl29_226])])).
fof(f4150,plain,(
  ( ! [X2,X0,X1] : (coll(X0,X1,X2) | ~coll(X2,X2,X0)) ) | ~spl29_226),
  inference(resolution,[],[f4138,f361])).
fof(f361,plain,(
  ( ! [X2,X0,X3,X1] : (~coll(X0,X1,X3) | coll(X2,X3,X0) | ~coll(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f171])).
fof(f171,plain,(
  ! [X0,X1,X2,X3] : (coll(X2,X3,X0) | ~coll(X0,X1,X3) | ~coll(X0,X1,X2))),
  inference(flattening,[],[f170])).
fof(f170,plain,(
  ! [X0,X1,X2,X3] : (coll(X2,X3,X0) | (~coll(X0,X1,X3) | ~coll(X0,X1,X2)))),
  inference(ennf_transformation,[],[f3])).
fof(f3,axiom,(
  ! [X0,X1,X2,X3] : ((coll(X0,X1,X3) & coll(X0,X1,X2)) => coll(X2,X3,X0))),
  file('./7_Book/00EE/11/E075-29.p',ruleD3)).
fof(f7725,plain,(
  ( ! [X28,X29,X27] : (~midp(X27,X28,X29) | ~coll(X28,X27,X29) | midp(X28,X27,X29)) ) | ~spl29_274),
  inference(resolution,[],[f2879,f426])).
fof(f426,plain,(
  ( ! [X2,X0,X1] : (~cong(X0,X1,X0,X2) | ~coll(X0,X1,X2) | midp(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f269])).
fof(f269,plain,(
  ! [X0,X1,X2] : (midp(X0,X1,X2) | ~coll(X0,X1,X2) | ~cong(X0,X1,X0,X2))),
  inference(flattening,[],[f268])).
fof(f268,plain,(
  ! [X0,X1,X2] : (midp(X0,X1,X2) | (~coll(X0,X1,X2) | ~cong(X0,X1,X0,X2)))),
  inference(ennf_transformation,[],[f68])).
fof(f68,axiom,(
  ! [X0,X1,X2] : ((coll(X0,X1,X2) & cong(X0,X1,X0,X2)) => midp(X0,X1,X2))),
  file('./7_Book/00EE/11/E075-29.p',ruleD67)).
fof(f2879,plain,(
  ( ! [X14,X15,X13] : (cong(X14,X13,X14,X15) | ~midp(X13,X14,X15)) ) | ~spl29_274),
  inference(avatar_component_clause,[],[f2878])).
fof(f2878,plain,(
  spl29_274 <=> ! [X13,X15,X14] : (~midp(X13,X14,X15) | cong(X14,X13,X14,X15))),
  introduced(avatar_definition,[new_symbols(naming,[spl29_274])])).
fof(f9954,plain,(
  ( ! [X28,X27] : (midp(X27,X28,X28)) ) | ~spl29_226),
  inference(resolution,[],[f9882,f4116])).
fof(f4116,plain,(
  ( ! [X12,X13,X11] : (~midp(X11,X13,X13) | midp(X11,X12,X12)) ) | ~spl29_226),
  inference(subsumption_resolution,[],[f4106,f2214])).
fof(f4106,plain,(
  ( ! [X12,X13,X11] : (midp(X11,X12,X12) | ~para(X13,X12,X13,X12) | ~midp(X11,X13,X13)) ) | ~spl29_226),
  inference(resolution,[],[f2214,f423])).
fof(f423,plain,(
  ( ! [X4,X2,X0,X3,X1] : (~para(X0,X3,X1,X2) | midp(X4,X2,X3) | ~para(X0,X2,X1,X3) | ~midp(X4,X0,X1)) )),
  inference(cnf_transformation,[],[f264])).
fof(f264,plain,(
  ! [X0,X1,X2,X3,X4] : (midp(X4,X2,X3) | ~para(X0,X3,X1,X2) | ~para(X0,X2,X1,X3) | ~midp(X4,X0,X1))),
  inference(flattening,[],[f263])).
fof(f263,plain,(
  ! [X0,X1,X2,X3,X4] : (midp(X4,X2,X3) | (~para(X0,X3,X1,X2) | ~para(X0,X2,X1,X3) | ~midp(X4,X0,X1)))),
  inference(ennf_transformation,[],[f143])).
fof(f143,plain,(
  ! [X0,X1,X2,X3,X4] : ((para(X0,X3,X1,X2) & para(X0,X2,X1,X3) & midp(X4,X0,X1)) => midp(X4,X2,X3))),
  inference(rectify,[],[f65])).
fof(f65,axiom,(
  ! [X0,X1,X2,X3,X6] : ((para(X0,X3,X1,X2) & para(X0,X2,X1,X3) & midp(X6,X0,X1)) => midp(X6,X2,X3))),
  file('./7_Book/00EE/11/E075-29.p',ruleD64)).
fof(f9882,plain,(
  ( ! [X14] : (midp(X14,X14,X14)) ) | ~spl29_226),
  inference(subsumption_resolution,[],[f9861,f4103])).
fof(f9861,plain,(
  ( ! [X14] : (~coll(X14,X14,X14) | midp(X14,X14,X14)) ) | ~spl29_226),
  inference(resolution,[],[f8924,f426])).
fof(f8924,plain,(
  ( ! [X35] : (cong(X35,X35,X35,X35)) ) | ~spl29_226),
  inference(subsumption_resolution,[],[f8920,f8835])).
fof(f8835,plain,(
  ( ! [X2,X0,X1] : (cyclic(X0,X0,X1,X2)) ) | ~spl29_226),
  inference(resolution,[],[f3339,f2214])).
fof(f3339,plain,(
  ( ! [X68,X66,X67] : (~para(X66,X67,X66,X67) | cyclic(X67,X67,X66,X68)) )),
  inference(subsumption_resolution,[],[f2389,f2390])).
fof(f2390,plain,(
  ( ! [X70,X71,X69] : (~para(X69,X70,X69,X70) | coll(X69,X71,X70) | cyclic(X70,X70,X69,X71)) )),
  inference(resolution,[],[f923,f400])).
fof(f400,plain,(
  ( ! [X2,X0,X3,X1] : (~eqangle(X2,X0,X2,X1,X3,X0,X3,X1) | coll(X2,X3,X0) | cyclic(X0,X1,X2,X3)) )),
  inference(cnf_transformation,[],[f222])).
fof(f222,plain,(
  ! [X0,X1,X2,X3] : (cyclic(X0,X1,X2,X3) | coll(X2,X3,X0) | ~eqangle(X2,X0,X2,X1,X3,X0,X3,X1))),
  inference(flattening,[],[f221])).
fof(f221,plain,(
  ! [X0,X1,X2,X3] : (cyclic(X0,X1,X2,X3) | (coll(X2,X3,X0) | ~eqangle(X2,X0,X2,X1,X3,X0,X3,X1)))),
  inference(ennf_transformation,[],[f121])).
fof(f121,plain,(
  ! [X0,X1,X2,X3] : ((~coll(X2,X3,X0) & eqangle(X2,X0,X2,X1,X3,X0,X3,X1)) => cyclic(X0,X1,X2,X3))),
  inference(rectify,[],[f42])).
fof(f42,axiom,(
  ! [X0,X1,X8,X9] : ((~coll(X8,X9,X0) & eqangle(X8,X0,X8,X1,X9,X0,X9,X1)) => cyclic(X0,X1,X8,X9))),
  file('./7_Book/00EE/11/E075-29.p',ruleD42a)).
fof(f923,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (eqangle(X0,X1,X2,X3,X4,X5,X4,X5) | ~para(X0,X1,X2,X3)) )),
  inference(resolution,[],[f379,f398])).
fof(f398,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (eqangle(X0,X1,X4,X5,X2,X3,X4,X5) | ~para(X0,X1,X2,X3)) )),
  inference(cnf_transformation,[],[f219])).
fof(f219,plain,(
  ! [X0,X1,X2,X3,X4,X5] : (eqangle(X0,X1,X4,X5,X2,X3,X4,X5) | ~para(X0,X1,X2,X3))),
  inference(ennf_transformation,[],[f119])).
fof(f119,plain,(
  ! [X0,X1,X2,X3,X4,X5] : (para(X0,X1,X2,X3) => eqangle(X0,X1,X4,X5,X2,X3,X4,X5))),
  inference(rectify,[],[f40])).
fof(f40,axiom,(
  ! [X0,X1,X2,X3,X8,X9] : (para(X0,X1,X2,X3) => eqangle(X0,X1,X8,X9,X2,X3,X8,X9))),
  file('./7_Book/00EE/11/E075-29.p',ruleD40)).
fof(f379,plain,(
  ( ! [X6,X4,X2,X0,X7,X5,X3,X1] : (~eqangle(X0,X1,X2,X3,X4,X5,X6,X7) | eqangle(X0,X1,X4,X5,X2,X3,X6,X7)) )),
  inference(cnf_transformation,[],[f195])).
fof(f195,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6,X7] : (eqangle(X0,X1,X4,X5,X2,X3,X6,X7) | ~eqangle(X0,X1,X2,X3,X4,X5,X6,X7))),
  inference(ennf_transformation,[],[f103])).
fof(f103,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6,X7] : (eqangle(X0,X1,X2,X3,X4,X5,X6,X7) => eqangle(X0,X1,X4,X5,X2,X3,X6,X7))),
  inference(rectify,[],[f21])).
fof(f21,axiom,(
  ! [X0,X1,X2,X3,X8,X9,X10,X11] : (eqangle(X0,X1,X2,X3,X8,X9,X10,X11) => eqangle(X0,X1,X8,X9,X2,X3,X10,X11))),
  file('./7_Book/00EE/11/E075-29.p',ruleD21)).
fof(f2389,plain,(
  ( ! [X68,X66,X67] : (~para(X66,X67,X66,X67) | ~coll(X66,X68,X67) | cyclic(X67,X67,X66,X68)) )),
  inference(resolution,[],[f923,f401])).
fof(f401,plain,(
  ( ! [X2,X0,X3,X1] : (~eqangle(X2,X0,X2,X1,X3,X0,X3,X1) | ~coll(X2,X3,X1) | cyclic(X0,X1,X2,X3)) )),
  inference(cnf_transformation,[],[f224])).
fof(f224,plain,(
  ! [X0,X1,X2,X3] : (cyclic(X0,X1,X2,X3) | ~coll(X2,X3,X1) | ~eqangle(X2,X0,X2,X1,X3,X0,X3,X1))),
  inference(flattening,[],[f223])).
fof(f223,plain,(
  ! [X0,X1,X2,X3] : (cyclic(X0,X1,X2,X3) | (~coll(X2,X3,X1) | ~eqangle(X2,X0,X2,X1,X3,X0,X3,X1)))),
  inference(ennf_transformation,[],[f122])).
fof(f122,plain,(
  ! [X0,X1,X2,X3] : ((coll(X2,X3,X1) & eqangle(X2,X0,X2,X1,X3,X0,X3,X1)) => cyclic(X0,X1,X2,X3))),
  inference(rectify,[],[f43])).
fof(f43,axiom,(
  ! [X0,X1,X8,X9] : ((coll(X8,X9,X1) & eqangle(X8,X0,X8,X1,X9,X0,X9,X1)) => cyclic(X0,X1,X8,X9))),
  file('./7_Book/00EE/11/E075-29.p',ruleD42b)).
fof(f8920,plain,(
  ( ! [X35,X36] : (cong(X35,X35,X35,X35) | ~cyclic(X35,X35,X36,X35)) ) | ~spl29_226),
  inference(duplicate_literal_removal,[],[f8907])).
fof(f8907,plain,(
  ( ! [X35,X36] : (cong(X35,X35,X35,X35) | ~cyclic(X35,X35,X36,X35) | ~cyclic(X35,X35,X36,X35)) ) | ~spl29_226),
  inference(resolution,[],[f8835,f1530])).
fof(f1530,plain,(
  ( ! [X10,X8,X7,X9] : (~cyclic(X7,X8,X9,X10) | cong(X7,X8,X7,X8) | ~cyclic(X7,X8,X9,X8) | ~cyclic(X7,X8,X9,X7)) )),
  inference(duplicate_literal_removal,[],[f1524])).
fof(f1524,plain,(
  ( ! [X10,X8,X7,X9] : (cong(X7,X8,X7,X8) | ~cyclic(X7,X8,X9,X10) | ~cyclic(X7,X8,X9,X8) | ~cyclic(X7,X8,X9,X7) | ~cyclic(X7,X8,X9,X10)) )),
  inference(resolution,[],[f402,f399])).
fof(f399,plain,(
  ( ! [X2,X0,X3,X1] : (eqangle(X2,X0,X2,X1,X3,X0,X3,X1) | ~cyclic(X0,X1,X2,X3)) )),
  inference(cnf_transformation,[],[f220])).
fof(f220,plain,(
  ! [X0,X1,X2,X3] : (eqangle(X2,X0,X2,X1,X3,X0,X3,X1) | ~cyclic(X0,X1,X2,X3))),
  inference(ennf_transformation,[],[f120])).
fof(f120,plain,(
  ! [X0,X1,X2,X3] : (cyclic(X0,X1,X2,X3) => eqangle(X2,X0,X2,X1,X3,X0,X3,X1))),
  inference(rectify,[],[f41])).
fof(f41,axiom,(
  ! [X0,X1,X8,X9] : (cyclic(X0,X1,X8,X9) => eqangle(X8,X0,X8,X1,X9,X0,X9,X1))),
  file('./7_Book/00EE/11/E075-29.p',ruleD41)).
fof(f402,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (~eqangle(X2,X0,X2,X1,X5,X3,X5,X4) | cong(X0,X1,X3,X4) | ~cyclic(X0,X1,X2,X5) | ~cyclic(X0,X1,X2,X4) | ~cyclic(X0,X1,X2,X3)) )),
  inference(cnf_transformation,[],[f226])).
fof(f226,plain,(
  ! [X0,X1,X2,X3,X4,X5] : (cong(X0,X1,X3,X4) | ~eqangle(X2,X0,X2,X1,X5,X3,X5,X4) | ~cyclic(X0,X1,X2,X5) | ~cyclic(X0,X1,X2,X4) | ~cyclic(X0,X1,X2,X3))),
  inference(flattening,[],[f225])).
fof(f225,plain,(
  ! [X0,X1,X2,X3,X4,X5] : (cong(X0,X1,X3,X4) | (~eqangle(X2,X0,X2,X1,X5,X3,X5,X4) | ~cyclic(X0,X1,X2,X5) | ~cyclic(X0,X1,X2,X4) | ~cyclic(X0,X1,X2,X3)))),
  inference(ennf_transformation,[],[f123])).
fof(f123,plain,(
  ! [X0,X1,X2,X3,X4,X5] : ((eqangle(X2,X0,X2,X1,X5,X3,X5,X4) & cyclic(X0,X1,X2,X5) & cyclic(X0,X1,X2,X4) & cyclic(X0,X1,X2,X3)) => cong(X0,X1,X3,X4))),
  inference(rectify,[],[f44])).
fof(f44,axiom,(
  ! [X0,X1,X2,X8,X9,X14] : ((eqangle(X2,X0,X2,X1,X14,X8,X14,X9) & cyclic(X0,X1,X2,X14) & cyclic(X0,X1,X2,X9) & cyclic(X0,X1,X2,X8)) => cong(X0,X1,X8,X9))),
  file('./7_Book/00EE/11/E075-29.p',ruleD43)).
fof(f9153,plain,(
  ( ! [X3] : (midp(sK22,X3,sK23) | ~midp(sK25,X3,sK25)) ) | (~spl29_226 | ~spl29_232)),
  inference(subsumption_resolution,[],[f9149,f4153])).
fof(f9149,plain,(
  ( ! [X3] : (~coll(sK22,X3,sK23) | midp(sK22,X3,sK23) | ~midp(sK25,X3,sK25)) ) | ~spl29_232),
  inference(resolution,[],[f9018,f404])).
fof(f404,plain,(
  ( ! [X4,X2,X0,X3,X1] : (~para(X3,X4,X1,X2) | ~coll(X4,X0,X2) | midp(X4,X0,X2) | ~midp(X3,X0,X1)) )),
  inference(cnf_transformation,[],[f230])).
fof(f230,plain,(
  ! [X0,X1,X2,X3,X4] : (midp(X4,X0,X2) | ~coll(X4,X0,X2) | ~para(X3,X4,X1,X2) | ~midp(X3,X0,X1))),
  inference(flattening,[],[f229])).
fof(f229,plain,(
  ! [X0,X1,X2,X3,X4] : (midp(X4,X0,X2) | (~coll(X4,X0,X2) | ~para(X3,X4,X1,X2) | ~midp(X3,X0,X1)))),
  inference(ennf_transformation,[],[f125])).
fof(f125,plain,(
  ! [X0,X1,X2,X3,X4] : ((coll(X4,X0,X2) & para(X3,X4,X1,X2) & midp(X3,X0,X1)) => midp(X4,X0,X2))),
  inference(rectify,[],[f46])).
fof(f46,axiom,(
  ! [X0,X1,X2,X4,X5] : ((coll(X5,X0,X2) & para(X4,X5,X1,X2) & midp(X4,X0,X1)) => midp(X5,X0,X2))),
  file('./7_Book/00EE/11/E075-29.p',ruleD45)).
fof(f9018,plain,(
  para(sK25,sK22,sK25,sK23) | ~spl29_232),
  inference(resolution,[],[f9004,f870])).
fof(f870,plain,(
  ( ! [X10,X11] : (~perp(X10,X11,sK20,sK23) | para(X10,X11,sK25,sK23)) )),
  inference(resolution,[],[f367,f512])).
fof(f512,plain,(
  perp(sK20,sK23,sK25,sK23)),
  inference(resolution,[],[f366,f476])).
fof(f476,plain,(
  perp(sK25,sK23,sK20,sK23)),
  inference(cnf_transformation,[],[f358])).
fof(f358,plain,(
  ~eqangle(sK22,sK25,sK25,sK26,sK25,sK26,sK26,sK22) & perp(sK26,sK24,sK20,sK24) & circle(sK24,sK20,sK26,sK28) & perp(sK20,sK21,sK21,sK26) & perp(sK25,sK23,sK25,sK22) & perp(sK25,sK23,sK20,sK23) & circle(sK23,sK20,sK25,sK27) & midp(sK24,sK22,sK20) & midp(sK23,sK21,sK20) & coll(sK22,sK20,sK21)),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK20,sK21,sK22,sK23,sK24,sK25,sK26,sK27,sK28])],[f320,f357])).
fof(f357,plain,(
  ? [X0,X1,X2,X3,X4,X5,X6,X7,X8] : (~eqangle(X2,X5,X5,X6,X5,X6,X6,X2) & perp(X6,X4,X0,X4) & circle(X4,X0,X6,X8) & perp(X0,X1,X1,X6) & perp(X5,X3,X5,X2) & perp(X5,X3,X0,X3) & circle(X3,X0,X5,X7) & midp(X4,X2,X0) & midp(X3,X1,X0) & coll(X2,X0,X1)) => (~eqangle(sK22,sK25,sK25,sK26,sK25,sK26,sK26,sK22) & perp(sK26,sK24,sK20,sK24) & circle(sK24,sK20,sK26,sK28) & perp(sK20,sK21,sK21,sK26) & perp(sK25,sK23,sK25,sK22) & perp(sK25,sK23,sK20,sK23) & circle(sK23,sK20,sK25,sK27) & midp(sK24,sK22,sK20) & midp(sK23,sK21,sK20) & coll(sK22,sK20,sK21))),
  introduced(choice_axiom,[])).
fof(f320,plain,(
  ? [X0,X1,X2,X3,X4,X5,X6,X7,X8] : (~eqangle(X2,X5,X5,X6,X5,X6,X6,X2) & perp(X6,X4,X0,X4) & circle(X4,X0,X6,X8) & perp(X0,X1,X1,X6) & perp(X5,X3,X5,X2) & perp(X5,X3,X0,X3) & circle(X3,X0,X5,X7) & midp(X4,X2,X0) & midp(X3,X1,X0) & coll(X2,X0,X1))),
  inference(flattening,[],[f319])).
fof(f319,plain,(
  ? [X0,X1,X2,X3,X4,X5,X6,X7,X8] : (~eqangle(X2,X5,X5,X6,X5,X6,X6,X2) & (perp(X6,X4,X0,X4) & circle(X4,X0,X6,X8) & perp(X0,X1,X1,X6) & perp(X5,X3,X5,X2) & perp(X5,X3,X0,X3) & circle(X3,X0,X5,X7) & midp(X4,X2,X0) & midp(X3,X1,X0) & coll(X2,X0,X1)))),
  inference(ennf_transformation,[],[f167])).
fof(f167,plain,(
  ~! [X0,X1,X2,X3,X4,X5,X6,X7,X8] : ((perp(X6,X4,X0,X4) & circle(X4,X0,X6,X8) & perp(X0,X1,X1,X6) & perp(X5,X3,X5,X2) & perp(X5,X3,X0,X3) & circle(X3,X0,X5,X7) & midp(X4,X2,X0) & midp(X3,X1,X0) & coll(X2,X0,X1)) => eqangle(X2,X5,X5,X6,X5,X6,X6,X2))),
  inference(rectify,[],[f96])).
fof(f96,negated_conjecture,(
  ~! [X0,X1,X2,X3,X4,X5,X12,X18,X19] : ((perp(X12,X4,X0,X4) & circle(X4,X0,X12,X19) & perp(X0,X1,X1,X12) & perp(X5,X3,X5,X2) & perp(X5,X3,X0,X3) & circle(X3,X0,X5,X18) & midp(X4,X2,X0) & midp(X3,X1,X0) & coll(X2,X0,X1)) => eqangle(X2,X5,X5,X12,X5,X12,X12,X2))),
  inference(negated_conjecture,[],[f95])).
fof(f95,conjecture,(
  ! [X0,X1,X2,X3,X4,X5,X12,X18,X19] : ((perp(X12,X4,X0,X4) & circle(X4,X0,X12,X19) & perp(X0,X1,X1,X12) & perp(X5,X3,X5,X2) & perp(X5,X3,X0,X3) & circle(X3,X0,X5,X18) & midp(X4,X2,X0) & midp(X3,X1,X0) & coll(X2,X0,X1)) => eqangle(X2,X5,X5,X12,X5,X12,X12,X2))),
  file('./7_Book/00EE/11/E075-29.p',exemplo7Book00EE11E07529)).
fof(f366,plain,(
  ( ! [X2,X0,X3,X1] : (~perp(X0,X1,X2,X3) | perp(X2,X3,X0,X1)) )),
  inference(cnf_transformation,[],[f177])).
fof(f177,plain,(
  ! [X0,X1,X2,X3] : (perp(X2,X3,X0,X1) | ~perp(X0,X1,X2,X3))),
  inference(ennf_transformation,[],[f8])).
fof(f8,axiom,(
  ! [X0,X1,X2,X3] : (perp(X0,X1,X2,X3) => perp(X2,X3,X0,X1))),
  file('./7_Book/00EE/11/E075-29.p',ruleD8)).
fof(f367,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (~perp(X2,X3,X4,X5) | para(X0,X1,X4,X5) | ~perp(X0,X1,X2,X3)) )),
  inference(cnf_transformation,[],[f179])).
fof(f179,plain,(
  ! [X0,X1,X2,X3,X4,X5] : (para(X0,X1,X4,X5) | ~perp(X2,X3,X4,X5) | ~perp(X0,X1,X2,X3))),
  inference(flattening,[],[f178])).
fof(f178,plain,(
  ! [X0,X1,X2,X3,X4,X5] : (para(X0,X1,X4,X5) | (~perp(X2,X3,X4,X5) | ~perp(X0,X1,X2,X3)))),
  inference(ennf_transformation,[],[f9])).
fof(f9,axiom,(
  ! [X0,X1,X2,X3,X4,X5] : ((perp(X2,X3,X4,X5) & perp(X0,X1,X2,X3)) => para(X0,X1,X4,X5))),
  file('./7_Book/00EE/11/E075-29.p',ruleD9)).
fof(f9004,plain,(
  perp(sK25,sK22,sK20,sK23) | ~spl29_232),
  inference(resolution,[],[f2405,f1667])).
fof(f1667,plain,(
  para(sK25,sK22,sK20,sK23)),
  inference(resolution,[],[f877,f513])).
fof(f513,plain,(
  perp(sK25,sK22,sK25,sK23)),
  inference(resolution,[],[f366,f477])).
fof(f477,plain,(
  perp(sK25,sK23,sK25,sK22)),
  inference(cnf_transformation,[],[f358])).
fof(f877,plain,(
  ( ! [X24,X25] : (~perp(X24,X25,sK25,sK23) | para(X24,X25,sK20,sK23)) )),
  inference(resolution,[],[f367,f476])).
fof(f2405,plain,(
  ( ! [X47,X45,X46,X44] : (~para(X44,X45,X46,X47) | perp(X44,X45,X46,X47)) ) | ~spl29_232),
  inference(avatar_component_clause,[],[f2404])).
fof(f2404,plain,(
  spl29_232 <=> ! [X44,X46,X45,X47] : (~para(X44,X45,X46,X47) | perp(X44,X45,X46,X47))),
  introduced(avatar_definition,[new_symbols(naming,[spl29_232])])).
fof(f12823,plain,(
  ( ! [X1] : (midp(sK25,X1,sK20) | ~midp(sK22,X1,sK23)) ) | ~spl29_226),
  inference(subsumption_resolution,[],[f3143,f4153])).
fof(f3143,plain,(
  ( ! [X1] : (~coll(sK25,X1,sK20) | midp(sK25,X1,sK20) | ~midp(sK22,X1,sK23)) )),
  inference(resolution,[],[f2624,f404])).
fof(f2624,plain,(
  para(sK22,sK25,sK23,sK20)),
  inference(resolution,[],[f1665,f362])).
fof(f362,plain,(
  ( ! [X2,X0,X3,X1] : (~para(X0,X1,X2,X3) | para(X0,X1,X3,X2)) )),
  inference(cnf_transformation,[],[f172])).
fof(f172,plain,(
  ! [X0,X1,X2,X3] : (para(X0,X1,X3,X2) | ~para(X0,X1,X2,X3))),
  inference(ennf_transformation,[],[f4])).
fof(f4,axiom,(
  ! [X0,X1,X2,X3] : (para(X0,X1,X2,X3) => para(X0,X1,X3,X2))),
  file('./7_Book/00EE/11/E075-29.p',ruleD4)).
fof(f1665,plain,(
  para(sK22,sK25,sK20,sK23)),
  inference(resolution,[],[f877,f565])).
fof(f565,plain,(
  perp(sK22,sK25,sK25,sK23)),
  inference(resolution,[],[f507,f366])).
fof(f507,plain,(
  perp(sK25,sK23,sK22,sK25)),
  inference(resolution,[],[f365,f477])).
fof(f365,plain,(
  ( ! [X2,X0,X3,X1] : (~perp(X0,X1,X2,X3) | perp(X0,X1,X3,X2)) )),
  inference(cnf_transformation,[],[f176])).
fof(f176,plain,(
  ! [X0,X1,X2,X3] : (perp(X0,X1,X3,X2) | ~perp(X0,X1,X2,X3))),
  inference(ennf_transformation,[],[f7])).
fof(f7,axiom,(
  ! [X0,X1,X2,X3] : (perp(X0,X1,X2,X3) => perp(X0,X1,X3,X2))),
  file('./7_Book/00EE/11/E075-29.p',ruleD7)).
fof(f24410,plain,(
  ( ! [X10,X9] : (midp(sK20,X9,X10) | ~midp(sK25,X9,sK20)) ) | (~spl29_226 | ~spl29_328)),
  inference(subsumption_resolution,[],[f24393,f4153])).
fof(f24393,plain,(
  ( ! [X10,X9] : (~coll(sK20,X9,X10) | midp(sK20,X9,X10) | ~midp(sK25,X9,sK20)) ) | (~spl29_226 | ~spl29_328)),
  inference(resolution,[],[f15437,f404])).
fof(f15437,plain,(
  ( ! [X4] : (para(sK25,sK20,sK20,X4)) ) | (~spl29_226 | ~spl29_328)),
  inference(resolution,[],[f10635,f362])).
fof(f10635,plain,(
  ( ! [X0] : (para(sK25,sK20,X0,sK20)) ) | (~spl29_226 | ~spl29_328)),
  inference(subsumption_resolution,[],[f10627,f10249])).
fof(f10249,plain,(
  ( ! [X14] : (midp(sK20,sK21,X14)) ) | ~spl29_226),
  inference(resolution,[],[f10028,f4210])).
fof(f4210,plain,(
  ( ! [X12,X10,X11] : (~midp(X12,X11,X10) | midp(X10,X11,X12)) ) | ~spl29_226),
  inference(subsumption_resolution,[],[f4201,f4153])).
fof(f4201,plain,(
  ( ! [X12,X10,X11] : (~coll(X10,X11,X12) | midp(X10,X11,X12) | ~midp(X12,X11,X10)) ) | ~spl29_226),
  inference(resolution,[],[f4110,f404])).
fof(f4110,plain,(
  ( ! [X24,X23] : (para(X23,X24,X24,X23)) ) | ~spl29_226),
  inference(resolution,[],[f2214,f362])).
fof(f10028,plain,(
  ( ! [X0] : (midp(X0,sK21,sK20)) ) | ~spl29_226),
  inference(resolution,[],[f9971,f369])).
fof(f369,plain,(
  ( ! [X2,X0,X1] : (~midp(X2,X1,X0) | midp(X2,X0,X1)) )),
  inference(cnf_transformation,[],[f182])).
fof(f182,plain,(
  ! [X0,X1,X2] : (midp(X2,X0,X1) | ~midp(X2,X1,X0))),
  inference(ennf_transformation,[],[f97])).
fof(f97,plain,(
  ! [X0,X1,X2] : (midp(X2,X1,X0) => midp(X2,X0,X1))),
  inference(rectify,[],[f11])).
fof(f11,axiom,(
  ! [X0,X1,X6] : (midp(X6,X1,X0) => midp(X6,X0,X1))),
  file('./7_Book/00EE/11/E075-29.p',ruleD11)).
fof(f9971,plain,(
  ( ! [X48] : (midp(X48,sK20,sK21)) ) | ~spl29_226),
  inference(subsumption_resolution,[],[f9966,f9954])).
fof(f9966,plain,(
  ( ! [X48] : (midp(X48,sK20,sK21) | ~midp(X48,sK23,sK23)) ) | ~spl29_226),
  inference(resolution,[],[f9882,f1598])).
fof(f1598,plain,(
  ( ! [X2,X3,X1] : (~midp(sK23,X1,X2) | midp(X3,sK20,sK21) | ~midp(X3,X1,X2)) )),
  inference(subsumption_resolution,[],[f1593,f560])).
fof(f560,plain,(
  ( ! [X4,X5] : (para(X4,sK20,X5,sK21) | ~midp(sK23,X4,X5)) )),
  inference(resolution,[],[f422,f488])).
fof(f488,plain,(
  midp(sK23,sK20,sK21)),
  inference(resolution,[],[f369,f473])).
fof(f473,plain,(
  midp(sK23,sK21,sK20)),
  inference(cnf_transformation,[],[f358])).
fof(f422,plain,(
  ( ! [X4,X2,X0,X3,X1] : (~midp(X4,X2,X3) | para(X0,X2,X1,X3) | ~midp(X4,X0,X1)) )),
  inference(cnf_transformation,[],[f262])).
fof(f262,plain,(
  ! [X0,X1,X2,X3,X4] : (para(X0,X2,X1,X3) | ~midp(X4,X2,X3) | ~midp(X4,X0,X1))),
  inference(flattening,[],[f261])).
fof(f261,plain,(
  ! [X0,X1,X2,X3,X4] : (para(X0,X2,X1,X3) | (~midp(X4,X2,X3) | ~midp(X4,X0,X1)))),
  inference(ennf_transformation,[],[f142])).
fof(f142,plain,(
  ! [X0,X1,X2,X3,X4] : ((midp(X4,X2,X3) & midp(X4,X0,X1)) => para(X0,X2,X1,X3))),
  inference(rectify,[],[f64])).
fof(f64,axiom,(
  ! [X0,X1,X2,X3,X6] : ((midp(X6,X2,X3) & midp(X6,X0,X1)) => para(X0,X2,X1,X3))),
  file('./7_Book/00EE/11/E075-29.p',ruleD63)).
fof(f1593,plain,(
  ( ! [X2,X3,X1] : (~midp(sK23,X1,X2) | midp(X3,sK20,sK21) | ~para(X1,sK20,X2,sK21) | ~midp(X3,X1,X2)) )),
  inference(resolution,[],[f558,f423])).
fof(f558,plain,(
  ( ! [X0,X1] : (para(X0,sK21,X1,sK20) | ~midp(sK23,X0,X1)) )),
  inference(resolution,[],[f422,f473])).
fof(f10627,plain,(
  ( ! [X0] : (para(sK25,sK20,X0,sK20) | ~midp(sK20,sK21,X0)) ) | ~spl29_328),
  inference(resolution,[],[f3401,f1576])).
fof(f1576,plain,(
  ( ! [X10,X8,X7,X9] : (~para(X9,X10,X7,sK23) | para(X9,X10,X8,sK20) | ~midp(X7,sK21,X8)) )),
  inference(resolution,[],[f551,f364])).
fof(f364,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (~para(X2,X3,X4,X5) | para(X0,X1,X4,X5) | ~para(X0,X1,X2,X3)) )),
  inference(cnf_transformation,[],[f175])).
fof(f175,plain,(
  ! [X0,X1,X2,X3,X4,X5] : (para(X0,X1,X4,X5) | ~para(X2,X3,X4,X5) | ~para(X0,X1,X2,X3))),
  inference(flattening,[],[f174])).
fof(f174,plain,(
  ! [X0,X1,X2,X3,X4,X5] : (para(X0,X1,X4,X5) | (~para(X2,X3,X4,X5) | ~para(X0,X1,X2,X3)))),
  inference(ennf_transformation,[],[f6])).
fof(f6,axiom,(
  ! [X0,X1,X2,X3,X4,X5] : ((para(X2,X3,X4,X5) & para(X0,X1,X2,X3)) => para(X0,X1,X4,X5))),
  file('./7_Book/00EE/11/E075-29.p',ruleD6)).
fof(f551,plain,(
  ( ! [X0,X1] : (para(X0,sK23,X1,sK20) | ~midp(X0,sK21,X1)) )),
  inference(resolution,[],[f403,f473])).
fof(f403,plain,(
  ( ! [X4,X2,X0,X3,X1] : (~midp(X4,X0,X2) | para(X3,X4,X1,X2) | ~midp(X3,X0,X1)) )),
  inference(cnf_transformation,[],[f228])).
fof(f228,plain,(
  ! [X0,X1,X2,X3,X4] : (para(X3,X4,X1,X2) | ~midp(X4,X0,X2) | ~midp(X3,X0,X1))),
  inference(flattening,[],[f227])).
fof(f227,plain,(
  ! [X0,X1,X2,X3,X4] : (para(X3,X4,X1,X2) | (~midp(X4,X0,X2) | ~midp(X3,X0,X1)))),
  inference(ennf_transformation,[],[f124])).
fof(f124,plain,(
  ! [X0,X1,X2,X3,X4] : ((midp(X4,X0,X2) & midp(X3,X0,X1)) => para(X3,X4,X1,X2))),
  inference(rectify,[],[f45])).
fof(f45,axiom,(
  ! [X0,X1,X2,X4,X5] : ((midp(X5,X0,X2) & midp(X4,X0,X1)) => para(X4,X5,X1,X2))),
  file('./7_Book/00EE/11/E075-29.p',ruleD44)).
fof(f3401,plain,(
  para(sK25,sK20,sK20,sK23) | ~spl29_328),
  inference(avatar_component_clause,[],[f3400])).
fof(f3400,plain,(
  spl29_328 <=> para(sK25,sK20,sK20,sK23)),
  introduced(avatar_definition,[new_symbols(naming,[spl29_328])])).
fof(f24419,plain,(
  ( ! [X6,X8,X7,X9] : (para(X6,X7,X8,X9) | ~midp(sK20,X6,X8)) ) | (~spl29_226 | ~spl29_232 | ~spl29_274 | ~spl29_328)),
  inference(resolution,[],[f24411,f422])).
fof(f31824,plain,(
  ~para(sK25,sK26,sK26,sK22) | (~spl29_226 | ~spl29_232 | ~spl29_274 | ~spl29_328)),
  inference(subsumption_resolution,[],[f31789,f24441])).
fof(f31789,plain,(
  ~para(sK22,sK25,sK25,sK26) | ~para(sK25,sK26,sK26,sK22)),
  inference(resolution,[],[f4616,f481])).
fof(f481,plain,(
  ~eqangle(sK22,sK25,sK25,sK26,sK25,sK26,sK26,sK22)),
  inference(cnf_transformation,[],[f358])).
fof(f4616,plain,(
  ( ! [X80,X87,X85,X83,X81,X86,X84,X82] : (eqangle(X80,X81,X82,X83,X84,X85,X86,X87) | ~para(X80,X81,X84,X85) | ~para(X82,X83,X86,X87)) )),
  inference(resolution,[],[f1456,f915])).
fof(f915,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (eqangle(X0,X1,X2,X3,X0,X1,X4,X5) | ~para(X2,X3,X4,X5)) )),
  inference(resolution,[],[f377,f398])).
fof(f377,plain,(
  ( ! [X6,X4,X2,X0,X7,X5,X3,X1] : (~eqangle(X0,X1,X2,X3,X4,X5,X6,X7) | eqangle(X2,X3,X0,X1,X6,X7,X4,X5)) )),
  inference(cnf_transformation,[],[f193])).
fof(f193,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6,X7] : (eqangle(X2,X3,X0,X1,X6,X7,X4,X5) | ~eqangle(X0,X1,X2,X3,X4,X5,X6,X7))),
  inference(ennf_transformation,[],[f101])).
fof(f101,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6,X7] : (eqangle(X0,X1,X2,X3,X4,X5,X6,X7) => eqangle(X2,X3,X0,X1,X6,X7,X4,X5))),
  inference(rectify,[],[f19])).
fof(f19,axiom,(
  ! [X0,X1,X2,X3,X8,X9,X10,X11] : (eqangle(X0,X1,X2,X3,X8,X9,X10,X11) => eqangle(X2,X3,X0,X1,X10,X11,X8,X9))),
  file('./7_Book/00EE/11/E075-29.p',ruleD19)).
fof(f1456,plain,(
  ( ! [X6,X4,X2,X0,X8,X7,X5,X3,X1,X9] : (~eqangle(X0,X1,X2,X3,X8,X9,X6,X7) | eqangle(X0,X1,X2,X3,X4,X5,X6,X7) | ~para(X8,X9,X4,X5)) )),
  inference(resolution,[],[f380,f398])).
fof(f380,plain,(
  ( ! [X6,X4,X2,X0,X10,X8,X7,X5,X3,X1,X11,X9] : (~eqangle(X4,X5,X6,X7,X8,X9,X10,X11) | eqangle(X0,X1,X2,X3,X8,X9,X10,X11) | ~eqangle(X0,X1,X2,X3,X4,X5,X6,X7)) )),
  inference(cnf_transformation,[],[f197])).
fof(f197,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11] : (eqangle(X0,X1,X2,X3,X8,X9,X10,X11) | ~eqangle(X4,X5,X6,X7,X8,X9,X10,X11) | ~eqangle(X0,X1,X2,X3,X4,X5,X6,X7))),
  inference(flattening,[],[f196])).
fof(f196,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11] : (eqangle(X0,X1,X2,X3,X8,X9,X10,X11) | (~eqangle(X4,X5,X6,X7,X8,X9,X10,X11) | ~eqangle(X0,X1,X2,X3,X4,X5,X6,X7)))),
  inference(ennf_transformation,[],[f104])).
fof(f104,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11] : ((eqangle(X4,X5,X6,X7,X8,X9,X10,X11) & eqangle(X0,X1,X2,X3,X4,X5,X6,X7)) => eqangle(X0,X1,X2,X3,X8,X9,X10,X11))),
  inference(rectify,[],[f22])).
fof(f22,axiom,(
  ! [X0,X1,X2,X3,X8,X9,X10,X11,X4,X5,X12,X13] : ((eqangle(X8,X9,X10,X11,X4,X5,X12,X13) & eqangle(X0,X1,X2,X3,X8,X9,X10,X11)) => eqangle(X0,X1,X2,X3,X4,X5,X12,X13))),
  file('./7_Book/00EE/11/E075-29.p',ruleD22)).
fof(f10554,plain,(
  spl29_328 | ~spl29_46 | ~spl29_232),
  inference(avatar_split_clause,[],[f10553,f2404,f769,f3400])).
fof(f769,plain,(
  spl29_46 <=> cong(sK25,sK20,sK25,sK23)),
  introduced(avatar_definition,[new_symbols(naming,[spl29_46])])).
fof(f10553,plain,(
  para(sK25,sK20,sK20,sK23) | (~spl29_46 | ~spl29_232)),
  inference(subsumption_resolution,[],[f10523,f9012])).
fof(f9012,plain,(
  para(sK20,sK23,sK25,sK23) | ~spl29_232),
  inference(resolution,[],[f8984,f881])).
fof(f881,plain,(
  ( ! [X33,X32] : (~perp(X32,X33,sK25,sK22) | para(X32,X33,sK25,sK23)) )),
  inference(resolution,[],[f367,f513])).
fof(f8984,plain,(
  perp(sK20,sK23,sK25,sK22) | ~spl29_232),
  inference(resolution,[],[f2405,f1674])).
fof(f1674,plain,(
  para(sK20,sK23,sK25,sK22)),
  inference(resolution,[],[f880,f512])).
fof(f880,plain,(
  ( ! [X30,X31] : (~perp(X30,X31,sK25,sK23) | para(X30,X31,sK25,sK22)) )),
  inference(resolution,[],[f367,f477])).
fof(f10523,plain,(
  para(sK25,sK20,sK20,sK23) | ~para(sK20,sK23,sK25,sK23) | ~spl29_46),
  inference(resolution,[],[f770,f1213])).
fof(f1213,plain,(
  ( ! [X21,X19,X20] : (~cong(X21,X19,X21,X20) | para(X21,X19,X19,X20) | ~para(X19,X20,X21,X20)) )),
  inference(resolution,[],[f432,f405])).
fof(f405,plain,(
  ( ! [X2,X0,X1] : (eqangle(X2,X0,X0,X1,X0,X1,X2,X1) | ~cong(X2,X0,X2,X1)) )),
  inference(cnf_transformation,[],[f231])).
fof(f231,plain,(
  ! [X0,X1,X2] : (eqangle(X2,X0,X0,X1,X0,X1,X2,X1) | ~cong(X2,X0,X2,X1))),
  inference(ennf_transformation,[],[f126])).
fof(f126,plain,(
  ! [X0,X1,X2] : (cong(X2,X0,X2,X1) => eqangle(X2,X0,X0,X1,X0,X1,X2,X1))),
  inference(rectify,[],[f47])).
fof(f47,axiom,(
  ! [X0,X1,X7] : (cong(X7,X0,X7,X1) => eqangle(X7,X0,X0,X1,X0,X1,X7,X1))),
  file('./7_Book/00EE/11/E075-29.p',ruleD46)).
fof(f432,plain,(
  ( ! [X6,X4,X2,X0,X7,X5,X3,X1] : (~eqangle(X0,X1,X2,X3,X4,X5,X6,X7) | ~para(X4,X5,X6,X7) | para(X0,X1,X2,X3)) )),
  inference(cnf_transformation,[],[f279])).
fof(f279,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6,X7] : (para(X0,X1,X2,X3) | ~para(X4,X5,X6,X7) | ~eqangle(X0,X1,X2,X3,X4,X5,X6,X7))),
  inference(flattening,[],[f278])).
fof(f278,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6,X7] : (para(X0,X1,X2,X3) | (~para(X4,X5,X6,X7) | ~eqangle(X0,X1,X2,X3,X4,X5,X6,X7)))),
  inference(ennf_transformation,[],[f146])).
fof(f146,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6,X7] : ((para(X4,X5,X6,X7) & eqangle(X0,X1,X2,X3,X4,X5,X6,X7)) => para(X0,X1,X2,X3))),
  inference(rectify,[],[f74])).
fof(f74,axiom,(
  ! [X0,X1,X2,X3,X8,X9,X10,X11] : ((para(X8,X9,X10,X11) & eqangle(X0,X1,X2,X3,X8,X9,X10,X11)) => para(X0,X1,X2,X3))),
  file('./7_Book/00EE/11/E075-29.p',ruleD73)).
fof(f770,plain,(
  cong(sK25,sK20,sK25,sK23) | ~spl29_46),
  inference(avatar_component_clause,[],[f769])).
fof(f10511,plain,(
  ~spl29_226 | spl29_545),
  inference(avatar_contradiction_clause,[],[f10493])).
fof(f10493,plain,(
  $false | (~spl29_226 | ~spl29_545)),
  inference(resolution,[],[f10110,f9029])).
fof(f9029,plain,(
  ~midp(sK22,sK20,sK23) | ~spl29_545),
  inference(avatar_component_clause,[],[f9028])).
fof(f9028,plain,(
  spl29_545 <=> ~midp(sK22,sK20,sK23)),
  introduced(avatar_definition,[new_symbols(naming,[spl29_545])])).
fof(f10110,plain,(
  ( ! [X14] : (midp(sK22,sK20,X14)) ) | ~spl29_226),
  inference(resolution,[],[f9972,f4210])).
fof(f9972,plain,(
  ( ! [X49] : (midp(X49,sK20,sK22)) ) | ~spl29_226),
  inference(subsumption_resolution,[],[f9968,f9954])).
fof(f9968,plain,(
  ( ! [X49] : (midp(X49,sK20,sK22) | ~midp(X49,sK24,sK24)) ) | ~spl29_226),
  inference(resolution,[],[f9882,f1608])).
fof(f1608,plain,(
  ( ! [X2,X3,X1] : (~midp(sK24,X1,X2) | midp(X3,sK20,sK22) | ~midp(X3,X1,X2)) )),
  inference(subsumption_resolution,[],[f1603,f561])).
fof(f561,plain,(
  ( ! [X6,X7] : (para(X6,sK20,X7,sK22) | ~midp(sK24,X6,X7)) )),
  inference(resolution,[],[f422,f489])).
fof(f489,plain,(
  midp(sK24,sK20,sK22)),
  inference(resolution,[],[f369,f474])).
fof(f474,plain,(
  midp(sK24,sK22,sK20)),
  inference(cnf_transformation,[],[f358])).
fof(f1603,plain,(
  ( ! [X2,X3,X1] : (~midp(sK24,X1,X2) | midp(X3,sK20,sK22) | ~para(X1,sK20,X2,sK22) | ~midp(X3,X1,X2)) )),
  inference(resolution,[],[f559,f423])).
fof(f559,plain,(
  ( ! [X2,X3] : (para(X2,sK22,X3,sK20) | ~midp(sK24,X2,X3)) )),
  inference(resolution,[],[f422,f474])).
fof(f9030,plain,(
  ~spl29_545 | spl29_46 | ~spl29_232),
  inference(avatar_split_clause,[],[f9021,f2404,f769,f9028])).
fof(f9021,plain,(
  cong(sK25,sK20,sK25,sK23) | ~midp(sK22,sK20,sK23) | ~spl29_232),
  inference(resolution,[],[f9004,f414])).
fof(f414,plain,(
  ( ! [X2,X0,X3,X1] : (~perp(X3,X2,X0,X1) | cong(X3,X0,X3,X1) | ~midp(X2,X0,X1)) )),
  inference(cnf_transformation,[],[f249])).
fof(f249,plain,(
  ! [X0,X1,X2,X3] : (cong(X3,X0,X3,X1) | ~perp(X3,X2,X0,X1) | ~midp(X2,X0,X1))),
  inference(flattening,[],[f248])).
fof(f248,plain,(
  ! [X0,X1,X2,X3] : (cong(X3,X0,X3,X1) | (~perp(X3,X2,X0,X1) | ~midp(X2,X0,X1)))),
  inference(ennf_transformation,[],[f134])).
fof(f134,plain,(
  ! [X0,X1,X2,X3] : ((perp(X3,X2,X0,X1) & midp(X2,X0,X1)) => cong(X3,X0,X3,X1))),
  inference(rectify,[],[f56])).
fof(f56,axiom,(
  ! [X0,X1,X6,X7] : ((perp(X7,X6,X0,X1) & midp(X6,X0,X1)) => cong(X7,X0,X7,X1))),
  file('./7_Book/00EE/11/E075-29.p',ruleD55)).
fof(f8943,plain,(
  ~spl29_230 | ~spl29_512),
  inference(avatar_contradiction_clause,[],[f8942])).
fof(f8942,plain,(
  $false | (~spl29_230 | ~spl29_512)),
  inference(subsumption_resolution,[],[f8940,f2402])).
fof(f2402,plain,(
  ( ! [X48,X49] : (~perp(X48,X49,X48,X49)) ) | ~spl29_230),
  inference(avatar_component_clause,[],[f2401])).
fof(f2401,plain,(
  spl29_230 <=> ! [X49,X48] : ~perp(X48,X49,X48,X49)),
  introduced(avatar_definition,[new_symbols(naming,[spl29_230])])).
fof(f8940,plain,(
  perp(sK20,sK23,sK20,sK23) | ~spl29_512),
  inference(resolution,[],[f8699,f365])).
fof(f8699,plain,(
  perp(sK20,sK23,sK23,sK20) | ~spl29_512),
  inference(avatar_component_clause,[],[f8698])).
fof(f8698,plain,(
  spl29_512 <=> perp(sK20,sK23,sK23,sK20)),
  introduced(avatar_definition,[new_symbols(naming,[spl29_512])])).
fof(f8864,plain,(
  ~spl29_124 | ~spl29_510),
  inference(avatar_contradiction_clause,[],[f8863])).
fof(f8863,plain,(
  $false | (~spl29_124 | ~spl29_510)),
  inference(resolution,[],[f8834,f488])).
fof(f8834,plain,(
  ( ! [X0] : (~midp(sK23,sK20,X0)) ) | (~spl29_124 | ~spl29_510)),
  inference(resolution,[],[f8693,f4037])).
fof(f4037,plain,(
  ( ! [X14,X12,X13] : (cong(X12,X13,X14,X13) | ~midp(X12,X13,X14)) ) | ~spl29_124),
  inference(resolution,[],[f1243,f381])).
fof(f381,plain,(
  ( ! [X2,X0,X3,X1] : (~cong(X0,X1,X2,X3) | cong(X0,X1,X3,X2)) )),
  inference(cnf_transformation,[],[f198])).
fof(f198,plain,(
  ! [X0,X1,X2,X3] : (cong(X0,X1,X3,X2) | ~cong(X0,X1,X2,X3))),
  inference(ennf_transformation,[],[f23])).
fof(f23,axiom,(
  ! [X0,X1,X2,X3] : (cong(X0,X1,X2,X3) => cong(X0,X1,X3,X2))),
  file('./7_Book/00EE/11/E075-29.p',ruleD23)).
fof(f1243,plain,(
  ( ! [X10,X11,X9] : (cong(X9,X10,X10,X11) | ~midp(X9,X10,X11)) ) | ~spl29_124),
  inference(avatar_component_clause,[],[f1242])).
fof(f1242,plain,(
  spl29_124 <=> ! [X9,X11,X10] : (cong(X9,X10,X10,X11) | ~midp(X9,X10,X11))),
  introduced(avatar_definition,[new_symbols(naming,[spl29_124])])).
fof(f8693,plain,(
  ( ! [X5] : (~cong(sK23,sK20,X5,sK20)) ) | ~spl29_510),
  inference(avatar_component_clause,[],[f8692])).
fof(f8692,plain,(
  spl29_510 <=> ! [X5] : ~cong(sK23,sK20,X5,sK20)),
  introduced(avatar_definition,[new_symbols(naming,[spl29_510])])).
fof(f8700,plain,(
  spl29_510 | spl29_512 | ~spl29_226),
  inference(avatar_split_clause,[],[f8690,f2213,f8698,f8692])).
fof(f8690,plain,(
  ( ! [X5] : (perp(sK20,sK23,sK23,sK20) | ~cong(sK23,sK20,X5,sK20)) ) | ~spl29_226),
  inference(subsumption_resolution,[],[f8686,f4153])).
fof(f8686,plain,(
  ( ! [X5] : (~coll(sK20,sK20,X5) | perp(sK20,sK23,sK23,sK20) | ~cong(sK23,sK20,X5,sK20)) )),
  inference(duplicate_literal_removal,[],[f8674])).
fof(f8674,plain,(
  ( ! [X5] : (~coll(sK20,sK20,X5) | perp(sK20,sK23,sK23,sK20) | ~cong(sK23,sK20,X5,sK20) | ~cong(sK23,sK20,X5,sK20)) )),
  inference(resolution,[],[f2542,f416])).
fof(f416,plain,(
  ( ! [X2,X0,X3,X1] : (~cyclic(X0,X1,X2,X3) | perp(X2,X0,X0,X3) | ~cong(X0,X3,X1,X3) | ~cong(X0,X2,X1,X2)) )),
  inference(cnf_transformation,[],[f253])).
fof(f253,plain,(
  ! [X0,X1,X2,X3] : (perp(X2,X0,X0,X3) | ~cyclic(X0,X1,X2,X3) | ~cong(X0,X3,X1,X3) | ~cong(X0,X2,X1,X2))),
  inference(flattening,[],[f252])).
fof(f252,plain,(
  ! [X0,X1,X2,X3] : (perp(X2,X0,X0,X3) | (~cyclic(X0,X1,X2,X3) | ~cong(X0,X3,X1,X3) | ~cong(X0,X2,X1,X2)))),
  inference(ennf_transformation,[],[f136])).
fof(f136,plain,(
  ! [X0,X1,X2,X3] : ((cyclic(X0,X1,X2,X3) & cong(X0,X3,X1,X3) & cong(X0,X2,X1,X2)) => perp(X2,X0,X0,X3))),
  inference(rectify,[],[f58])).
fof(f58,axiom,(
  ! [X0,X1,X8,X9] : ((cyclic(X0,X1,X8,X9) & cong(X0,X9,X1,X9) & cong(X0,X8,X1,X8)) => perp(X8,X0,X0,X9))),
  file('./7_Book/00EE/11/E075-29.p',ruleD57)).
fof(f2542,plain,(
  ( ! [X3] : (cyclic(sK23,X3,sK20,sK20) | ~coll(sK20,sK20,X3)) )),
  inference(resolution,[],[f994,f1664])).
fof(f1664,plain,(
  para(sK20,sK23,sK20,sK23)),
  inference(resolution,[],[f877,f512])).
fof(f994,plain,(
  ( ! [X2,X0,X1] : (~para(X0,X2,X0,X2) | cyclic(X2,X1,X0,X0) | ~coll(X0,X0,X1)) )),
  inference(resolution,[],[f401,f398])).
fof(f3951,plain,(
  ~spl29_228),
  inference(avatar_contradiction_clause,[],[f3944])).
fof(f3944,plain,(
  $false | ~spl29_228),
  inference(resolution,[],[f2217,f1664])).
fof(f2217,plain,(
  ( ! [X0,X1] : (~para(X0,X1,X0,X1)) ) | ~spl29_228),
  inference(avatar_component_clause,[],[f2216])).
fof(f2216,plain,(
  spl29_228 <=> ! [X1,X0] : ~para(X0,X1,X0,X1)),
  introduced(avatar_definition,[new_symbols(naming,[spl29_228])])).
fof(f3300,plain,(
  ~spl29_126 | ~spl29_226),
  inference(avatar_contradiction_clause,[],[f3295])).
fof(f3295,plain,(
  $false | (~spl29_126 | ~spl29_226)),
  inference(resolution,[],[f3269,f475])).
fof(f475,plain,(
  circle(sK23,sK20,sK25,sK27)),
  inference(cnf_transformation,[],[f358])).
fof(f3269,plain,(
  ( ! [X6,X8,X7,X9] : (~circle(X8,X9,X7,X6)) ) | (~spl29_126 | ~spl29_226)),
  inference(subsumption_resolution,[],[f3268,f2891])).
fof(f2891,plain,(
  ( ! [X0,X1] : (~midp(X0,X1,X0)) ) | ~spl29_126),
  inference(subsumption_resolution,[],[f2885,f369])).
fof(f2885,plain,(
  ( ! [X0,X1] : (~midp(X0,X0,X1) | ~midp(X0,X1,X0)) ) | ~spl29_126),
  inference(resolution,[],[f1246,f527])).
fof(f527,plain,(
  ( ! [X2,X0,X1] : (cong(X0,X1,X0,X2) | ~midp(X0,X2,X1)) )),
  inference(resolution,[],[f382,f427])).
fof(f427,plain,(
  ( ! [X2,X0,X1] : (cong(X0,X1,X0,X2) | ~midp(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f270])).
fof(f270,plain,(
  ! [X0,X1,X2] : (cong(X0,X1,X0,X2) | ~midp(X0,X1,X2))),
  inference(ennf_transformation,[],[f69])).
fof(f69,axiom,(
  ! [X0,X1,X2] : (midp(X0,X1,X2) => cong(X0,X1,X0,X2))),
  file('./7_Book/00EE/11/E075-29.p',ruleD68)).
fof(f382,plain,(
  ( ! [X2,X0,X3,X1] : (~cong(X0,X1,X2,X3) | cong(X2,X3,X0,X1)) )),
  inference(cnf_transformation,[],[f199])).
fof(f199,plain,(
  ! [X0,X1,X2,X3] : (cong(X2,X3,X0,X1) | ~cong(X0,X1,X2,X3))),
  inference(ennf_transformation,[],[f24])).
fof(f24,axiom,(
  ! [X0,X1,X2,X3] : (cong(X0,X1,X2,X3) => cong(X2,X3,X0,X1))),
  file('./7_Book/00EE/11/E075-29.p',ruleD24)).
fof(f1246,plain,(
  ( ! [X6,X8,X7] : (~cong(X6,X7,X7,X8) | ~midp(X6,X7,X8)) ) | ~spl29_126),
  inference(avatar_component_clause,[],[f1245])).
fof(f1245,plain,(
  spl29_126 <=> ! [X7,X8,X6] : (~cong(X6,X7,X7,X8) | ~midp(X6,X7,X8))),
  introduced(avatar_definition,[new_symbols(naming,[spl29_126])])).
fof(f3268,plain,(
  ( ! [X6,X8,X7,X9] : (~circle(X8,X9,X7,X6) | midp(X6,X7,X6)) ) | ~spl29_226),
  inference(subsumption_resolution,[],[f3263,f2305])).
fof(f2305,plain,(
  ( ! [X4,X3] : (coll(X3,X4,X3)) ) | ~spl29_226),
  inference(resolution,[],[f2270,f360])).
fof(f2270,plain,(
  ( ! [X0,X1] : (coll(X0,X1,X1)) ) | ~spl29_226),
  inference(resolution,[],[f2214,f425])).
fof(f3263,plain,(
  ( ! [X6,X8,X7,X9] : (~coll(X6,X7,X6) | ~circle(X8,X9,X7,X6) | midp(X6,X7,X6)) ) | ~spl29_226),
  inference(resolution,[],[f1361,f2600])).
fof(f2600,plain,(
  ( ! [X6,X4,X5,X3] : (cyclic(X3,X4,X5,X6)) ) | ~spl29_226),
  inference(subsumption_resolution,[],[f2593,f2408])).
fof(f2408,plain,(
  ( ! [X68,X66,X67] : (cyclic(X67,X67,X66,X68)) ) | ~spl29_226),
  inference(subsumption_resolution,[],[f2407,f2375])).
fof(f2375,plain,(
  ( ! [X2,X0,X1] : (coll(X0,X1,X2)) ) | ~spl29_226),
  inference(subsumption_resolution,[],[f2370,f2338])).
fof(f2338,plain,(
  ( ! [X6,X5] : (coll(X5,X5,X6)) ) | ~spl29_226),
  inference(resolution,[],[f2305,f359])).
fof(f2370,plain,(
  ( ! [X2,X0,X1] : (coll(X0,X1,X2) | ~coll(X2,X2,X0)) ) | ~spl29_226),
  inference(resolution,[],[f2338,f361])).
fof(f2407,plain,(
  ( ! [X68,X66,X67] : (~coll(X66,X68,X67) | cyclic(X67,X67,X66,X68)) ) | ~spl29_226),
  inference(subsumption_resolution,[],[f2389,f2214])).
fof(f2593,plain,(
  ( ! [X6,X4,X5,X3] : (cyclic(X3,X4,X5,X6) | ~cyclic(X3,X3,X4,X5)) ) | ~spl29_226),
  inference(resolution,[],[f2408,f375])).
fof(f375,plain,(
  ( ! [X4,X2,X0,X3,X1] : (~cyclic(X0,X1,X2,X4) | cyclic(X1,X2,X3,X4) | ~cyclic(X0,X1,X2,X3)) )),
  inference(cnf_transformation,[],[f191])).
fof(f191,plain,(
  ! [X0,X1,X2,X3,X4] : (cyclic(X1,X2,X3,X4) | ~cyclic(X0,X1,X2,X4) | ~cyclic(X0,X1,X2,X3))),
  inference(flattening,[],[f190])).
fof(f190,plain,(
  ! [X0,X1,X2,X3,X4] : (cyclic(X1,X2,X3,X4) | (~cyclic(X0,X1,X2,X4) | ~cyclic(X0,X1,X2,X3)))),
  inference(ennf_transformation,[],[f17])).
fof(f17,axiom,(
  ! [X0,X1,X2,X3,X4] : ((cyclic(X0,X1,X2,X4) & cyclic(X0,X1,X2,X3)) => cyclic(X1,X2,X3,X4))),
  file('./7_Book/00EE/11/E075-29.p',ruleD17)).
fof(f1361,plain,(
  ( ! [X6,X8,X7,X9] : (~cyclic(X7,X6,X9,X8) | ~coll(X6,X7,X6) | ~circle(X8,X9,X7,X6) | midp(X6,X7,X6)) )),
  inference(resolution,[],[f410,f399])).
fof(f410,plain,(
  ( ! [X4,X2,X0,X3,X1] : (~eqangle(X0,X1,X0,X2,X3,X1,X3,X4) | midp(X4,X1,X2) | ~coll(X4,X1,X2) | ~circle(X3,X0,X1,X2)) )),
  inference(cnf_transformation,[],[f241])).
fof(f241,plain,(
  ! [X0,X1,X2,X3,X4] : (midp(X4,X1,X2) | ~eqangle(X0,X1,X0,X2,X3,X1,X3,X4) | ~coll(X4,X1,X2) | ~circle(X3,X0,X1,X2))),
  inference(flattening,[],[f240])).
fof(f240,plain,(
  ! [X0,X1,X2,X3,X4] : (midp(X4,X1,X2) | (~eqangle(X0,X1,X0,X2,X3,X1,X3,X4) | ~coll(X4,X1,X2) | ~circle(X3,X0,X1,X2)))),
  inference(ennf_transformation,[],[f131])).
fof(f131,plain,(
  ! [X0,X1,X2,X3,X4] : ((eqangle(X0,X1,X0,X2,X3,X1,X3,X4) & coll(X4,X1,X2) & circle(X3,X0,X1,X2)) => midp(X4,X1,X2))),
  inference(rectify,[],[f52])).
fof(f52,axiom,(
  ! [X0,X1,X2,X7,X6] : ((eqangle(X0,X1,X0,X2,X7,X1,X7,X6) & coll(X6,X1,X2) & circle(X7,X0,X1,X2)) => midp(X6,X1,X2))),
  file('./7_Book/00EE/11/E075-29.p',ruleD51)).
fof(f2884,plain,(
  spl29_274 | spl29_126),
  inference(avatar_split_clause,[],[f2871,f1245,f2878])).
fof(f2871,plain,(
  ( ! [X14,X12,X10,X15,X13,X11] : (~midp(X10,X11,X12) | ~midp(X13,X14,X15) | ~cong(X10,X11,X11,X12) | cong(X14,X13,X14,X15)) )),
  inference(resolution,[],[f928,f434])).
fof(f434,plain,(
  ( ! [X6,X4,X2,X0,X7,X5,X3,X1] : (~eqratio(X0,X1,X2,X3,X4,X5,X6,X7) | ~cong(X4,X5,X6,X7) | cong(X0,X1,X2,X3)) )),
  inference(cnf_transformation,[],[f283])).
fof(f283,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6,X7] : (cong(X0,X1,X2,X3) | ~cong(X4,X5,X6,X7) | ~eqratio(X0,X1,X2,X3,X4,X5,X6,X7))),
  inference(flattening,[],[f282])).
fof(f282,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6,X7] : (cong(X0,X1,X2,X3) | (~cong(X4,X5,X6,X7) | ~eqratio(X0,X1,X2,X3,X4,X5,X6,X7)))),
  inference(ennf_transformation,[],[f148])).
fof(f148,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6,X7] : ((cong(X4,X5,X6,X7) & eqratio(X0,X1,X2,X3,X4,X5,X6,X7)) => cong(X0,X1,X2,X3))),
  inference(rectify,[],[f76])).
fof(f76,axiom,(
  ! [X0,X1,X2,X3,X8,X9,X10,X11] : ((cong(X8,X9,X10,X11) & eqratio(X0,X1,X2,X3,X8,X9,X10,X11)) => cong(X0,X1,X2,X3))),
  file('./7_Book/00EE/11/E075-29.p',ruleD75)).
fof(f928,plain,(
  ( ! [X6,X10,X8,X7,X11,X9] : (eqratio(X6,X7,X6,X8,X9,X10,X10,X11) | ~midp(X9,X10,X11) | ~midp(X7,X6,X8)) )),
  inference(resolution,[],[f384,f429])).
fof(f429,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (eqratio(X4,X0,X0,X1,X5,X2,X2,X3) | ~midp(X5,X2,X3) | ~midp(X4,X0,X1)) )),
  inference(cnf_transformation,[],[f273])).
fof(f273,plain,(
  ! [X0,X1,X2,X3,X4,X5] : (eqratio(X4,X0,X0,X1,X5,X2,X2,X3) | ~midp(X5,X2,X3) | ~midp(X4,X0,X1))),
  inference(flattening,[],[f272])).
fof(f272,plain,(
  ! [X0,X1,X2,X3,X4,X5] : (eqratio(X4,X0,X0,X1,X5,X2,X2,X3) | (~midp(X5,X2,X3) | ~midp(X4,X0,X1)))),
  inference(ennf_transformation,[],[f145])).
fof(f145,plain,(
  ! [X0,X1,X2,X3,X4,X5] : ((midp(X5,X2,X3) & midp(X4,X0,X1)) => eqratio(X4,X0,X0,X1,X5,X2,X2,X3))),
  inference(rectify,[],[f71])).
fof(f71,axiom,(
  ! [X0,X1,X2,X3,X6,X16] : ((midp(X16,X2,X3) & midp(X6,X0,X1)) => eqratio(X6,X0,X0,X1,X16,X2,X2,X3))),
  file('./7_Book/00EE/11/E075-29.p',ruleD70)).
fof(f384,plain,(
  ( ! [X6,X4,X2,X0,X7,X5,X3,X1] : (~eqratio(X0,X1,X2,X3,X4,X5,X6,X7) | eqratio(X1,X0,X2,X3,X4,X5,X6,X7)) )),
  inference(cnf_transformation,[],[f202])).
fof(f202,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6,X7] : (eqratio(X1,X0,X2,X3,X4,X5,X6,X7) | ~eqratio(X0,X1,X2,X3,X4,X5,X6,X7))),
  inference(ennf_transformation,[],[f105])).
fof(f105,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6,X7] : (eqratio(X0,X1,X2,X3,X4,X5,X6,X7) => eqratio(X1,X0,X2,X3,X4,X5,X6,X7))),
  inference(rectify,[],[f26])).
fof(f26,axiom,(
  ! [X0,X1,X2,X3,X8,X9,X10,X11] : (eqratio(X0,X1,X2,X3,X8,X9,X10,X11) => eqratio(X1,X0,X2,X3,X8,X9,X10,X11))),
  file('./7_Book/00EE/11/E075-29.p',ruleD26)).
fof(f2406,plain,(
  spl29_230 | spl29_232),
  inference(avatar_split_clause,[],[f2383,f2404,f2401])).
fof(f2383,plain,(
  ( ! [X47,X45,X48,X46,X44,X49] : (~para(X44,X45,X46,X47) | ~perp(X48,X49,X48,X49) | perp(X44,X45,X46,X47)) )),
  inference(resolution,[],[f923,f433])).
fof(f433,plain,(
  ( ! [X6,X4,X2,X0,X7,X5,X3,X1] : (~eqangle(X0,X1,X2,X3,X4,X5,X6,X7) | ~perp(X4,X5,X6,X7) | perp(X0,X1,X2,X3)) )),
  inference(cnf_transformation,[],[f281])).
fof(f281,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6,X7] : (perp(X0,X1,X2,X3) | ~perp(X4,X5,X6,X7) | ~eqangle(X0,X1,X2,X3,X4,X5,X6,X7))),
  inference(flattening,[],[f280])).
fof(f280,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6,X7] : (perp(X0,X1,X2,X3) | (~perp(X4,X5,X6,X7) | ~eqangle(X0,X1,X2,X3,X4,X5,X6,X7)))),
  inference(ennf_transformation,[],[f147])).
fof(f147,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6,X7] : ((perp(X4,X5,X6,X7) & eqangle(X0,X1,X2,X3,X4,X5,X6,X7)) => perp(X0,X1,X2,X3))),
  inference(rectify,[],[f75])).
fof(f75,axiom,(
  ! [X0,X1,X2,X3,X8,X9,X10,X11] : ((perp(X8,X9,X10,X11) & eqangle(X0,X1,X2,X3,X8,X9,X10,X11)) => perp(X0,X1,X2,X3))),
  file('./7_Book/00EE/11/E075-29.p',ruleD74)).
fof(f2218,plain,(
  spl29_226 | spl29_228),
  inference(avatar_split_clause,[],[f2189,f2216,f2213])).
fof(f2189,plain,(
  ( ! [X2,X0,X3,X1] : (~para(X0,X1,X0,X1) | para(X2,X3,X2,X3)) )),
  inference(resolution,[],[f915,f397])).
fof(f397,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (~eqangle(X0,X1,X4,X5,X2,X3,X4,X5) | para(X0,X1,X2,X3)) )),
  inference(cnf_transformation,[],[f218])).
fof(f218,plain,(
  ! [X0,X1,X2,X3,X4,X5] : (para(X0,X1,X2,X3) | ~eqangle(X0,X1,X4,X5,X2,X3,X4,X5))),
  inference(ennf_transformation,[],[f118])).
fof(f118,plain,(
  ! [X0,X1,X2,X3,X4,X5] : (eqangle(X0,X1,X4,X5,X2,X3,X4,X5) => para(X0,X1,X2,X3))),
  inference(rectify,[],[f39])).
fof(f39,axiom,(
  ! [X0,X1,X2,X3,X8,X9] : (eqangle(X0,X1,X8,X9,X2,X3,X8,X9) => para(X0,X1,X2,X3))),
  file('./7_Book/00EE/11/E075-29.p',ruleD39)).
fof(f1247,plain,(
  spl29_124 | spl29_126),
  inference(avatar_split_clause,[],[f1240,f1245,f1242])).
fof(f1240,plain,(
  ( ! [X6,X10,X8,X7,X11,X9] : (~cong(X6,X7,X7,X8) | cong(X9,X10,X10,X11) | ~midp(X6,X7,X8) | ~midp(X9,X10,X11)) )),
  inference(resolution,[],[f434,f429])).
% SZS output end Proof for E075-29
% ------------------------------
% Version: Vampire 4.4.0 (commit 7916d27 on 2019-08-23 08:50:16 +0100)
% Termination reason: Refutation

% Memory used [KB]: 20084
% Time elapsed: 0.515 s
% ------------------------------
% ------------------------------
