% Refutation found. Thanks to Tanya!
% SZS status Theorem for E037-22
% SZS output start Proof for E037-22
fof(f29419,plain,(
  $false),
  inference(avatar_sat_refutation,[],[f1918,f2184,f3669,f7917,f8489,f9013,f29418])).
fof(f29418,plain,(
  ~spl29_128 | ~spl29_142),
  inference(avatar_contradiction_clause,[],[f29411])).
fof(f29411,plain,(
  $false | (~spl29_128 | ~spl29_142)),
  inference(resolution,[],[f27361,f480])).
fof(f480,plain,(
  ~perp(sK20,sK22,sK25,sK23)),
  inference(cnf_transformation,[],[f358])).
fof(f358,plain,(
  ~perp(sK20,sK22,sK25,sK23) & circle(sK22,sK24,sK25,sK28) & coll(sK25,sK20,sK22) & circle(sK22,sK24,sK26,sK27) & midp(sK24,sK21,sK20) & cong(sK22,sK21,sK21,sK23) & coll(sK23,sK20,sK21) & cong(sK20,sK22,sK22,sK21) & perp(sK22,sK20,sK22,sK21)),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK20,sK21,sK22,sK23,sK24,sK25,sK26,sK27,sK28])],[f320,f357])).
fof(f357,plain,(
  ? [X0,X1,X2,X3,X4,X5,X6,X7,X8] : (~perp(X0,X2,X5,X3) & circle(X2,X4,X5,X8) & coll(X5,X0,X2) & circle(X2,X4,X6,X7) & midp(X4,X1,X0) & cong(X2,X1,X1,X3) & coll(X3,X0,X1) & cong(X0,X2,X2,X1) & perp(X2,X0,X2,X1)) => (~perp(sK20,sK22,sK25,sK23) & circle(sK22,sK24,sK25,sK28) & coll(sK25,sK20,sK22) & circle(sK22,sK24,sK26,sK27) & midp(sK24,sK21,sK20) & cong(sK22,sK21,sK21,sK23) & coll(sK23,sK20,sK21) & cong(sK20,sK22,sK22,sK21) & perp(sK22,sK20,sK22,sK21))),
  introduced(choice_axiom,[])).
fof(f320,plain,(
  ? [X0,X1,X2,X3,X4,X5,X6,X7,X8] : (~perp(X0,X2,X5,X3) & circle(X2,X4,X5,X8) & coll(X5,X0,X2) & circle(X2,X4,X6,X7) & midp(X4,X1,X0) & cong(X2,X1,X1,X3) & coll(X3,X0,X1) & cong(X0,X2,X2,X1) & perp(X2,X0,X2,X1))),
  inference(flattening,[],[f319])).
fof(f319,plain,(
  ? [X0,X1,X2,X3,X4,X5,X6,X7,X8] : (~perp(X0,X2,X5,X3) & (circle(X2,X4,X5,X8) & coll(X5,X0,X2) & circle(X2,X4,X6,X7) & midp(X4,X1,X0) & cong(X2,X1,X1,X3) & coll(X3,X0,X1) & cong(X0,X2,X2,X1) & perp(X2,X0,X2,X1)))),
  inference(ennf_transformation,[],[f167])).
fof(f167,plain,(
  ~! [X0,X1,X2,X3,X4,X5,X6,X7,X8] : ((circle(X2,X4,X5,X8) & coll(X5,X0,X2) & circle(X2,X4,X6,X7) & midp(X4,X1,X0) & cong(X2,X1,X1,X3) & coll(X3,X0,X1) & cong(X0,X2,X2,X1) & perp(X2,X0,X2,X1)) => perp(X0,X2,X5,X3))),
  inference(rectify,[],[f96])).
fof(f96,negated_conjecture,(
  ~! [X0,X1,X2,X3,X4,X5,X18,X19,X20] : ((circle(X2,X4,X5,X20) & coll(X5,X0,X2) & circle(X2,X4,X18,X19) & midp(X4,X1,X0) & cong(X2,X1,X1,X3) & coll(X3,X0,X1) & cong(X0,X2,X2,X1) & perp(X2,X0,X2,X1)) => perp(X0,X2,X5,X3))),
  inference(negated_conjecture,[],[f95])).
fof(f95,conjecture,(
  ! [X0,X1,X2,X3,X4,X5,X18,X19,X20] : ((circle(X2,X4,X5,X20) & coll(X5,X0,X2) & circle(X2,X4,X18,X19) & midp(X4,X1,X0) & cong(X2,X1,X1,X3) & coll(X3,X0,X1) & cong(X0,X2,X2,X1) & perp(X2,X0,X2,X1)) => perp(X0,X2,X5,X3))),
  file('./7_Book/00EE/03/E037-22.p',exemplo7Book00EE03E03722)).
fof(f27361,plain,(
  ( ! [X66,X64,X65,X63] : (perp(X63,X64,X65,X66)) ) | (~spl29_128 | ~spl29_142)),
  inference(subsumption_resolution,[],[f27345,f27282])).
fof(f27282,plain,(
  ( ! [X12,X10,X13,X11] : (para(X10,X11,X12,X13)) ) | ~spl29_128),
  inference(subsumption_resolution,[],[f27256,f27208])).
fof(f27208,plain,(
  ( ! [X59,X57,X58] : (midp(X57,X58,X59)) ) | ~spl29_128),
  inference(subsumption_resolution,[],[f27207,f10886])).
fof(f10886,plain,(
  ( ! [X37,X36] : (midp(X36,X37,X36)) ) | ~spl29_128),
  inference(resolution,[],[f10818,f4755])).
fof(f4755,plain,(
  ( ! [X14,X15,X16] : (~midp(X16,X15,X16) | midp(X14,X15,X14)) ) | ~spl29_128),
  inference(subsumption_resolution,[],[f3729,f3737])).
fof(f3737,plain,(
  ( ! [X4,X3] : (coll(X3,X4,X3)) ) | ~spl29_128),
  inference(resolution,[],[f3725,f360])).
fof(f360,plain,(
  ( ! [X2,X0,X1] : (~coll(X0,X1,X2) | coll(X1,X0,X2)) )),
  inference(cnf_transformation,[],[f169])).
fof(f169,plain,(
  ! [X0,X1,X2] : (coll(X1,X0,X2) | ~coll(X0,X1,X2))),
  inference(ennf_transformation,[],[f2])).
fof(f2,axiom,(
  ! [X0,X1,X2] : (coll(X0,X1,X2) => coll(X1,X0,X2))),
  file('./7_Book/00EE/03/E037-22.p',ruleD2)).
fof(f3725,plain,(
  ( ! [X4,X3] : (coll(X3,X4,X4)) ) | ~spl29_128),
  inference(resolution,[],[f1914,f425])).
fof(f425,plain,(
  ( ! [X2,X0,X1] : (~para(X0,X1,X0,X2) | coll(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f267])).
fof(f267,plain,(
  ! [X0,X1,X2] : (coll(X0,X1,X2) | ~para(X0,X1,X0,X2))),
  inference(ennf_transformation,[],[f67])).
fof(f67,axiom,(
  ! [X0,X1,X2] : (para(X0,X1,X0,X2) => coll(X0,X1,X2))),
  file('./7_Book/00EE/03/E037-22.p',ruleD66)).
fof(f1914,plain,(
  ( ! [X2,X3] : (para(X2,X3,X2,X3)) ) | ~spl29_128),
  inference(avatar_component_clause,[],[f1913])).
fof(f1913,plain,(
  spl29_128 <=> ! [X3,X2] : para(X2,X3,X2,X3)),
  introduced(avatar_definition,[new_symbols(naming,[spl29_128])])).
fof(f3729,plain,(
  ( ! [X14,X15,X16] : (~coll(X14,X15,X14) | midp(X14,X15,X14) | ~midp(X16,X15,X16)) ) | ~spl29_128),
  inference(resolution,[],[f1914,f404])).
fof(f404,plain,(
  ( ! [X4,X2,X0,X3,X1] : (~para(X3,X4,X1,X2) | ~coll(X4,X0,X2) | midp(X4,X0,X2) | ~midp(X3,X0,X1)) )),
  inference(cnf_transformation,[],[f230])).
fof(f230,plain,(
  ! [X0,X1,X2,X3,X4] : (midp(X4,X0,X2) | ~coll(X4,X0,X2) | ~para(X3,X4,X1,X2) | ~midp(X3,X0,X1))),
  inference(flattening,[],[f229])).
fof(f229,plain,(
  ! [X0,X1,X2,X3,X4] : (midp(X4,X0,X2) | (~coll(X4,X0,X2) | ~para(X3,X4,X1,X2) | ~midp(X3,X0,X1)))),
  inference(ennf_transformation,[],[f125])).
fof(f125,plain,(
  ! [X0,X1,X2,X3,X4] : ((coll(X4,X0,X2) & para(X3,X4,X1,X2) & midp(X3,X0,X1)) => midp(X4,X0,X2))),
  inference(rectify,[],[f46])).
fof(f46,axiom,(
  ! [X0,X1,X2,X4,X5] : ((coll(X5,X0,X2) & para(X4,X5,X1,X2) & midp(X4,X0,X1)) => midp(X5,X0,X2))),
  file('./7_Book/00EE/03/E037-22.p',ruleD45)).
fof(f10818,plain,(
  ( ! [X8,X7] : (midp(X7,X8,X8)) ) | ~spl29_128),
  inference(subsumption_resolution,[],[f10768,f3725])).
fof(f10768,plain,(
  ( ! [X8,X7] : (~coll(X7,X8,X8) | midp(X7,X8,X8)) ) | ~spl29_128),
  inference(resolution,[],[f10765,f426])).
fof(f426,plain,(
  ( ! [X2,X0,X1] : (~cong(X0,X1,X0,X2) | ~coll(X0,X1,X2) | midp(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f269])).
fof(f269,plain,(
  ! [X0,X1,X2] : (midp(X0,X1,X2) | ~coll(X0,X1,X2) | ~cong(X0,X1,X0,X2))),
  inference(flattening,[],[f268])).
fof(f268,plain,(
  ! [X0,X1,X2] : (midp(X0,X1,X2) | (~coll(X0,X1,X2) | ~cong(X0,X1,X0,X2)))),
  inference(ennf_transformation,[],[f68])).
fof(f68,axiom,(
  ! [X0,X1,X2] : ((coll(X0,X1,X2) & cong(X0,X1,X0,X2)) => midp(X0,X1,X2))),
  file('./7_Book/00EE/03/E037-22.p',ruleD67)).
fof(f10765,plain,(
  ( ! [X83,X84] : (cong(X83,X84,X83,X84)) ) | ~spl29_128),
  inference(subsumption_resolution,[],[f10764,f9361])).
fof(f9361,plain,(
  ( ! [X12,X10,X11,X9] : (cyclic(X9,X10,X11,X12)) ) | ~spl29_128),
  inference(subsumption_resolution,[],[f9339,f8344])).
fof(f8344,plain,(
  ( ! [X2,X0,X1] : (cyclic(X0,X0,X1,X2)) ) | ~spl29_128),
  inference(resolution,[],[f2918,f1914])).
fof(f2918,plain,(
  ( ! [X68,X66,X67] : (~para(X66,X67,X66,X67) | cyclic(X67,X67,X66,X68)) )),
  inference(subsumption_resolution,[],[f2167,f2168])).
fof(f2168,plain,(
  ( ! [X70,X71,X69] : (~para(X69,X70,X69,X70) | coll(X69,X71,X70) | cyclic(X70,X70,X69,X71)) )),
  inference(resolution,[],[f854,f400])).
fof(f400,plain,(
  ( ! [X2,X0,X3,X1] : (~eqangle(X2,X0,X2,X1,X3,X0,X3,X1) | coll(X2,X3,X0) | cyclic(X0,X1,X2,X3)) )),
  inference(cnf_transformation,[],[f222])).
fof(f222,plain,(
  ! [X0,X1,X2,X3] : (cyclic(X0,X1,X2,X3) | coll(X2,X3,X0) | ~eqangle(X2,X0,X2,X1,X3,X0,X3,X1))),
  inference(flattening,[],[f221])).
fof(f221,plain,(
  ! [X0,X1,X2,X3] : (cyclic(X0,X1,X2,X3) | (coll(X2,X3,X0) | ~eqangle(X2,X0,X2,X1,X3,X0,X3,X1)))),
  inference(ennf_transformation,[],[f121])).
fof(f121,plain,(
  ! [X0,X1,X2,X3] : ((~coll(X2,X3,X0) & eqangle(X2,X0,X2,X1,X3,X0,X3,X1)) => cyclic(X0,X1,X2,X3))),
  inference(rectify,[],[f42])).
fof(f42,axiom,(
  ! [X0,X1,X8,X9] : ((~coll(X8,X9,X0) & eqangle(X8,X0,X8,X1,X9,X0,X9,X1)) => cyclic(X0,X1,X8,X9))),
  file('./7_Book/00EE/03/E037-22.p',ruleD42a)).
fof(f854,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (eqangle(X0,X1,X2,X3,X4,X5,X4,X5) | ~para(X0,X1,X2,X3)) )),
  inference(resolution,[],[f379,f398])).
fof(f398,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (eqangle(X0,X1,X4,X5,X2,X3,X4,X5) | ~para(X0,X1,X2,X3)) )),
  inference(cnf_transformation,[],[f219])).
fof(f219,plain,(
  ! [X0,X1,X2,X3,X4,X5] : (eqangle(X0,X1,X4,X5,X2,X3,X4,X5) | ~para(X0,X1,X2,X3))),
  inference(ennf_transformation,[],[f119])).
fof(f119,plain,(
  ! [X0,X1,X2,X3,X4,X5] : (para(X0,X1,X2,X3) => eqangle(X0,X1,X4,X5,X2,X3,X4,X5))),
  inference(rectify,[],[f40])).
fof(f40,axiom,(
  ! [X0,X1,X2,X3,X8,X9] : (para(X0,X1,X2,X3) => eqangle(X0,X1,X8,X9,X2,X3,X8,X9))),
  file('./7_Book/00EE/03/E037-22.p',ruleD40)).
fof(f379,plain,(
  ( ! [X6,X4,X2,X0,X7,X5,X3,X1] : (~eqangle(X0,X1,X2,X3,X4,X5,X6,X7) | eqangle(X0,X1,X4,X5,X2,X3,X6,X7)) )),
  inference(cnf_transformation,[],[f195])).
fof(f195,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6,X7] : (eqangle(X0,X1,X4,X5,X2,X3,X6,X7) | ~eqangle(X0,X1,X2,X3,X4,X5,X6,X7))),
  inference(ennf_transformation,[],[f103])).
fof(f103,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6,X7] : (eqangle(X0,X1,X2,X3,X4,X5,X6,X7) => eqangle(X0,X1,X4,X5,X2,X3,X6,X7))),
  inference(rectify,[],[f21])).
fof(f21,axiom,(
  ! [X0,X1,X2,X3,X8,X9,X10,X11] : (eqangle(X0,X1,X2,X3,X8,X9,X10,X11) => eqangle(X0,X1,X8,X9,X2,X3,X10,X11))),
  file('./7_Book/00EE/03/E037-22.p',ruleD21)).
fof(f2167,plain,(
  ( ! [X68,X66,X67] : (~para(X66,X67,X66,X67) | ~coll(X66,X68,X67) | cyclic(X67,X67,X66,X68)) )),
  inference(resolution,[],[f854,f401])).
fof(f401,plain,(
  ( ! [X2,X0,X3,X1] : (~eqangle(X2,X0,X2,X1,X3,X0,X3,X1) | ~coll(X2,X3,X1) | cyclic(X0,X1,X2,X3)) )),
  inference(cnf_transformation,[],[f224])).
fof(f224,plain,(
  ! [X0,X1,X2,X3] : (cyclic(X0,X1,X2,X3) | ~coll(X2,X3,X1) | ~eqangle(X2,X0,X2,X1,X3,X0,X3,X1))),
  inference(flattening,[],[f223])).
fof(f223,plain,(
  ! [X0,X1,X2,X3] : (cyclic(X0,X1,X2,X3) | (~coll(X2,X3,X1) | ~eqangle(X2,X0,X2,X1,X3,X0,X3,X1)))),
  inference(ennf_transformation,[],[f122])).
fof(f122,plain,(
  ! [X0,X1,X2,X3] : ((coll(X2,X3,X1) & eqangle(X2,X0,X2,X1,X3,X0,X3,X1)) => cyclic(X0,X1,X2,X3))),
  inference(rectify,[],[f43])).
fof(f43,axiom,(
  ! [X0,X1,X8,X9] : ((coll(X8,X9,X1) & eqangle(X8,X0,X8,X1,X9,X0,X9,X1)) => cyclic(X0,X1,X8,X9))),
  file('./7_Book/00EE/03/E037-22.p',ruleD42b)).
fof(f9339,plain,(
  ( ! [X12,X10,X11,X9] : (cyclic(X9,X10,X11,X12) | ~cyclic(X9,X9,X10,X11)) ) | ~spl29_128),
  inference(resolution,[],[f8344,f375])).
fof(f375,plain,(
  ( ! [X4,X2,X0,X3,X1] : (~cyclic(X0,X1,X2,X4) | cyclic(X1,X2,X3,X4) | ~cyclic(X0,X1,X2,X3)) )),
  inference(cnf_transformation,[],[f191])).
fof(f191,plain,(
  ! [X0,X1,X2,X3,X4] : (cyclic(X1,X2,X3,X4) | ~cyclic(X0,X1,X2,X4) | ~cyclic(X0,X1,X2,X3))),
  inference(flattening,[],[f190])).
fof(f190,plain,(
  ! [X0,X1,X2,X3,X4] : (cyclic(X1,X2,X3,X4) | (~cyclic(X0,X1,X2,X4) | ~cyclic(X0,X1,X2,X3)))),
  inference(ennf_transformation,[],[f17])).
fof(f17,axiom,(
  ! [X0,X1,X2,X3,X4] : ((cyclic(X0,X1,X2,X4) & cyclic(X0,X1,X2,X3)) => cyclic(X1,X2,X3,X4))),
  file('./7_Book/00EE/03/E037-22.p',ruleD17)).
fof(f10764,plain,(
  ( ! [X85,X83,X84] : (cong(X83,X84,X83,X84) | ~cyclic(X83,X84,X85,X84)) ) | ~spl29_128),
  inference(subsumption_resolution,[],[f10763,f9361])).
fof(f10763,plain,(
  ( ! [X85,X83,X84] : (cong(X83,X84,X83,X84) | ~cyclic(X83,X84,X85,X83) | ~cyclic(X83,X84,X85,X84)) ) | ~spl29_128),
  inference(subsumption_resolution,[],[f10752,f9337])).
fof(f9337,plain,(
  ( ! [X4,X5,X3] : (cyclic(X3,X4,X3,X5)) ) | ~spl29_128),
  inference(resolution,[],[f8344,f373])).
fof(f373,plain,(
  ( ! [X2,X0,X3,X1] : (~cyclic(X0,X1,X2,X3) | cyclic(X0,X2,X1,X3)) )),
  inference(cnf_transformation,[],[f188])).
fof(f188,plain,(
  ! [X0,X1,X2,X3] : (cyclic(X0,X2,X1,X3) | ~cyclic(X0,X1,X2,X3))),
  inference(ennf_transformation,[],[f15])).
fof(f15,axiom,(
  ! [X0,X1,X2,X3] : (cyclic(X0,X1,X2,X3) => cyclic(X0,X2,X1,X3))),
  file('./7_Book/00EE/03/E037-22.p',ruleD15)).
fof(f10752,plain,(
  ( ! [X85,X83,X84] : (~cyclic(X83,X84,X83,X85) | cong(X83,X84,X83,X84) | ~cyclic(X83,X84,X85,X83) | ~cyclic(X83,X84,X85,X84)) )),
  inference(duplicate_literal_removal,[],[f10742])).
fof(f10742,plain,(
  ( ! [X85,X83,X84] : (~cyclic(X83,X84,X83,X85) | cong(X83,X84,X83,X84) | ~cyclic(X83,X84,X85,X83) | ~cyclic(X83,X84,X85,X84) | ~cyclic(X83,X84,X85,X83)) )),
  inference(resolution,[],[f1837,f402])).
fof(f402,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (~eqangle(X2,X0,X2,X1,X5,X3,X5,X4) | cong(X0,X1,X3,X4) | ~cyclic(X0,X1,X2,X5) | ~cyclic(X0,X1,X2,X4) | ~cyclic(X0,X1,X2,X3)) )),
  inference(cnf_transformation,[],[f226])).
fof(f226,plain,(
  ! [X0,X1,X2,X3,X4,X5] : (cong(X0,X1,X3,X4) | ~eqangle(X2,X0,X2,X1,X5,X3,X5,X4) | ~cyclic(X0,X1,X2,X5) | ~cyclic(X0,X1,X2,X4) | ~cyclic(X0,X1,X2,X3))),
  inference(flattening,[],[f225])).
fof(f225,plain,(
  ! [X0,X1,X2,X3,X4,X5] : (cong(X0,X1,X3,X4) | (~eqangle(X2,X0,X2,X1,X5,X3,X5,X4) | ~cyclic(X0,X1,X2,X5) | ~cyclic(X0,X1,X2,X4) | ~cyclic(X0,X1,X2,X3)))),
  inference(ennf_transformation,[],[f123])).
fof(f123,plain,(
  ! [X0,X1,X2,X3,X4,X5] : ((eqangle(X2,X0,X2,X1,X5,X3,X5,X4) & cyclic(X0,X1,X2,X5) & cyclic(X0,X1,X2,X4) & cyclic(X0,X1,X2,X3)) => cong(X0,X1,X3,X4))),
  inference(rectify,[],[f44])).
fof(f44,axiom,(
  ! [X0,X1,X2,X8,X9,X14] : ((eqangle(X2,X0,X2,X1,X14,X8,X14,X9) & cyclic(X0,X1,X2,X14) & cyclic(X0,X1,X2,X9) & cyclic(X0,X1,X2,X8)) => cong(X0,X1,X8,X9))),
  file('./7_Book/00EE/03/E037-22.p',ruleD43)).
fof(f1837,plain,(
  ( ! [X14,X12,X13,X11] : (eqangle(X14,X11,X14,X12,X11,X13,X13,X12) | ~cyclic(X11,X12,X13,X14)) )),
  inference(resolution,[],[f807,f378])).
fof(f378,plain,(
  ( ! [X6,X4,X2,X0,X7,X5,X3,X1] : (~eqangle(X0,X1,X2,X3,X4,X5,X6,X7) | eqangle(X4,X5,X6,X7,X0,X1,X2,X3)) )),
  inference(cnf_transformation,[],[f194])).
fof(f194,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6,X7] : (eqangle(X4,X5,X6,X7,X0,X1,X2,X3) | ~eqangle(X0,X1,X2,X3,X4,X5,X6,X7))),
  inference(ennf_transformation,[],[f102])).
fof(f102,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6,X7] : (eqangle(X0,X1,X2,X3,X4,X5,X6,X7) => eqangle(X4,X5,X6,X7,X0,X1,X2,X3))),
  inference(rectify,[],[f20])).
fof(f20,axiom,(
  ! [X0,X1,X2,X3,X8,X9,X10,X11] : (eqangle(X0,X1,X2,X3,X8,X9,X10,X11) => eqangle(X8,X9,X10,X11,X0,X1,X2,X3))),
  file('./7_Book/00EE/03/E037-22.p',ruleD20)).
fof(f807,plain,(
  ( ! [X6,X8,X7,X9] : (eqangle(X6,X7,X7,X8,X9,X6,X9,X8) | ~cyclic(X6,X8,X7,X9)) )),
  inference(resolution,[],[f376,f399])).
fof(f399,plain,(
  ( ! [X2,X0,X3,X1] : (eqangle(X2,X0,X2,X1,X3,X0,X3,X1) | ~cyclic(X0,X1,X2,X3)) )),
  inference(cnf_transformation,[],[f220])).
fof(f220,plain,(
  ! [X0,X1,X2,X3] : (eqangle(X2,X0,X2,X1,X3,X0,X3,X1) | ~cyclic(X0,X1,X2,X3))),
  inference(ennf_transformation,[],[f120])).
fof(f120,plain,(
  ! [X0,X1,X2,X3] : (cyclic(X0,X1,X2,X3) => eqangle(X2,X0,X2,X1,X3,X0,X3,X1))),
  inference(rectify,[],[f41])).
fof(f41,axiom,(
  ! [X0,X1,X8,X9] : (cyclic(X0,X1,X8,X9) => eqangle(X8,X0,X8,X1,X9,X0,X9,X1))),
  file('./7_Book/00EE/03/E037-22.p',ruleD41)).
fof(f376,plain,(
  ( ! [X6,X4,X2,X0,X7,X5,X3,X1] : (~eqangle(X0,X1,X2,X3,X4,X5,X6,X7) | eqangle(X1,X0,X2,X3,X4,X5,X6,X7)) )),
  inference(cnf_transformation,[],[f192])).
fof(f192,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6,X7] : (eqangle(X1,X0,X2,X3,X4,X5,X6,X7) | ~eqangle(X0,X1,X2,X3,X4,X5,X6,X7))),
  inference(ennf_transformation,[],[f100])).
fof(f100,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6,X7] : (eqangle(X0,X1,X2,X3,X4,X5,X6,X7) => eqangle(X1,X0,X2,X3,X4,X5,X6,X7))),
  inference(rectify,[],[f18])).
fof(f18,axiom,(
  ! [X0,X1,X2,X3,X8,X9,X10,X11] : (eqangle(X0,X1,X2,X3,X8,X9,X10,X11) => eqangle(X1,X0,X2,X3,X8,X9,X10,X11))),
  file('./7_Book/00EE/03/E037-22.p',ruleD18)).
fof(f27207,plain,(
  ( ! [X59,X57,X60,X58] : (midp(X57,X58,X59) | ~midp(X60,X58,X60)) ) | ~spl29_128),
  inference(subsumption_resolution,[],[f27175,f3789])).
fof(f3789,plain,(
  ( ! [X2,X0,X1] : (coll(X0,X1,X2)) ) | ~spl29_128),
  inference(subsumption_resolution,[],[f3786,f3771])).
fof(f3771,plain,(
  ( ! [X6,X5] : (coll(X5,X5,X6)) ) | ~spl29_128),
  inference(resolution,[],[f3737,f359])).
fof(f359,plain,(
  ( ! [X2,X0,X1] : (~coll(X0,X1,X2) | coll(X0,X2,X1)) )),
  inference(cnf_transformation,[],[f168])).
fof(f168,plain,(
  ! [X0,X1,X2] : (coll(X0,X2,X1) | ~coll(X0,X1,X2))),
  inference(ennf_transformation,[],[f1])).
fof(f1,axiom,(
  ! [X0,X1,X2] : (coll(X0,X1,X2) => coll(X0,X2,X1))),
  file('./7_Book/00EE/03/E037-22.p',ruleD1)).
fof(f3786,plain,(
  ( ! [X2,X0,X1] : (coll(X0,X1,X2) | ~coll(X2,X2,X0)) ) | ~spl29_128),
  inference(resolution,[],[f3771,f361])).
fof(f361,plain,(
  ( ! [X2,X0,X3,X1] : (~coll(X0,X1,X3) | coll(X2,X3,X0) | ~coll(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f171])).
fof(f171,plain,(
  ! [X0,X1,X2,X3] : (coll(X2,X3,X0) | ~coll(X0,X1,X3) | ~coll(X0,X1,X2))),
  inference(flattening,[],[f170])).
fof(f170,plain,(
  ! [X0,X1,X2,X3] : (coll(X2,X3,X0) | (~coll(X0,X1,X3) | ~coll(X0,X1,X2)))),
  inference(ennf_transformation,[],[f3])).
fof(f3,axiom,(
  ! [X0,X1,X2,X3] : ((coll(X0,X1,X3) & coll(X0,X1,X2)) => coll(X2,X3,X0))),
  file('./7_Book/00EE/03/E037-22.p',ruleD3)).
fof(f27175,plain,(
  ( ! [X59,X57,X60,X58] : (~coll(X57,X58,X59) | midp(X57,X58,X59) | ~midp(X60,X58,X60)) ) | ~spl29_128),
  inference(resolution,[],[f27145,f404])).
fof(f27145,plain,(
  ( ! [X39,X37,X38] : (para(X39,X37,X39,X38)) ) | ~spl29_128),
  inference(subsumption_resolution,[],[f27144,f1914])).
fof(f27144,plain,(
  ( ! [X39,X37,X38] : (~para(X37,X38,X37,X38) | para(X39,X37,X39,X38)) ) | ~spl29_128),
  inference(subsumption_resolution,[],[f27143,f9337])).
fof(f27143,plain,(
  ( ! [X39,X37,X38] : (~cyclic(X37,X38,X37,X39) | ~para(X37,X38,X37,X38) | para(X39,X37,X39,X38)) ) | ~spl29_128),
  inference(subsumption_resolution,[],[f27102,f12378])).
fof(f12378,plain,(
  ( ! [X14,X12,X13] : (cong(X12,X13,X12,X14)) ) | ~spl29_128),
  inference(subsumption_resolution,[],[f12356,f11025])).
fof(f11025,plain,(
  ( ! [X0,X1] : (midp(X0,X0,X1)) ) | ~spl29_128),
  inference(resolution,[],[f10886,f369])).
fof(f369,plain,(
  ( ! [X2,X0,X1] : (~midp(X2,X1,X0) | midp(X2,X0,X1)) )),
  inference(cnf_transformation,[],[f182])).
fof(f182,plain,(
  ! [X0,X1,X2] : (midp(X2,X0,X1) | ~midp(X2,X1,X0))),
  inference(ennf_transformation,[],[f97])).
fof(f97,plain,(
  ! [X0,X1,X2] : (midp(X2,X1,X0) => midp(X2,X0,X1))),
  inference(rectify,[],[f11])).
fof(f11,axiom,(
  ! [X0,X1,X6] : (midp(X6,X1,X0) => midp(X6,X0,X1))),
  file('./7_Book/00EE/03/E037-22.p',ruleD11)).
fof(f12356,plain,(
  ( ! [X14,X12,X13] : (cong(X12,X13,X12,X14) | ~midp(X12,X12,X13)) ) | ~spl29_128),
  inference(resolution,[],[f11025,f1781])).
fof(f1781,plain,(
  ( ! [X2,X0,X3,X1] : (~midp(X0,X3,X2) | cong(X0,X1,X0,X2) | ~midp(X0,X3,X1)) )),
  inference(resolution,[],[f743,f522])).
fof(f522,plain,(
  ( ! [X2,X0,X1] : (cong(X0,X1,X0,X2) | ~midp(X0,X2,X1)) )),
  inference(resolution,[],[f382,f427])).
fof(f427,plain,(
  ( ! [X2,X0,X1] : (cong(X0,X1,X0,X2) | ~midp(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f270])).
fof(f270,plain,(
  ! [X0,X1,X2] : (cong(X0,X1,X0,X2) | ~midp(X0,X1,X2))),
  inference(ennf_transformation,[],[f69])).
fof(f69,axiom,(
  ! [X0,X1,X2] : (midp(X0,X1,X2) => cong(X0,X1,X0,X2))),
  file('./7_Book/00EE/03/E037-22.p',ruleD68)).
fof(f382,plain,(
  ( ! [X2,X0,X3,X1] : (~cong(X0,X1,X2,X3) | cong(X2,X3,X0,X1)) )),
  inference(cnf_transformation,[],[f199])).
fof(f199,plain,(
  ! [X0,X1,X2,X3] : (cong(X2,X3,X0,X1) | ~cong(X0,X1,X2,X3))),
  inference(ennf_transformation,[],[f24])).
fof(f24,axiom,(
  ! [X0,X1,X2,X3] : (cong(X0,X1,X2,X3) => cong(X2,X3,X0,X1))),
  file('./7_Book/00EE/03/E037-22.p',ruleD24)).
fof(f743,plain,(
  ( ! [X4,X2,X0,X3,X1] : (~cong(X0,X1,X2,X4) | cong(X0,X1,X2,X3) | ~midp(X2,X4,X3)) )),
  inference(resolution,[],[f383,f427])).
fof(f383,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (~cong(X2,X3,X4,X5) | cong(X0,X1,X4,X5) | ~cong(X0,X1,X2,X3)) )),
  inference(cnf_transformation,[],[f201])).
fof(f201,plain,(
  ! [X0,X1,X2,X3,X4,X5] : (cong(X0,X1,X4,X5) | ~cong(X2,X3,X4,X5) | ~cong(X0,X1,X2,X3))),
  inference(flattening,[],[f200])).
fof(f200,plain,(
  ! [X0,X1,X2,X3,X4,X5] : (cong(X0,X1,X4,X5) | (~cong(X2,X3,X4,X5) | ~cong(X0,X1,X2,X3)))),
  inference(ennf_transformation,[],[f25])).
fof(f25,axiom,(
  ! [X0,X1,X2,X3,X4,X5] : ((cong(X2,X3,X4,X5) & cong(X0,X1,X2,X3)) => cong(X0,X1,X4,X5))),
  file('./7_Book/00EE/03/E037-22.p',ruleD25)).
fof(f27102,plain,(
  ( ! [X39,X37,X38] : (~cong(X37,X37,X37,X38) | ~cyclic(X37,X38,X37,X39) | ~para(X37,X38,X37,X38) | para(X39,X37,X39,X38)) )),
  inference(resolution,[],[f4297,f432])).
fof(f432,plain,(
  ( ! [X6,X4,X2,X0,X7,X5,X3,X1] : (~eqangle(X0,X1,X2,X3,X4,X5,X6,X7) | ~para(X4,X5,X6,X7) | para(X0,X1,X2,X3)) )),
  inference(cnf_transformation,[],[f279])).
fof(f279,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6,X7] : (para(X0,X1,X2,X3) | ~para(X4,X5,X6,X7) | ~eqangle(X0,X1,X2,X3,X4,X5,X6,X7))),
  inference(flattening,[],[f278])).
fof(f278,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6,X7] : (para(X0,X1,X2,X3) | (~para(X4,X5,X6,X7) | ~eqangle(X0,X1,X2,X3,X4,X5,X6,X7)))),
  inference(ennf_transformation,[],[f146])).
fof(f146,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6,X7] : ((para(X4,X5,X6,X7) & eqangle(X0,X1,X2,X3,X4,X5,X6,X7)) => para(X0,X1,X2,X3))),
  inference(rectify,[],[f74])).
fof(f74,axiom,(
  ! [X0,X1,X2,X3,X8,X9,X10,X11] : ((para(X8,X9,X10,X11) & eqangle(X0,X1,X2,X3,X8,X9,X10,X11)) => para(X0,X1,X2,X3))),
  file('./7_Book/00EE/03/E037-22.p',ruleD73)).
fof(f4297,plain,(
  ( ! [X57,X58,X56] : (eqangle(X56,X57,X56,X58,X57,X58,X57,X58) | ~cong(X57,X57,X57,X58) | ~cyclic(X57,X58,X57,X56)) )),
  inference(resolution,[],[f1225,f833])).
fof(f833,plain,(
  ( ! [X6,X8,X7,X9] : (eqangle(X6,X7,X6,X8,X9,X7,X9,X8) | ~cyclic(X7,X8,X9,X6)) )),
  inference(resolution,[],[f378,f399])).
fof(f1225,plain,(
  ( ! [X39,X37,X35,X41,X38,X36,X40] : (~eqangle(X35,X36,X37,X38,X41,X39,X39,X40) | eqangle(X35,X36,X37,X38,X39,X40,X41,X40) | ~cong(X41,X39,X41,X40)) )),
  inference(resolution,[],[f380,f405])).
fof(f405,plain,(
  ( ! [X2,X0,X1] : (eqangle(X2,X0,X0,X1,X0,X1,X2,X1) | ~cong(X2,X0,X2,X1)) )),
  inference(cnf_transformation,[],[f231])).
fof(f231,plain,(
  ! [X0,X1,X2] : (eqangle(X2,X0,X0,X1,X0,X1,X2,X1) | ~cong(X2,X0,X2,X1))),
  inference(ennf_transformation,[],[f126])).
fof(f126,plain,(
  ! [X0,X1,X2] : (cong(X2,X0,X2,X1) => eqangle(X2,X0,X0,X1,X0,X1,X2,X1))),
  inference(rectify,[],[f47])).
fof(f47,axiom,(
  ! [X0,X1,X7] : (cong(X7,X0,X7,X1) => eqangle(X7,X0,X0,X1,X0,X1,X7,X1))),
  file('./7_Book/00EE/03/E037-22.p',ruleD46)).
fof(f380,plain,(
  ( ! [X6,X4,X2,X0,X10,X8,X7,X5,X3,X1,X11,X9] : (~eqangle(X4,X5,X6,X7,X8,X9,X10,X11) | eqangle(X0,X1,X2,X3,X8,X9,X10,X11) | ~eqangle(X0,X1,X2,X3,X4,X5,X6,X7)) )),
  inference(cnf_transformation,[],[f197])).
fof(f197,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11] : (eqangle(X0,X1,X2,X3,X8,X9,X10,X11) | ~eqangle(X4,X5,X6,X7,X8,X9,X10,X11) | ~eqangle(X0,X1,X2,X3,X4,X5,X6,X7))),
  inference(flattening,[],[f196])).
fof(f196,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11] : (eqangle(X0,X1,X2,X3,X8,X9,X10,X11) | (~eqangle(X4,X5,X6,X7,X8,X9,X10,X11) | ~eqangle(X0,X1,X2,X3,X4,X5,X6,X7)))),
  inference(ennf_transformation,[],[f104])).
fof(f104,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11] : ((eqangle(X4,X5,X6,X7,X8,X9,X10,X11) & eqangle(X0,X1,X2,X3,X4,X5,X6,X7)) => eqangle(X0,X1,X2,X3,X8,X9,X10,X11))),
  inference(rectify,[],[f22])).
fof(f22,axiom,(
  ! [X0,X1,X2,X3,X8,X9,X10,X11,X4,X5,X12,X13] : ((eqangle(X8,X9,X10,X11,X4,X5,X12,X13) & eqangle(X0,X1,X2,X3,X8,X9,X10,X11)) => eqangle(X0,X1,X2,X3,X4,X5,X12,X13))),
  file('./7_Book/00EE/03/E037-22.p',ruleD22)).
fof(f27256,plain,(
  ( ! [X14,X12,X10,X13,X11] : (para(X10,X11,X12,X13) | ~midp(X10,X14,X12)) ) | ~spl29_128),
  inference(resolution,[],[f27208,f403])).
fof(f403,plain,(
  ( ! [X4,X2,X0,X3,X1] : (~midp(X4,X0,X2) | para(X3,X4,X1,X2) | ~midp(X3,X0,X1)) )),
  inference(cnf_transformation,[],[f228])).
fof(f228,plain,(
  ! [X0,X1,X2,X3,X4] : (para(X3,X4,X1,X2) | ~midp(X4,X0,X2) | ~midp(X3,X0,X1))),
  inference(flattening,[],[f227])).
fof(f227,plain,(
  ! [X0,X1,X2,X3,X4] : (para(X3,X4,X1,X2) | (~midp(X4,X0,X2) | ~midp(X3,X0,X1)))),
  inference(ennf_transformation,[],[f124])).
fof(f124,plain,(
  ! [X0,X1,X2,X3,X4] : ((midp(X4,X0,X2) & midp(X3,X0,X1)) => para(X3,X4,X1,X2))),
  inference(rectify,[],[f45])).
fof(f45,axiom,(
  ! [X0,X1,X2,X4,X5] : ((midp(X5,X0,X2) & midp(X4,X0,X1)) => para(X4,X5,X1,X2))),
  file('./7_Book/00EE/03/E037-22.p',ruleD44)).
fof(f27345,plain,(
  ( ! [X66,X64,X67,X65,X63] : (perp(X63,X64,X65,X66) | ~para(X63,X64,X65,X67)) ) | (~spl29_128 | ~spl29_142)),
  inference(resolution,[],[f27148,f368])).
fof(f368,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (~perp(X2,X3,X4,X5) | perp(X0,X1,X4,X5) | ~para(X0,X1,X2,X3)) )),
  inference(cnf_transformation,[],[f181])).
fof(f181,plain,(
  ! [X0,X1,X2,X3,X4,X5] : (perp(X0,X1,X4,X5) | ~perp(X2,X3,X4,X5) | ~para(X0,X1,X2,X3))),
  inference(flattening,[],[f180])).
fof(f180,plain,(
  ! [X0,X1,X2,X3,X4,X5] : (perp(X0,X1,X4,X5) | (~perp(X2,X3,X4,X5) | ~para(X0,X1,X2,X3)))),
  inference(ennf_transformation,[],[f10])).
fof(f10,axiom,(
  ! [X0,X1,X2,X3,X4,X5] : ((perp(X2,X3,X4,X5) & para(X0,X1,X2,X3)) => perp(X0,X1,X4,X5))),
  file('./7_Book/00EE/03/E037-22.p',ruleD10)).
fof(f27148,plain,(
  ( ! [X41,X42,X40] : (perp(X42,X40,X42,X41)) ) | (~spl29_128 | ~spl29_142)),
  inference(subsumption_resolution,[],[f27147,f9030])).
fof(f9030,plain,(
  ( ! [X0,X1] : (perp(X0,X1,X0,X1)) ) | (~spl29_128 | ~spl29_142)),
  inference(resolution,[],[f2183,f1914])).
fof(f2183,plain,(
  ( ! [X47,X45,X46,X44] : (~para(X44,X45,X46,X47) | perp(X44,X45,X46,X47)) ) | ~spl29_142),
  inference(avatar_component_clause,[],[f2182])).
fof(f2182,plain,(
  spl29_142 <=> ! [X44,X46,X45,X47] : (~para(X44,X45,X46,X47) | perp(X44,X45,X46,X47))),
  introduced(avatar_definition,[new_symbols(naming,[spl29_142])])).
fof(f27147,plain,(
  ( ! [X41,X42,X40] : (~perp(X40,X41,X40,X41) | perp(X42,X40,X42,X41)) ) | ~spl29_128),
  inference(subsumption_resolution,[],[f27146,f9337])).
fof(f27146,plain,(
  ( ! [X41,X42,X40] : (~cyclic(X40,X41,X40,X42) | ~perp(X40,X41,X40,X41) | perp(X42,X40,X42,X41)) ) | ~spl29_128),
  inference(subsumption_resolution,[],[f27103,f12378])).
fof(f27103,plain,(
  ( ! [X41,X42,X40] : (~cong(X40,X40,X40,X41) | ~cyclic(X40,X41,X40,X42) | ~perp(X40,X41,X40,X41) | perp(X42,X40,X42,X41)) )),
  inference(resolution,[],[f4297,f433])).
fof(f433,plain,(
  ( ! [X6,X4,X2,X0,X7,X5,X3,X1] : (~eqangle(X0,X1,X2,X3,X4,X5,X6,X7) | ~perp(X4,X5,X6,X7) | perp(X0,X1,X2,X3)) )),
  inference(cnf_transformation,[],[f281])).
fof(f281,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6,X7] : (perp(X0,X1,X2,X3) | ~perp(X4,X5,X6,X7) | ~eqangle(X0,X1,X2,X3,X4,X5,X6,X7))),
  inference(flattening,[],[f280])).
fof(f280,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6,X7] : (perp(X0,X1,X2,X3) | (~perp(X4,X5,X6,X7) | ~eqangle(X0,X1,X2,X3,X4,X5,X6,X7)))),
  inference(ennf_transformation,[],[f147])).
fof(f147,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6,X7] : ((perp(X4,X5,X6,X7) & eqangle(X0,X1,X2,X3,X4,X5,X6,X7)) => perp(X0,X1,X2,X3))),
  inference(rectify,[],[f75])).
fof(f75,axiom,(
  ! [X0,X1,X2,X3,X8,X9,X10,X11] : ((perp(X8,X9,X10,X11) & eqangle(X0,X1,X2,X3,X8,X9,X10,X11)) => perp(X0,X1,X2,X3))),
  file('./7_Book/00EE/03/E037-22.p',ruleD74)).
fof(f9013,plain,(
  ~spl29_18 | ~spl29_128 | ~spl29_140),
  inference(avatar_contradiction_clause,[],[f8997])).
fof(f8997,plain,(
  $false | (~spl29_18 | ~spl29_128 | ~spl29_140)),
  inference(resolution,[],[f8657,f2180])).
fof(f2180,plain,(
  ( ! [X48,X49] : (~perp(X48,X49,X48,X49)) ) | ~spl29_140),
  inference(avatar_component_clause,[],[f2179])).
fof(f2179,plain,(
  spl29_140 <=> ! [X49,X48] : ~perp(X48,X49,X48,X49)),
  introduced(avatar_definition,[new_symbols(naming,[spl29_140])])).
fof(f8657,plain,(
  ( ! [X33] : (perp(sK22,sK22,sK22,X33)) ) | (~spl29_18 | ~spl29_128)),
  inference(subsumption_resolution,[],[f8656,f6689])).
fof(f6689,plain,(
  ( ! [X39] : (midp(sK22,X39,X39)) ) | ~spl29_128),
  inference(resolution,[],[f3735,f1459])).
fof(f1459,plain,(
  midp(sK22,sK21,sK21)),
  inference(subsumption_resolution,[],[f1450,f1413])).
fof(f1413,plain,(
  coll(sK22,sK21,sK21)),
  inference(resolution,[],[f1412,f425])).
fof(f1412,plain,(
  para(sK22,sK21,sK22,sK21)),
  inference(resolution,[],[f710,f506])).
fof(f506,plain,(
  perp(sK22,sK21,sK22,sK20)),
  inference(resolution,[],[f366,f472])).
fof(f472,plain,(
  perp(sK22,sK20,sK22,sK21)),
  inference(cnf_transformation,[],[f358])).
fof(f366,plain,(
  ( ! [X2,X0,X3,X1] : (~perp(X0,X1,X2,X3) | perp(X2,X3,X0,X1)) )),
  inference(cnf_transformation,[],[f177])).
fof(f177,plain,(
  ! [X0,X1,X2,X3] : (perp(X2,X3,X0,X1) | ~perp(X0,X1,X2,X3))),
  inference(ennf_transformation,[],[f8])).
fof(f8,axiom,(
  ! [X0,X1,X2,X3] : (perp(X0,X1,X2,X3) => perp(X2,X3,X0,X1))),
  file('./7_Book/00EE/03/E037-22.p',ruleD8)).
fof(f710,plain,(
  ( ! [X0,X1] : (~perp(X0,X1,sK22,sK20) | para(X0,X1,sK22,sK21)) )),
  inference(resolution,[],[f367,f472])).
fof(f367,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (~perp(X2,X3,X4,X5) | para(X0,X1,X4,X5) | ~perp(X0,X1,X2,X3)) )),
  inference(cnf_transformation,[],[f179])).
fof(f179,plain,(
  ! [X0,X1,X2,X3,X4,X5] : (para(X0,X1,X4,X5) | ~perp(X2,X3,X4,X5) | ~perp(X0,X1,X2,X3))),
  inference(flattening,[],[f178])).
fof(f178,plain,(
  ! [X0,X1,X2,X3,X4,X5] : (para(X0,X1,X4,X5) | (~perp(X2,X3,X4,X5) | ~perp(X0,X1,X2,X3)))),
  inference(ennf_transformation,[],[f9])).
fof(f9,axiom,(
  ! [X0,X1,X2,X3,X4,X5] : ((perp(X2,X3,X4,X5) & perp(X0,X1,X2,X3)) => para(X0,X1,X4,X5))),
  file('./7_Book/00EE/03/E037-22.p',ruleD9)).
fof(f1450,plain,(
  ~coll(sK22,sK21,sK21) | midp(sK22,sK21,sK21)),
  inference(resolution,[],[f1432,f426])).
fof(f1432,plain,(
  cong(sK22,sK21,sK22,sK21)),
  inference(resolution,[],[f744,f520])).
fof(f520,plain,(
  cong(sK22,sK21,sK20,sK22)),
  inference(resolution,[],[f382,f473])).
fof(f473,plain,(
  cong(sK20,sK22,sK22,sK21)),
  inference(cnf_transformation,[],[f358])).
fof(f744,plain,(
  ( ! [X6,X5] : (~cong(X5,X6,sK20,sK22) | cong(X5,X6,sK22,sK21)) )),
  inference(resolution,[],[f383,f473])).
fof(f3735,plain,(
  ( ! [X12,X13,X11] : (~midp(X11,X13,X13) | midp(X11,X12,X12)) ) | ~spl29_128),
  inference(subsumption_resolution,[],[f3728,f1914])).
fof(f3728,plain,(
  ( ! [X12,X13,X11] : (midp(X11,X12,X12) | ~para(X13,X12,X13,X12) | ~midp(X11,X13,X13)) ) | ~spl29_128),
  inference(resolution,[],[f1914,f423])).
fof(f423,plain,(
  ( ! [X4,X2,X0,X3,X1] : (~para(X0,X3,X1,X2) | midp(X4,X2,X3) | ~para(X0,X2,X1,X3) | ~midp(X4,X0,X1)) )),
  inference(cnf_transformation,[],[f264])).
fof(f264,plain,(
  ! [X0,X1,X2,X3,X4] : (midp(X4,X2,X3) | ~para(X0,X3,X1,X2) | ~para(X0,X2,X1,X3) | ~midp(X4,X0,X1))),
  inference(flattening,[],[f263])).
fof(f263,plain,(
  ! [X0,X1,X2,X3,X4] : (midp(X4,X2,X3) | (~para(X0,X3,X1,X2) | ~para(X0,X2,X1,X3) | ~midp(X4,X0,X1)))),
  inference(ennf_transformation,[],[f143])).
fof(f143,plain,(
  ! [X0,X1,X2,X3,X4] : ((para(X0,X3,X1,X2) & para(X0,X2,X1,X3) & midp(X4,X0,X1)) => midp(X4,X2,X3))),
  inference(rectify,[],[f65])).
fof(f65,axiom,(
  ! [X0,X1,X2,X3,X6] : ((para(X0,X3,X1,X2) & para(X0,X2,X1,X3) & midp(X6,X0,X1)) => midp(X6,X2,X3))),
  file('./7_Book/00EE/03/E037-22.p',ruleD64)).
fof(f8656,plain,(
  ( ! [X33] : (perp(sK22,sK22,sK22,X33) | ~midp(sK22,X33,X33)) ) | (~spl29_18 | ~spl29_128)),
  inference(subsumption_resolution,[],[f8643,f6830])).
fof(f6830,plain,(
  ( ! [X0] : (midp(X0,sK21,X0)) ) | ~spl29_128),
  inference(resolution,[],[f4755,f6686])).
fof(f6686,plain,(
  ( ! [X36] : (midp(sK21,X36,X36)) ) | ~spl29_128),
  inference(resolution,[],[f3735,f1487])).
fof(f1487,plain,(
  midp(sK21,sK23,sK23)),
  inference(subsumption_resolution,[],[f1478,f1213])).
fof(f1213,plain,(
  coll(sK21,sK23,sK23)),
  inference(resolution,[],[f1184,f360])).
fof(f1184,plain,(
  coll(sK23,sK21,sK23)),
  inference(resolution,[],[f1166,f527])).
fof(f527,plain,(
  ( ! [X2] : (~coll(sK23,sK20,X2) | coll(X2,sK21,sK23)) )),
  inference(resolution,[],[f361,f474])).
fof(f474,plain,(
  coll(sK23,sK20,sK21)),
  inference(cnf_transformation,[],[f358])).
fof(f1166,plain,(
  coll(sK23,sK20,sK23)),
  inference(resolution,[],[f1153,f359])).
fof(f1153,plain,(
  coll(sK23,sK23,sK20)),
  inference(resolution,[],[f532,f491])).
fof(f491,plain,(
  coll(sK20,sK21,sK23)),
  inference(resolution,[],[f484,f359])).
fof(f484,plain,(
  coll(sK20,sK23,sK21)),
  inference(resolution,[],[f360,f474])).
fof(f532,plain,(
  ( ! [X7] : (~coll(sK20,sK21,X7) | coll(X7,sK23,sK20)) )),
  inference(resolution,[],[f361,f491])).
fof(f1478,plain,(
  ~coll(sK21,sK23,sK23) | midp(sK21,sK23,sK23)),
  inference(resolution,[],[f1441,f426])).
fof(f1441,plain,(
  cong(sK21,sK23,sK21,sK23)),
  inference(resolution,[],[f750,f521])).
fof(f521,plain,(
  cong(sK21,sK23,sK22,sK21)),
  inference(resolution,[],[f382,f475])).
fof(f475,plain,(
  cong(sK22,sK21,sK21,sK23)),
  inference(cnf_transformation,[],[f358])).
fof(f750,plain,(
  ( ! [X17,X18] : (~cong(X17,X18,sK22,sK21) | cong(X17,X18,sK21,sK23)) )),
  inference(resolution,[],[f383,f475])).
fof(f8643,plain,(
  ( ! [X33] : (~midp(sK22,sK21,sK22) | perp(sK22,sK22,sK22,X33) | ~midp(sK22,X33,X33)) ) | ~spl29_18),
  inference(resolution,[],[f8509,f754])).
fof(f754,plain,(
  ( ! [X2,X0,X1] : (~cong(X0,X1,X0,X1) | perp(X0,X0,X1,X2) | ~midp(X0,X2,X2)) )),
  inference(resolution,[],[f415,f427])).
fof(f415,plain,(
  ( ! [X2,X0,X3,X1] : (~cong(X0,X3,X1,X3) | perp(X0,X1,X2,X3) | ~cong(X0,X2,X1,X2)) )),
  inference(cnf_transformation,[],[f251])).
fof(f251,plain,(
  ! [X0,X1,X2,X3] : (perp(X0,X1,X2,X3) | ~cong(X0,X3,X1,X3) | ~cong(X0,X2,X1,X2))),
  inference(flattening,[],[f250])).
fof(f250,plain,(
  ! [X0,X1,X2,X3] : (perp(X0,X1,X2,X3) | (~cong(X0,X3,X1,X3) | ~cong(X0,X2,X1,X2)))),
  inference(ennf_transformation,[],[f135])).
fof(f135,plain,(
  ! [X0,X1,X2,X3] : ((cong(X0,X3,X1,X3) & cong(X0,X2,X1,X2)) => perp(X0,X1,X2,X3))),
  inference(rectify,[],[f57])).
fof(f57,axiom,(
  ! [X0,X1,X8,X9] : ((cong(X0,X9,X1,X9) & cong(X0,X8,X1,X8)) => perp(X0,X1,X8,X9))),
  file('./7_Book/00EE/03/E037-22.p',ruleD56)).
fof(f8509,plain,(
  ( ! [X20] : (cong(sK22,sK22,sK22,X20) | ~midp(sK22,sK21,X20)) ) | ~spl29_18),
  inference(resolution,[],[f652,f743])).
fof(f652,plain,(
  cong(sK22,sK22,sK22,sK21) | ~spl29_18),
  inference(avatar_component_clause,[],[f651])).
fof(f651,plain,(
  spl29_18 <=> cong(sK22,sK22,sK22,sK21)),
  introduced(avatar_definition,[new_symbols(naming,[spl29_18])])).
fof(f8489,plain,(
  spl29_18 | ~spl29_26),
  inference(avatar_split_clause,[],[f8483,f677,f651])).
fof(f677,plain,(
  spl29_26 <=> cong(sK22,sK22,sK22,sK20)),
  introduced(avatar_definition,[new_symbols(naming,[spl29_26])])).
fof(f8483,plain,(
  cong(sK22,sK22,sK22,sK21) | ~spl29_26),
  inference(resolution,[],[f7937,f744])).
fof(f7937,plain,(
  cong(sK22,sK22,sK20,sK22) | ~spl29_26),
  inference(resolution,[],[f678,f381])).
fof(f381,plain,(
  ( ! [X2,X0,X3,X1] : (~cong(X0,X1,X2,X3) | cong(X0,X1,X3,X2)) )),
  inference(cnf_transformation,[],[f198])).
fof(f198,plain,(
  ! [X0,X1,X2,X3] : (cong(X0,X1,X3,X2) | ~cong(X0,X1,X2,X3))),
  inference(ennf_transformation,[],[f23])).
fof(f23,axiom,(
  ! [X0,X1,X2,X3] : (cong(X0,X1,X2,X3) => cong(X0,X1,X3,X2))),
  file('./7_Book/00EE/03/E037-22.p',ruleD23)).
fof(f678,plain,(
  cong(sK22,sK22,sK22,sK20) | ~spl29_26),
  inference(avatar_component_clause,[],[f677])).
fof(f7917,plain,(
  spl29_27 | ~spl29_128),
  inference(avatar_contradiction_clause,[],[f7916])).
fof(f7916,plain,(
  $false | (~spl29_27 | ~spl29_128)),
  inference(subsumption_resolution,[],[f1632,f7870])).
fof(f7870,plain,(
  ( ! [X11] : (midp(X11,sK20,X11)) ) | ~spl29_128),
  inference(resolution,[],[f7736,f4755])).
fof(f7736,plain,(
  ( ! [X5] : (midp(sK20,X5,X5)) ) | ~spl29_128),
  inference(resolution,[],[f7671,f3735])).
fof(f7671,plain,(
  midp(sK20,sK22,sK22)),
  inference(subsumption_resolution,[],[f7651,f1082])).
fof(f1082,plain,(
  coll(sK20,sK22,sK22)),
  inference(resolution,[],[f1080,f360])).
fof(f1080,plain,(
  coll(sK22,sK20,sK22)),
  inference(resolution,[],[f1069,f359])).
fof(f1069,plain,(
  coll(sK22,sK22,sK20)),
  inference(resolution,[],[f539,f485])).
fof(f485,plain,(
  coll(sK20,sK25,sK22)),
  inference(resolution,[],[f360,f478])).
fof(f478,plain,(
  coll(sK25,sK20,sK22)),
  inference(cnf_transformation,[],[f358])).
fof(f539,plain,(
  ( ! [X14] : (~coll(sK20,sK25,X14) | coll(X14,sK22,sK20)) )),
  inference(resolution,[],[f361,f485])).
fof(f7651,plain,(
  ~coll(sK20,sK22,sK22) | midp(sK20,sK22,sK22)),
  inference(resolution,[],[f7438,f426])).
fof(f7438,plain,(
  cong(sK20,sK22,sK20,sK22)),
  inference(resolution,[],[f749,f473])).
fof(f749,plain,(
  ( ! [X15,X16] : (~cong(X15,X16,sK22,sK21) | cong(X15,X16,sK20,sK22)) )),
  inference(resolution,[],[f383,f520])).
fof(f1632,plain,(
  ~midp(sK22,sK20,sK22) | ~spl29_27),
  inference(resolution,[],[f675,f522])).
fof(f675,plain,(
  ~cong(sK22,sK22,sK22,sK20) | ~spl29_27),
  inference(avatar_component_clause,[],[f674])).
fof(f674,plain,(
  spl29_27 <=> ~cong(sK22,sK22,sK22,sK20)),
  introduced(avatar_definition,[new_symbols(naming,[spl29_27])])).
fof(f3669,plain,(
  ~spl29_130),
  inference(avatar_contradiction_clause,[],[f3668])).
fof(f3668,plain,(
  $false | ~spl29_130),
  inference(resolution,[],[f1917,f1412])).
fof(f1917,plain,(
  ( ! [X0,X1] : (~para(X0,X1,X0,X1)) ) | ~spl29_130),
  inference(avatar_component_clause,[],[f1916])).
fof(f1916,plain,(
  spl29_130 <=> ! [X1,X0] : ~para(X0,X1,X0,X1)),
  introduced(avatar_definition,[new_symbols(naming,[spl29_130])])).
fof(f2184,plain,(
  spl29_140 | spl29_142),
  inference(avatar_split_clause,[],[f2161,f2182,f2179])).
fof(f2161,plain,(
  ( ! [X47,X45,X48,X46,X44,X49] : (~para(X44,X45,X46,X47) | ~perp(X48,X49,X48,X49) | perp(X44,X45,X46,X47)) )),
  inference(resolution,[],[f854,f433])).
fof(f1918,plain,(
  spl29_128 | spl29_130),
  inference(avatar_split_clause,[],[f1889,f1916,f1913])).
fof(f1889,plain,(
  ( ! [X2,X0,X3,X1] : (~para(X0,X1,X0,X1) | para(X2,X3,X2,X3)) )),
  inference(resolution,[],[f828,f397])).
fof(f397,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (~eqangle(X0,X1,X4,X5,X2,X3,X4,X5) | para(X0,X1,X2,X3)) )),
  inference(cnf_transformation,[],[f218])).
fof(f218,plain,(
  ! [X0,X1,X2,X3,X4,X5] : (para(X0,X1,X2,X3) | ~eqangle(X0,X1,X4,X5,X2,X3,X4,X5))),
  inference(ennf_transformation,[],[f118])).
fof(f118,plain,(
  ! [X0,X1,X2,X3,X4,X5] : (eqangle(X0,X1,X4,X5,X2,X3,X4,X5) => para(X0,X1,X2,X3))),
  inference(rectify,[],[f39])).
fof(f39,axiom,(
  ! [X0,X1,X2,X3,X8,X9] : (eqangle(X0,X1,X8,X9,X2,X3,X8,X9) => para(X0,X1,X2,X3))),
  file('./7_Book/00EE/03/E037-22.p',ruleD39)).
fof(f828,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (eqangle(X0,X1,X2,X3,X0,X1,X4,X5) | ~para(X2,X3,X4,X5)) )),
  inference(resolution,[],[f377,f398])).
fof(f377,plain,(
  ( ! [X6,X4,X2,X0,X7,X5,X3,X1] : (~eqangle(X0,X1,X2,X3,X4,X5,X6,X7) | eqangle(X2,X3,X0,X1,X6,X7,X4,X5)) )),
  inference(cnf_transformation,[],[f193])).
fof(f193,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6,X7] : (eqangle(X2,X3,X0,X1,X6,X7,X4,X5) | ~eqangle(X0,X1,X2,X3,X4,X5,X6,X7))),
  inference(ennf_transformation,[],[f101])).
fof(f101,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6,X7] : (eqangle(X0,X1,X2,X3,X4,X5,X6,X7) => eqangle(X2,X3,X0,X1,X6,X7,X4,X5))),
  inference(rectify,[],[f19])).
fof(f19,axiom,(
  ! [X0,X1,X2,X3,X8,X9,X10,X11] : (eqangle(X0,X1,X2,X3,X8,X9,X10,X11) => eqangle(X2,X3,X0,X1,X10,X11,X8,X9))),
  file('./7_Book/00EE/03/E037-22.p',ruleD19)).
% SZS output end Proof for E037-22
% ------------------------------
% Version: Vampire 4.4.0 (commit 7916d27 on 2019-08-23 08:50:16 +0100)
% Termination reason: Refutation

% Memory used [KB]: 19317
% Time elapsed: 0.628 s
% ------------------------------
% ------------------------------
