% Refutation found. Thanks to Tanya!
% SZS status Theorem for E037-24
% SZS output start Proof for E037-24
fof(f1000,plain,(
  $false),
  inference(resolution,[],[f999,f480])).
fof(f480,plain,(
  ~perp(sK23,sK22,sK21,sK20)),
  inference(cnf_transformation,[],[f357])).
fof(f357,plain,(
  ~perp(sK23,sK22,sK21,sK20) & ! [X4,X5,X6,X7,X8,X9,X10,X11] : (coll(X11,X5,X6) & eqangle(X11,X4,X4,X6,X11,X4,X4,X5) & coll(X10,X8,X9) & coll(X10,X4,X6) & coll(X9,X4,X5) & eqangle(X9,X8,X8,X4,X9,X8,X8,X5) & coll(X8,X5,X6) & perp(X8,X4,X4,X7) & circle(X7,X5,X4,X6))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK20,sK21,sK22,sK23])],[f319,f356])).
fof(f356,plain,(
  ? [X0,X1,X2,X3] : (~perp(X3,X2,X1,X0) & ! [X4,X5,X6,X7,X8,X9,X10,X11] : (coll(X11,X5,X6) & eqangle(X11,X4,X4,X6,X11,X4,X4,X5) & coll(X10,X8,X9) & coll(X10,X4,X6) & coll(X9,X4,X5) & eqangle(X9,X8,X8,X4,X9,X8,X8,X5) & coll(X8,X5,X6) & perp(X8,X4,X4,X7) & circle(X7,X5,X4,X6))) => (~perp(sK23,sK22,sK21,sK20) & ! [X4,X5,X6,X7,X8,X9,X10,X11] : (coll(X11,X5,X6) & eqangle(X11,X4,X4,X6,X11,X4,X4,X5) & coll(X10,X8,X9) & coll(X10,X4,X6) & coll(X9,X4,X5) & eqangle(X9,X8,X8,X4,X9,X8,X8,X5) & coll(X8,X5,X6) & perp(X8,X4,X4,X7) & circle(X7,X5,X4,X6)))),
  introduced(choice_axiom,[])).
fof(f319,plain,(
  ? [X0,X1,X2,X3] : (~perp(X3,X2,X1,X0) & ! [X4,X5,X6,X7,X8,X9,X10,X11] : (coll(X11,X5,X6) & eqangle(X11,X4,X4,X6,X11,X4,X4,X5) & coll(X10,X8,X9) & coll(X10,X4,X6) & coll(X9,X4,X5) & eqangle(X9,X8,X8,X4,X9,X8,X8,X5) & coll(X8,X5,X6) & perp(X8,X4,X4,X7) & circle(X7,X5,X4,X6)))),
  inference(ennf_transformation,[],[f167])).
fof(f167,plain,(
  ~! [X0,X1,X2,X3] : (! [X4,X5,X6,X7,X8,X9,X10,X11] : (coll(X11,X5,X6) & eqangle(X11,X4,X4,X6,X11,X4,X4,X5) & coll(X10,X8,X9) & coll(X10,X4,X6) & coll(X9,X4,X5) & eqangle(X9,X8,X8,X4,X9,X8,X8,X5) & coll(X8,X5,X6) & perp(X8,X4,X4,X7) & circle(X7,X5,X4,X6)) => perp(X3,X2,X1,X0))),
  inference(rectify,[],[f96])).
fof(f96,negated_conjecture,(
  ~! [X13,X0,X4,X5] : (! [X0,X1,X2,X3,X4,X5,X12,X13] : (coll(X13,X1,X2) & eqangle(X13,X0,X0,X2,X13,X0,X0,X1) & coll(X12,X4,X5) & coll(X12,X0,X2) & coll(X5,X0,X1) & eqangle(X5,X4,X4,X0,X5,X4,X4,X1) & coll(X4,X1,X2) & perp(X4,X0,X0,X3) & circle(X3,X1,X0,X2)) => perp(X5,X4,X0,X13))),
  inference(negated_conjecture,[],[f95])).
fof(f95,conjecture,(
  ( ! [X4,X0,X5,X13] : (! [X0,X1,X2,X3,X4,X5,X12,X13] : (coll(X13,X1,X2) & eqangle(X13,X0,X0,X2,X13,X0,X0,X1) & coll(X12,X4,X5) & coll(X12,X0,X2) & coll(X5,X0,X1) & eqangle(X5,X4,X4,X0,X5,X4,X4,X1) & coll(X4,X1,X2) & perp(X4,X0,X0,X3) & circle(X3,X1,X0,X2)) => perp(X5,X4,X0,X13)) )),
  file('./7_Book/00EE/03/E037-24.p',exemplo7Book00EE03E03724)).
fof(f999,plain,(
  ( ! [X85,X83,X84,X82] : (perp(X82,X83,X84,X85)) )),
  inference(subsumption_resolution,[],[f998,f483])).
fof(f483,plain,(
  ( ! [X2,X0,X1] : (perp(X0,X1,X2,X1)) )),
  inference(resolution,[],[f364,f472])).
fof(f472,plain,(
  ( ! [X4,X8,X7] : (perp(X8,X4,X4,X7)) )),
  inference(cnf_transformation,[],[f357])).
fof(f364,plain,(
  ( ! [X2,X0,X3,X1] : (~perp(X0,X1,X2,X3) | perp(X0,X1,X3,X2)) )),
  inference(cnf_transformation,[],[f176])).
fof(f176,plain,(
  ! [X0,X1,X2,X3] : (perp(X0,X1,X3,X2) | ~perp(X0,X1,X2,X3))),
  inference(ennf_transformation,[],[f7])).
fof(f7,axiom,(
  ! [X0,X1,X2,X3] : (perp(X0,X1,X2,X3) => perp(X0,X1,X3,X2))),
  file('./7_Book/00EE/03/E037-24.p',ruleD7)).
fof(f998,plain,(
  ( ! [X80,X85,X83,X81,X84,X82] : (~perp(X80,X81,X80,X81) | perp(X82,X83,X84,X85)) )),
  inference(resolution,[],[f432,f678])).
fof(f678,plain,(
  ( ! [X6,X4,X8,X7,X5,X9] : (eqangle(X4,X5,X6,X7,X8,X9,X8,X9)) )),
  inference(subsumption_resolution,[],[f673,f640])).
fof(f640,plain,(
  ( ! [X2,X0,X3,X1] : (para(X0,X1,X2,X3)) )),
  inference(resolution,[],[f554,f472])).
fof(f554,plain,(
  ( ! [X4,X2,X0,X3,X1] : (~perp(X0,X1,X4,X2) | para(X0,X1,X2,X3)) )),
  inference(resolution,[],[f366,f472])).
fof(f366,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (~perp(X2,X3,X4,X5) | para(X0,X1,X4,X5) | ~perp(X0,X1,X2,X3)) )),
  inference(cnf_transformation,[],[f179])).
fof(f179,plain,(
  ! [X0,X1,X2,X3,X4,X5] : (para(X0,X1,X4,X5) | ~perp(X2,X3,X4,X5) | ~perp(X0,X1,X2,X3))),
  inference(flattening,[],[f178])).
fof(f178,plain,(
  ! [X0,X1,X2,X3,X4,X5] : (para(X0,X1,X4,X5) | (~perp(X2,X3,X4,X5) | ~perp(X0,X1,X2,X3)))),
  inference(ennf_transformation,[],[f9])).
fof(f9,axiom,(
  ! [X0,X1,X2,X3,X4,X5] : ((perp(X2,X3,X4,X5) & perp(X0,X1,X2,X3)) => para(X0,X1,X4,X5))),
  file('./7_Book/00EE/03/E037-24.p',ruleD9)).
fof(f673,plain,(
  ( ! [X6,X4,X8,X7,X5,X9] : (eqangle(X4,X5,X6,X7,X8,X9,X8,X9) | ~para(X4,X5,X6,X7)) )),
  inference(resolution,[],[f378,f397])).
fof(f397,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (eqangle(X0,X1,X4,X5,X2,X3,X4,X5) | ~para(X0,X1,X2,X3)) )),
  inference(cnf_transformation,[],[f219])).
fof(f219,plain,(
  ! [X0,X1,X2,X3,X4,X5] : (eqangle(X0,X1,X4,X5,X2,X3,X4,X5) | ~para(X0,X1,X2,X3))),
  inference(ennf_transformation,[],[f119])).
fof(f119,plain,(
  ! [X0,X1,X2,X3,X4,X5] : (para(X0,X1,X2,X3) => eqangle(X0,X1,X4,X5,X2,X3,X4,X5))),
  inference(rectify,[],[f40])).
fof(f40,axiom,(
  ! [X0,X1,X2,X3,X8,X9] : (para(X0,X1,X2,X3) => eqangle(X0,X1,X8,X9,X2,X3,X8,X9))),
  file('./7_Book/00EE/03/E037-24.p',ruleD40)).
fof(f378,plain,(
  ( ! [X6,X4,X2,X0,X7,X5,X3,X1] : (~eqangle(X0,X1,X2,X3,X4,X5,X6,X7) | eqangle(X0,X1,X4,X5,X2,X3,X6,X7)) )),
  inference(cnf_transformation,[],[f195])).
fof(f195,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6,X7] : (eqangle(X0,X1,X4,X5,X2,X3,X6,X7) | ~eqangle(X0,X1,X2,X3,X4,X5,X6,X7))),
  inference(ennf_transformation,[],[f103])).
fof(f103,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6,X7] : (eqangle(X0,X1,X2,X3,X4,X5,X6,X7) => eqangle(X0,X1,X4,X5,X2,X3,X6,X7))),
  inference(rectify,[],[f21])).
fof(f21,axiom,(
  ! [X0,X1,X2,X3,X8,X9,X10,X11] : (eqangle(X0,X1,X2,X3,X8,X9,X10,X11) => eqangle(X0,X1,X8,X9,X2,X3,X10,X11))),
  file('./7_Book/00EE/03/E037-24.p',ruleD21)).
fof(f432,plain,(
  ( ! [X6,X4,X2,X0,X7,X5,X3,X1] : (~eqangle(X0,X1,X2,X3,X4,X5,X6,X7) | ~perp(X4,X5,X6,X7) | perp(X0,X1,X2,X3)) )),
  inference(cnf_transformation,[],[f281])).
fof(f281,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6,X7] : (perp(X0,X1,X2,X3) | ~perp(X4,X5,X6,X7) | ~eqangle(X0,X1,X2,X3,X4,X5,X6,X7))),
  inference(flattening,[],[f280])).
fof(f280,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6,X7] : (perp(X0,X1,X2,X3) | (~perp(X4,X5,X6,X7) | ~eqangle(X0,X1,X2,X3,X4,X5,X6,X7)))),
  inference(ennf_transformation,[],[f147])).
fof(f147,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6,X7] : ((perp(X4,X5,X6,X7) & eqangle(X0,X1,X2,X3,X4,X5,X6,X7)) => perp(X0,X1,X2,X3))),
  inference(rectify,[],[f75])).
fof(f75,axiom,(
  ! [X0,X1,X2,X3,X8,X9,X10,X11] : ((perp(X8,X9,X10,X11) & eqangle(X0,X1,X2,X3,X8,X9,X10,X11)) => perp(X0,X1,X2,X3))),
  file('./7_Book/00EE/03/E037-24.p',ruleD74)).
% SZS output end Proof for E037-24
% ------------------------------
% Version: Vampire 4.4.0 (commit 7916d27 on 2019-08-23 08:50:16 +0100)
% Termination reason: Refutation

% Memory used [KB]: 5245
% Time elapsed: 0.033 s
% ------------------------------
% ------------------------------
