% Refutation found. Thanks to Tanya!
% SZS status Theorem for E040-28-1
% SZS output start Proof for E040-28-1
44. ! [X0,X1,X2,X4,X5] : ((midp(X5,X0,X2) & midp(X4,X0,X1)) => para(X4,X5,X1,X2)) [input]
53. ! [X0,X1,X2,X7] : ((coll(X7,X0,X2) & circle(X7,X0,X1,X2)) => perp(X0,X1,X1,X2)) [input]
55. ! [X0,X1,X6,X7] : ((perp(X7,X6,X0,X1) & midp(X6,X0,X1)) => cong(X7,X0,X7,X1)) [input]
67. ! [X0,X1,X2] : ((coll(X0,X1,X2) & cong(X0,X1,X0,X2)) => midp(X0,X1,X2)) [input]
94. ! [X0,X1,X2,X3,X4,X5,X18] : (coll(X5,X0,X2) & circle(X4,X0,X5,X18) & circle(X4,X0,X1,X3) & coll(X3,X1,X2) & (cong(X2,X0,X2,X1) | cong(X2,X0,X0,X1))) => para(X0,X1,X5,X3) [input]
95. ~! [X3,X5,X1,X0] : (! [X0,X1,X2,X3,X4,X5,X18] : (coll(X5,X0,X2) & circle(X4,X0,X5,X18) & circle(X4,X0,X1,X3) & coll(X3,X1,X2) & (cong(X2,X0,X2,X1) | cong(X2,X0,X0,X1))) => para(X0,X1,X5,X3)) [negated conjecture 94]
122. ! [X0,X1,X2,X3,X4] : ((midp(X4,X0,X2) & midp(X3,X0,X1)) => para(X3,X4,X1,X2)) [rectify 44]
131. ! [X0,X1,X2,X3] : ((coll(X3,X0,X2) & circle(X3,X0,X1,X2)) => perp(X0,X1,X1,X2)) [rectify 53]
132. ! [X0,X1,X2,X3] : ((perp(X3,X2,X0,X1) & midp(X2,X0,X1)) => cong(X3,X0,X3,X1)) [rectify 55]
165. ~! [X0,X1,X2,X3] : (! [X4,X5,X6,X7,X8,X9,X10] : (coll(X9,X4,X6) & circle(X8,X4,X9,X10) & circle(X8,X4,X5,X7) & coll(X7,X5,X6) & (cong(X6,X4,X6,X5) | cong(X6,X4,X4,X5))) => para(X3,X2,X1,X0)) [rectify 95]
223. ! [X0,X1,X2,X3,X4] : (para(X3,X4,X1,X2) | (~midp(X4,X0,X2) | ~midp(X3,X0,X1))) [ennf transformation 122]
224. ! [X0,X1,X2,X3,X4] : (para(X3,X4,X1,X2) | ~midp(X4,X0,X2) | ~midp(X3,X0,X1)) [flattening 223]
240. ! [X0,X1,X2,X3] : (perp(X0,X1,X1,X2) | (~coll(X3,X0,X2) | ~circle(X3,X0,X1,X2))) [ennf transformation 131]
241. ! [X0,X1,X2,X3] : (perp(X0,X1,X1,X2) | ~coll(X3,X0,X2) | ~circle(X3,X0,X1,X2)) [flattening 240]
244. ! [X0,X1,X2,X3] : (cong(X3,X0,X3,X1) | (~perp(X3,X2,X0,X1) | ~midp(X2,X0,X1))) [ennf transformation 132]
245. ! [X0,X1,X2,X3] : (cong(X3,X0,X3,X1) | ~perp(X3,X2,X0,X1) | ~midp(X2,X0,X1)) [flattening 244]
264. ! [X0,X1,X2] : (midp(X0,X1,X2) | (~coll(X0,X1,X2) | ~cong(X0,X1,X0,X2))) [ennf transformation 67]
265. ! [X0,X1,X2] : (midp(X0,X1,X2) | ~coll(X0,X1,X2) | ~cong(X0,X1,X0,X2)) [flattening 264]
315. ? [X0,X1,X2,X3] : (~para(X3,X2,X1,X0) & ! [X4,X5,X6,X7,X8,X9,X10] : (coll(X9,X4,X6) & circle(X8,X4,X9,X10) & circle(X8,X4,X5,X7) & coll(X7,X5,X6) & (cong(X6,X4,X6,X5) | cong(X6,X4,X4,X5)))) [ennf transformation 165]
352. ? [X0,X1,X2,X3] : (~para(X3,X2,X1,X0) & ! [X4,X5,X6,X7,X8,X9,X10] : (coll(X9,X4,X6) & circle(X8,X4,X9,X10) & circle(X8,X4,X5,X7) & coll(X7,X5,X6) & (cong(X6,X4,X6,X5) | cong(X6,X4,X4,X5)))) => (~para(sK23,sK22,sK21,sK20) & ! [X4,X5,X6,X7,X8,X9,X10] : (coll(X9,X4,X6) & circle(X8,X4,X9,X10) & circle(X8,X4,X5,X7) & coll(X7,X5,X6) & (cong(X6,X4,X6,X5) | cong(X6,X4,X4,X5)))) [choice axiom]
353. ~para(sK23,sK22,sK21,sK20) & ! [X4,X5,X6,X7,X8,X9,X10] : (coll(X9,X4,X6) & circle(X8,X4,X9,X10) & circle(X8,X4,X5,X7) & coll(X7,X5,X6) & (cong(X6,X4,X6,X5) | cong(X6,X4,X4,X5))) [skolemisation 315,352]
398. ~midp(X4,X0,X2) | para(X3,X4,X1,X2) | ~midp(X3,X0,X1) [cnf transformation 224]
407. ~circle(X3,X0,X1,X2) | ~coll(X3,X0,X2) | perp(X0,X1,X1,X2) [cnf transformation 241]
409. ~perp(X3,X2,X0,X1) | cong(X3,X0,X3,X1) | ~midp(X2,X0,X1) [cnf transformation 245]
421. ~cong(X0,X1,X0,X2) | ~coll(X0,X1,X2) | midp(X0,X1,X2) [cnf transformation 265]
467. cong(X6,X4,X6,X5) | cong(X6,X4,X4,X5) [cnf transformation 353]
470. circle(X8,X4,X9,X10) [cnf transformation 353]
471. coll(X9,X4,X6) [cnf transformation 353]
472. ~para(sK23,sK22,sK21,sK20) [cnf transformation 353]
475. cong(X0,X0,X0,X1) [factoring 467]
502. ~coll(X16,X16,X17) | midp(X16,X16,X17) [resolution 421,475]
506. midp(X16,X16,X17) [subsumption resolution 502,471]
508. para(X4,X5,X6,X7) | ~midp(X4,X5,X6) [resolution 506,398]
537. ~midp(sK23,sK22,sK21) [resolution 508,472]
546. ~coll(X0,X1,X2) | perp(X1,X3,X3,X2) [resolution 407,470]
547. perp(X1,X3,X3,X2) [subsumption resolution 546,471]
551. cong(X0,X1,X0,X2) | ~midp(X1,X1,X2) [resolution 409,547]
552. cong(X0,X1,X0,X2) [subsumption resolution 551,506]
586. ~coll(X6,X7,X8) | midp(X6,X7,X8) [resolution 552,421]
587. midp(X6,X7,X8) [subsumption resolution 586,471]
608. $false [resolution 587,537]
% SZS output end Proof for E040-28-1
% ------------------------------
% Version: Vampire 4.4.0 (commit 7916d27 on 2019-08-23 08:50:16 +0100)
% Termination reason: Refutation

% Memory used [KB]: 5117
% Time elapsed: 0.028 s
% ------------------------------
% ------------------------------
