% Refutation found. Thanks to Tanya!
% SZS status Theorem for E042-2
% SZS output start Proof for E042-2
fof(f23499,plain,(
  $false),
  inference(avatar_sat_refutation,[],[f866,f951,f2064,f2565,f3950,f5369,f5418,f7107,f7123,f7131,f8665,f9382,f9660,f9702,f9739,f9947,f10629,f10702,f10715,f13644,f13660,f13705,f13706,f13774,f13851,f23498])).
fof(f23498,plain,(
  ~spl28_16 | ~spl28_18 | ~spl28_24 | ~spl28_136 | ~spl28_520),
  inference(avatar_contradiction_clause,[],[f23483])).
fof(f23483,plain,(
  $false | (~spl28_16 | ~spl28_18 | ~spl28_24 | ~spl28_136 | ~spl28_520)),
  inference(resolution,[],[f23438,f5452])).
fof(f5452,plain,(
  ~midp(sK25,sK26,sK27) | ~spl28_24),
  inference(resolution,[],[f947,f479])).
fof(f479,plain,(
  ~cong(sK25,sK26,sK26,sK27)),
  inference(cnf_transformation,[],[f358])).
fof(f358,plain,(
  ~cong(sK25,sK26,sK26,sK27) & midp(sK27,sK24,sK22) & midp(sK26,sK22,sK21) & midp(sK25,sK21,sK23) & eqangle(sK22,sK20,sK20,sK24,sK20,sK24,sK24,sK22) & eqangle(sK22,sK20,sK20,sK24,sK20,sK22,sK22,sK24) & eqangle(sK21,sK20,sK20,sK23,sK20,sK23,sK23,sK21) & eqangle(sK21,sK20,sK20,sK23,sK20,sK21,sK21,sK23)),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK20,sK21,sK22,sK23,sK24,sK25,sK26,sK27])],[f320,f357])).
fof(f357,plain,(
  ? [X0,X1,X2,X3,X4,X5,X6,X7] : (~cong(X5,X6,X6,X7) & midp(X7,X4,X2) & midp(X6,X2,X1) & midp(X5,X1,X3) & eqangle(X2,X0,X0,X4,X0,X4,X4,X2) & eqangle(X2,X0,X0,X4,X0,X2,X2,X4) & eqangle(X1,X0,X0,X3,X0,X3,X3,X1) & eqangle(X1,X0,X0,X3,X0,X1,X1,X3)) => (~cong(sK25,sK26,sK26,sK27) & midp(sK27,sK24,sK22) & midp(sK26,sK22,sK21) & midp(sK25,sK21,sK23) & eqangle(sK22,sK20,sK20,sK24,sK20,sK24,sK24,sK22) & eqangle(sK22,sK20,sK20,sK24,sK20,sK22,sK22,sK24) & eqangle(sK21,sK20,sK20,sK23,sK20,sK23,sK23,sK21) & eqangle(sK21,sK20,sK20,sK23,sK20,sK21,sK21,sK23))),
  introduced(choice_axiom,[])).
fof(f320,plain,(
  ? [X0,X1,X2,X3,X4,X5,X6,X7] : (~cong(X5,X6,X6,X7) & midp(X7,X4,X2) & midp(X6,X2,X1) & midp(X5,X1,X3) & eqangle(X2,X0,X0,X4,X0,X4,X4,X2) & eqangle(X2,X0,X0,X4,X0,X2,X2,X4) & eqangle(X1,X0,X0,X3,X0,X3,X3,X1) & eqangle(X1,X0,X0,X3,X0,X1,X1,X3))),
  inference(flattening,[],[f319])).
fof(f319,plain,(
  ? [X0,X1,X2,X3,X4,X5,X6,X7] : (~cong(X5,X6,X6,X7) & (midp(X7,X4,X2) & midp(X6,X2,X1) & midp(X5,X1,X3) & eqangle(X2,X0,X0,X4,X0,X4,X4,X2) & eqangle(X2,X0,X0,X4,X0,X2,X2,X4) & eqangle(X1,X0,X0,X3,X0,X3,X3,X1) & eqangle(X1,X0,X0,X3,X0,X1,X1,X3)))),
  inference(ennf_transformation,[],[f167])).
fof(f167,plain,(
  ~! [X0,X1,X2,X3,X4,X5,X6,X7] : ((midp(X7,X4,X2) & midp(X6,X2,X1) & midp(X5,X1,X3) & eqangle(X2,X0,X0,X4,X0,X4,X4,X2) & eqangle(X2,X0,X0,X4,X0,X2,X2,X4) & eqangle(X1,X0,X0,X3,X0,X3,X3,X1) & eqangle(X1,X0,X0,X3,X0,X1,X1,X3)) => cong(X5,X6,X6,X7))),
  inference(rectify,[],[f96])).
fof(f96,negated_conjecture,(
  ~! [X0,X1,X2,X3,X4,X5,X12,X13] : ((midp(X13,X4,X2) & midp(X12,X2,X1) & midp(X5,X1,X3) & eqangle(X2,X0,X0,X4,X0,X4,X4,X2) & eqangle(X2,X0,X0,X4,X0,X2,X2,X4) & eqangle(X1,X0,X0,X3,X0,X3,X3,X1) & eqangle(X1,X0,X0,X3,X0,X1,X1,X3)) => cong(X5,X12,X12,X13))),
  inference(negated_conjecture,[],[f95])).
fof(f95,conjecture,(
  ! [X0,X1,X2,X3,X4,X5,X12,X13] : ((midp(X13,X4,X2) & midp(X12,X2,X1) & midp(X5,X1,X3) & eqangle(X2,X0,X0,X4,X0,X4,X4,X2) & eqangle(X2,X0,X0,X4,X0,X2,X2,X4) & eqangle(X1,X0,X0,X3,X0,X3,X3,X1) & eqangle(X1,X0,X0,X3,X0,X1,X1,X3)) => cong(X5,X12,X12,X13))),
  file('./7_Book/00EE/03/E042-2.p',exemplo7Book00EE03E0422)).
fof(f947,plain,(
  ( ! [X10,X11,X9] : (cong(X9,X10,X10,X11) | ~midp(X9,X10,X11)) ) | ~spl28_24),
  inference(avatar_component_clause,[],[f946])).
fof(f946,plain,(
  spl28_24 <=> ! [X9,X11,X10] : (cong(X9,X10,X10,X11) | ~midp(X9,X10,X11))),
  introduced(avatar_definition,[new_symbols(naming,[spl28_24])])).
fof(f23438,plain,(
  ( ! [X6,X4,X5] : (midp(X4,X5,X6)) ) | (~spl28_16 | ~spl28_18 | ~spl28_136 | ~spl28_520)),
  inference(subsumption_resolution,[],[f23437,f21777])).
fof(f21777,plain,(
  ( ! [X0,X1] : (midp(X0,X0,X1)) ) | (~spl28_16 | ~spl28_18 | ~spl28_136 | ~spl28_520)),
  inference(resolution,[],[f21698,f369])).
fof(f369,plain,(
  ( ! [X2,X0,X1] : (~midp(X2,X1,X0) | midp(X2,X0,X1)) )),
  inference(cnf_transformation,[],[f182])).
fof(f182,plain,(
  ! [X0,X1,X2] : (midp(X2,X0,X1) | ~midp(X2,X1,X0))),
  inference(ennf_transformation,[],[f97])).
fof(f97,plain,(
  ! [X0,X1,X2] : (midp(X2,X1,X0) => midp(X2,X0,X1))),
  inference(rectify,[],[f11])).
fof(f11,axiom,(
  ! [X0,X1,X6] : (midp(X6,X1,X0) => midp(X6,X0,X1))),
  file('./7_Book/00EE/03/E042-2.p',ruleD11)).
fof(f21698,plain,(
  ( ! [X47,X46] : (midp(X46,X47,X46)) ) | (~spl28_16 | ~spl28_18 | ~spl28_136 | ~spl28_520)),
  inference(subsumption_resolution,[],[f21697,f2060])).
fof(f2060,plain,(
  ( ! [X2,X3] : (para(X2,X3,X2,X3)) ) | ~spl28_136),
  inference(avatar_component_clause,[],[f2059])).
fof(f2059,plain,(
  spl28_136 <=> ! [X3,X2] : para(X2,X3,X2,X3)),
  introduced(avatar_definition,[new_symbols(naming,[spl28_136])])).
fof(f21697,plain,(
  ( ! [X47,X46] : (midp(X46,X47,X46) | ~para(sK22,X47,sK22,X47)) ) | (~spl28_16 | ~spl28_18 | ~spl28_136 | ~spl28_520)),
  inference(subsumption_resolution,[],[f21693,f6566])).
fof(f6566,plain,(
  ( ! [X4,X3] : (coll(X3,X4,X3)) ) | ~spl28_136),
  inference(resolution,[],[f6550,f360])).
fof(f360,plain,(
  ( ! [X2,X0,X1] : (~coll(X0,X1,X2) | coll(X1,X0,X2)) )),
  inference(cnf_transformation,[],[f169])).
fof(f169,plain,(
  ! [X0,X1,X2] : (coll(X1,X0,X2) | ~coll(X0,X1,X2))),
  inference(ennf_transformation,[],[f2])).
fof(f2,axiom,(
  ! [X0,X1,X2] : (coll(X0,X1,X2) => coll(X1,X0,X2))),
  file('./7_Book/00EE/03/E042-2.p',ruleD2)).
fof(f6550,plain,(
  ( ! [X4,X3] : (coll(X3,X4,X4)) ) | ~spl28_136),
  inference(resolution,[],[f2060,f425])).
fof(f425,plain,(
  ( ! [X2,X0,X1] : (~para(X0,X1,X0,X2) | coll(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f267])).
fof(f267,plain,(
  ! [X0,X1,X2] : (coll(X0,X1,X2) | ~para(X0,X1,X0,X2))),
  inference(ennf_transformation,[],[f67])).
fof(f67,axiom,(
  ! [X0,X1,X2] : (para(X0,X1,X0,X2) => coll(X0,X1,X2))),
  file('./7_Book/00EE/03/E042-2.p',ruleD66)).
fof(f21693,plain,(
  ( ! [X47,X46] : (~coll(X46,X47,X46) | midp(X46,X47,X46) | ~para(sK22,X47,sK22,X47)) ) | (~spl28_16 | ~spl28_18 | ~spl28_136 | ~spl28_520)),
  inference(resolution,[],[f19224,f1005])).
fof(f1005,plain,(
  ( ! [X2,X0,X1] : (~circle(X2,X2,X1,X0) | ~coll(X0,X1,X0) | midp(X0,X1,X0) | ~para(X2,X1,X2,X1)) )),
  inference(resolution,[],[f410,f398])).
fof(f398,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (eqangle(X0,X1,X4,X5,X2,X3,X4,X5) | ~para(X0,X1,X2,X3)) )),
  inference(cnf_transformation,[],[f219])).
fof(f219,plain,(
  ! [X0,X1,X2,X3,X4,X5] : (eqangle(X0,X1,X4,X5,X2,X3,X4,X5) | ~para(X0,X1,X2,X3))),
  inference(ennf_transformation,[],[f119])).
fof(f119,plain,(
  ! [X0,X1,X2,X3,X4,X5] : (para(X0,X1,X2,X3) => eqangle(X0,X1,X4,X5,X2,X3,X4,X5))),
  inference(rectify,[],[f40])).
fof(f40,axiom,(
  ! [X0,X1,X2,X3,X8,X9] : (para(X0,X1,X2,X3) => eqangle(X0,X1,X8,X9,X2,X3,X8,X9))),
  file('./7_Book/00EE/03/E042-2.p',ruleD40)).
fof(f410,plain,(
  ( ! [X4,X2,X0,X3,X1] : (~eqangle(X0,X1,X0,X2,X3,X1,X3,X4) | midp(X4,X1,X2) | ~coll(X4,X1,X2) | ~circle(X3,X0,X1,X2)) )),
  inference(cnf_transformation,[],[f241])).
fof(f241,plain,(
  ! [X0,X1,X2,X3,X4] : (midp(X4,X1,X2) | ~eqangle(X0,X1,X0,X2,X3,X1,X3,X4) | ~coll(X4,X1,X2) | ~circle(X3,X0,X1,X2))),
  inference(flattening,[],[f240])).
fof(f240,plain,(
  ! [X0,X1,X2,X3,X4] : (midp(X4,X1,X2) | (~eqangle(X0,X1,X0,X2,X3,X1,X3,X4) | ~coll(X4,X1,X2) | ~circle(X3,X0,X1,X2)))),
  inference(ennf_transformation,[],[f131])).
fof(f131,plain,(
  ! [X0,X1,X2,X3,X4] : ((eqangle(X0,X1,X0,X2,X3,X1,X3,X4) & coll(X4,X1,X2) & circle(X3,X0,X1,X2)) => midp(X4,X1,X2))),
  inference(rectify,[],[f52])).
fof(f52,axiom,(
  ! [X0,X1,X2,X7,X6] : ((eqangle(X0,X1,X0,X2,X7,X1,X7,X6) & coll(X6,X1,X2) & circle(X7,X0,X1,X2)) => midp(X6,X1,X2))),
  file('./7_Book/00EE/03/E042-2.p',ruleD51)).
fof(f19224,plain,(
  ( ! [X24,X25] : (circle(sK22,sK22,X24,X25)) ) | (~spl28_16 | ~spl28_18 | ~spl28_136 | ~spl28_520)),
  inference(subsumption_resolution,[],[f19192,f16055])).
fof(f16055,plain,(
  ( ! [X33,X34] : (cong(sK22,X33,X33,X34)) ) | (~spl28_16 | ~spl28_18 | ~spl28_136 | ~spl28_520)),
  inference(subsumption_resolution,[],[f16054,f6712])).
fof(f6712,plain,(
  ( ! [X2,X0,X1] : (coll(X0,X1,X2)) ) | ~spl28_136),
  inference(subsumption_resolution,[],[f6697,f6628])).
fof(f6628,plain,(
  ( ! [X6,X5] : (coll(X5,X5,X6)) ) | ~spl28_136),
  inference(resolution,[],[f6566,f359])).
fof(f359,plain,(
  ( ! [X2,X0,X1] : (~coll(X0,X1,X2) | coll(X0,X2,X1)) )),
  inference(cnf_transformation,[],[f168])).
fof(f168,plain,(
  ! [X0,X1,X2] : (coll(X0,X2,X1) | ~coll(X0,X1,X2))),
  inference(ennf_transformation,[],[f1])).
fof(f1,axiom,(
  ! [X0,X1,X2] : (coll(X0,X1,X2) => coll(X0,X2,X1))),
  file('./7_Book/00EE/03/E042-2.p',ruleD1)).
fof(f6697,plain,(
  ( ! [X2,X0,X1] : (coll(X0,X1,X2) | ~coll(X2,X2,X0)) ) | ~spl28_136),
  inference(resolution,[],[f6628,f361])).
fof(f361,plain,(
  ( ! [X2,X0,X3,X1] : (~coll(X0,X1,X3) | coll(X2,X3,X0) | ~coll(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f171])).
fof(f171,plain,(
  ! [X0,X1,X2,X3] : (coll(X2,X3,X0) | ~coll(X0,X1,X3) | ~coll(X0,X1,X2))),
  inference(flattening,[],[f170])).
fof(f170,plain,(
  ! [X0,X1,X2,X3] : (coll(X2,X3,X0) | (~coll(X0,X1,X3) | ~coll(X0,X1,X2)))),
  inference(ennf_transformation,[],[f3])).
fof(f3,axiom,(
  ! [X0,X1,X2,X3] : ((coll(X0,X1,X3) & coll(X0,X1,X2)) => coll(X2,X3,X0))),
  file('./7_Book/00EE/03/E042-2.p',ruleD3)).
fof(f16054,plain,(
  ( ! [X35,X33,X34] : (~coll(sK22,X35,sK20) | cong(sK22,X33,X33,X34)) ) | (~spl28_16 | ~spl28_18 | ~spl28_136 | ~spl28_520)),
  inference(subsumption_resolution,[],[f16053,f15970])).
fof(f15970,plain,(
  ( ! [X28,X26,X27,X25] : (para(X25,X26,X27,X28)) ) | (~spl28_16 | ~spl28_136 | ~spl28_520)),
  inference(subsumption_resolution,[],[f15955,f15954])).
fof(f15954,plain,(
  ( ! [X24,X23] : (perp(X23,X24,sK20,sK23)) ) | (~spl28_16 | ~spl28_136 | ~spl28_520)),
  inference(resolution,[],[f14168,f366])).
fof(f366,plain,(
  ( ! [X2,X0,X3,X1] : (~perp(X0,X1,X2,X3) | perp(X2,X3,X0,X1)) )),
  inference(cnf_transformation,[],[f177])).
fof(f177,plain,(
  ! [X0,X1,X2,X3] : (perp(X2,X3,X0,X1) | ~perp(X0,X1,X2,X3))),
  inference(ennf_transformation,[],[f8])).
fof(f8,axiom,(
  ! [X0,X1,X2,X3] : (perp(X0,X1,X2,X3) => perp(X2,X3,X0,X1))),
  file('./7_Book/00EE/03/E042-2.p',ruleD8)).
fof(f14168,plain,(
  ( ! [X15,X16] : (perp(sK20,sK23,X15,X16)) ) | (~spl28_16 | ~spl28_136 | ~spl28_520)),
  inference(subsumption_resolution,[],[f14164,f13945])).
fof(f13945,plain,(
  ( ! [X6] : (cong(sK20,X6,sK23,X6)) ) | (~spl28_16 | ~spl28_136 | ~spl28_520)),
  inference(subsumption_resolution,[],[f13938,f10716])).
fof(f10716,plain,(
  ( ! [X0] : (midp(X0,sK20,sK21)) ) | (~spl28_136 | ~spl28_520)),
  inference(subsumption_resolution,[],[f8664,f9399])).
fof(f9399,plain,(
  ( ! [X80] : (midp(X80,sK23,sK21)) ) | ~spl28_136),
  inference(subsumption_resolution,[],[f9374,f9319])).
fof(f9319,plain,(
  ( ! [X21,X22] : (midp(X21,X22,X22)) ) | ~spl28_136),
  inference(subsumption_resolution,[],[f9290,f6550])).
fof(f9290,plain,(
  ( ! [X21,X22] : (~coll(X21,X22,X22) | midp(X21,X22,X22)) ) | ~spl28_136),
  inference(resolution,[],[f8840,f426])).
fof(f426,plain,(
  ( ! [X2,X0,X1] : (~cong(X0,X1,X0,X2) | ~coll(X0,X1,X2) | midp(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f269])).
fof(f269,plain,(
  ! [X0,X1,X2] : (midp(X0,X1,X2) | ~coll(X0,X1,X2) | ~cong(X0,X1,X0,X2))),
  inference(flattening,[],[f268])).
fof(f268,plain,(
  ! [X0,X1,X2] : (midp(X0,X1,X2) | (~coll(X0,X1,X2) | ~cong(X0,X1,X0,X2)))),
  inference(ennf_transformation,[],[f68])).
fof(f68,axiom,(
  ! [X0,X1,X2] : ((coll(X0,X1,X2) & cong(X0,X1,X0,X2)) => midp(X0,X1,X2))),
  file('./7_Book/00EE/03/E042-2.p',ruleD67)).
fof(f8840,plain,(
  ( ! [X35,X36] : (cong(X35,X36,X35,X36)) ) | ~spl28_136),
  inference(subsumption_resolution,[],[f8839,f8592])).
fof(f8592,plain,(
  ( ! [X4,X5,X3] : (cyclic(X3,X4,X3,X5)) ) | ~spl28_136),
  inference(resolution,[],[f8581,f373])).
fof(f373,plain,(
  ( ! [X2,X0,X3,X1] : (~cyclic(X0,X1,X2,X3) | cyclic(X0,X2,X1,X3)) )),
  inference(cnf_transformation,[],[f188])).
fof(f188,plain,(
  ! [X0,X1,X2,X3] : (cyclic(X0,X2,X1,X3) | ~cyclic(X0,X1,X2,X3))),
  inference(ennf_transformation,[],[f15])).
fof(f15,axiom,(
  ! [X0,X1,X2,X3] : (cyclic(X0,X1,X2,X3) => cyclic(X0,X2,X1,X3))),
  file('./7_Book/00EE/03/E042-2.p',ruleD15)).
fof(f8581,plain,(
  ( ! [X2,X0,X1] : (cyclic(X0,X0,X1,X2)) ) | ~spl28_136),
  inference(resolution,[],[f4115,f2060])).
fof(f4115,plain,(
  ( ! [X68,X66,X67] : (~para(X66,X67,X66,X67) | cyclic(X67,X67,X66,X68)) )),
  inference(subsumption_resolution,[],[f2202,f2203])).
fof(f2203,plain,(
  ( ! [X70,X71,X69] : (~para(X69,X70,X69,X70) | coll(X69,X71,X70) | cyclic(X70,X70,X69,X71)) )),
  inference(resolution,[],[f688,f400])).
fof(f400,plain,(
  ( ! [X2,X0,X3,X1] : (~eqangle(X2,X0,X2,X1,X3,X0,X3,X1) | coll(X2,X3,X0) | cyclic(X0,X1,X2,X3)) )),
  inference(cnf_transformation,[],[f222])).
fof(f222,plain,(
  ! [X0,X1,X2,X3] : (cyclic(X0,X1,X2,X3) | coll(X2,X3,X0) | ~eqangle(X2,X0,X2,X1,X3,X0,X3,X1))),
  inference(flattening,[],[f221])).
fof(f221,plain,(
  ! [X0,X1,X2,X3] : (cyclic(X0,X1,X2,X3) | (coll(X2,X3,X0) | ~eqangle(X2,X0,X2,X1,X3,X0,X3,X1)))),
  inference(ennf_transformation,[],[f121])).
fof(f121,plain,(
  ! [X0,X1,X2,X3] : ((~coll(X2,X3,X0) & eqangle(X2,X0,X2,X1,X3,X0,X3,X1)) => cyclic(X0,X1,X2,X3))),
  inference(rectify,[],[f42])).
fof(f42,axiom,(
  ! [X0,X1,X8,X9] : ((~coll(X8,X9,X0) & eqangle(X8,X0,X8,X1,X9,X0,X9,X1)) => cyclic(X0,X1,X8,X9))),
  file('./7_Book/00EE/03/E042-2.p',ruleD42a)).
fof(f688,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (eqangle(X0,X1,X2,X3,X4,X5,X4,X5) | ~para(X0,X1,X2,X3)) )),
  inference(resolution,[],[f379,f398])).
fof(f379,plain,(
  ( ! [X6,X4,X2,X0,X7,X5,X3,X1] : (~eqangle(X0,X1,X2,X3,X4,X5,X6,X7) | eqangle(X0,X1,X4,X5,X2,X3,X6,X7)) )),
  inference(cnf_transformation,[],[f195])).
fof(f195,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6,X7] : (eqangle(X0,X1,X4,X5,X2,X3,X6,X7) | ~eqangle(X0,X1,X2,X3,X4,X5,X6,X7))),
  inference(ennf_transformation,[],[f103])).
fof(f103,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6,X7] : (eqangle(X0,X1,X2,X3,X4,X5,X6,X7) => eqangle(X0,X1,X4,X5,X2,X3,X6,X7))),
  inference(rectify,[],[f21])).
fof(f21,axiom,(
  ! [X0,X1,X2,X3,X8,X9,X10,X11] : (eqangle(X0,X1,X2,X3,X8,X9,X10,X11) => eqangle(X0,X1,X8,X9,X2,X3,X10,X11))),
  file('./7_Book/00EE/03/E042-2.p',ruleD21)).
fof(f2202,plain,(
  ( ! [X68,X66,X67] : (~para(X66,X67,X66,X67) | ~coll(X66,X68,X67) | cyclic(X67,X67,X66,X68)) )),
  inference(resolution,[],[f688,f401])).
fof(f401,plain,(
  ( ! [X2,X0,X3,X1] : (~eqangle(X2,X0,X2,X1,X3,X0,X3,X1) | ~coll(X2,X3,X1) | cyclic(X0,X1,X2,X3)) )),
  inference(cnf_transformation,[],[f224])).
fof(f224,plain,(
  ! [X0,X1,X2,X3] : (cyclic(X0,X1,X2,X3) | ~coll(X2,X3,X1) | ~eqangle(X2,X0,X2,X1,X3,X0,X3,X1))),
  inference(flattening,[],[f223])).
fof(f223,plain,(
  ! [X0,X1,X2,X3] : (cyclic(X0,X1,X2,X3) | (~coll(X2,X3,X1) | ~eqangle(X2,X0,X2,X1,X3,X0,X3,X1)))),
  inference(ennf_transformation,[],[f122])).
fof(f122,plain,(
  ! [X0,X1,X2,X3] : ((coll(X2,X3,X1) & eqangle(X2,X0,X2,X1,X3,X0,X3,X1)) => cyclic(X0,X1,X2,X3))),
  inference(rectify,[],[f43])).
fof(f43,axiom,(
  ! [X0,X1,X8,X9] : ((coll(X8,X9,X1) & eqangle(X8,X0,X8,X1,X9,X0,X9,X1)) => cyclic(X0,X1,X8,X9))),
  file('./7_Book/00EE/03/E042-2.p',ruleD42b)).
fof(f8839,plain,(
  ( ! [X35,X36] : (cong(X35,X36,X35,X36) | ~cyclic(X35,X36,X35,X35)) ) | ~spl28_136),
  inference(subsumption_resolution,[],[f8825,f8592])).
fof(f8825,plain,(
  ( ! [X35,X36] : (cong(X35,X36,X35,X36) | ~cyclic(X35,X36,X35,X36) | ~cyclic(X35,X36,X35,X35)) ) | ~spl28_136),
  inference(resolution,[],[f8592,f1199])).
fof(f1199,plain,(
  ( ! [X10,X8,X7,X9] : (~cyclic(X7,X8,X9,X10) | cong(X7,X8,X7,X8) | ~cyclic(X7,X8,X9,X8) | ~cyclic(X7,X8,X9,X7)) )),
  inference(duplicate_literal_removal,[],[f1193])).
fof(f1193,plain,(
  ( ! [X10,X8,X7,X9] : (cong(X7,X8,X7,X8) | ~cyclic(X7,X8,X9,X10) | ~cyclic(X7,X8,X9,X8) | ~cyclic(X7,X8,X9,X7) | ~cyclic(X7,X8,X9,X10)) )),
  inference(resolution,[],[f402,f399])).
fof(f399,plain,(
  ( ! [X2,X0,X3,X1] : (eqangle(X2,X0,X2,X1,X3,X0,X3,X1) | ~cyclic(X0,X1,X2,X3)) )),
  inference(cnf_transformation,[],[f220])).
fof(f220,plain,(
  ! [X0,X1,X2,X3] : (eqangle(X2,X0,X2,X1,X3,X0,X3,X1) | ~cyclic(X0,X1,X2,X3))),
  inference(ennf_transformation,[],[f120])).
fof(f120,plain,(
  ! [X0,X1,X2,X3] : (cyclic(X0,X1,X2,X3) => eqangle(X2,X0,X2,X1,X3,X0,X3,X1))),
  inference(rectify,[],[f41])).
fof(f41,axiom,(
  ! [X0,X1,X8,X9] : (cyclic(X0,X1,X8,X9) => eqangle(X8,X0,X8,X1,X9,X0,X9,X1))),
  file('./7_Book/00EE/03/E042-2.p',ruleD41)).
fof(f402,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (~eqangle(X2,X0,X2,X1,X5,X3,X5,X4) | cong(X0,X1,X3,X4) | ~cyclic(X0,X1,X2,X5) | ~cyclic(X0,X1,X2,X4) | ~cyclic(X0,X1,X2,X3)) )),
  inference(cnf_transformation,[],[f226])).
fof(f226,plain,(
  ! [X0,X1,X2,X3,X4,X5] : (cong(X0,X1,X3,X4) | ~eqangle(X2,X0,X2,X1,X5,X3,X5,X4) | ~cyclic(X0,X1,X2,X5) | ~cyclic(X0,X1,X2,X4) | ~cyclic(X0,X1,X2,X3))),
  inference(flattening,[],[f225])).
fof(f225,plain,(
  ! [X0,X1,X2,X3,X4,X5] : (cong(X0,X1,X3,X4) | (~eqangle(X2,X0,X2,X1,X5,X3,X5,X4) | ~cyclic(X0,X1,X2,X5) | ~cyclic(X0,X1,X2,X4) | ~cyclic(X0,X1,X2,X3)))),
  inference(ennf_transformation,[],[f123])).
fof(f123,plain,(
  ! [X0,X1,X2,X3,X4,X5] : ((eqangle(X2,X0,X2,X1,X5,X3,X5,X4) & cyclic(X0,X1,X2,X5) & cyclic(X0,X1,X2,X4) & cyclic(X0,X1,X2,X3)) => cong(X0,X1,X3,X4))),
  inference(rectify,[],[f44])).
fof(f44,axiom,(
  ! [X0,X1,X2,X8,X9,X14] : ((eqangle(X2,X0,X2,X1,X14,X8,X14,X9) & cyclic(X0,X1,X2,X14) & cyclic(X0,X1,X2,X9) & cyclic(X0,X1,X2,X8)) => cong(X0,X1,X8,X9))),
  file('./7_Book/00EE/03/E042-2.p',ruleD43)).
fof(f9374,plain,(
  ( ! [X80,X81] : (midp(X80,sK23,sK21) | ~midp(X80,X81,X81)) ) | ~spl28_136),
  inference(resolution,[],[f9319,f1273])).
fof(f1273,plain,(
  ( ! [X2,X3,X1] : (~midp(sK25,X1,X2) | midp(X3,sK23,sK21) | ~midp(X3,X1,X2)) )),
  inference(subsumption_resolution,[],[f1267,f549])).
fof(f549,plain,(
  ( ! [X2,X3] : (para(X2,sK23,X3,sK21) | ~midp(sK25,X2,X3)) )),
  inference(resolution,[],[f422,f480])).
fof(f480,plain,(
  midp(sK25,sK23,sK21)),
  inference(resolution,[],[f369,f476])).
fof(f476,plain,(
  midp(sK25,sK21,sK23)),
  inference(cnf_transformation,[],[f358])).
fof(f422,plain,(
  ( ! [X4,X2,X0,X3,X1] : (~midp(X4,X2,X3) | para(X0,X2,X1,X3) | ~midp(X4,X0,X1)) )),
  inference(cnf_transformation,[],[f262])).
fof(f262,plain,(
  ! [X0,X1,X2,X3,X4] : (para(X0,X2,X1,X3) | ~midp(X4,X2,X3) | ~midp(X4,X0,X1))),
  inference(flattening,[],[f261])).
fof(f261,plain,(
  ! [X0,X1,X2,X3,X4] : (para(X0,X2,X1,X3) | (~midp(X4,X2,X3) | ~midp(X4,X0,X1)))),
  inference(ennf_transformation,[],[f142])).
fof(f142,plain,(
  ! [X0,X1,X2,X3,X4] : ((midp(X4,X2,X3) & midp(X4,X0,X1)) => para(X0,X2,X1,X3))),
  inference(rectify,[],[f64])).
fof(f64,axiom,(
  ! [X0,X1,X2,X3,X6] : ((midp(X6,X2,X3) & midp(X6,X0,X1)) => para(X0,X2,X1,X3))),
  file('./7_Book/00EE/03/E042-2.p',ruleD63)).
fof(f1267,plain,(
  ( ! [X2,X3,X1] : (~midp(sK25,X1,X2) | midp(X3,sK23,sK21) | ~para(X1,sK23,X2,sK21) | ~midp(X3,X1,X2)) )),
  inference(resolution,[],[f548,f423])).
fof(f423,plain,(
  ( ! [X4,X2,X0,X3,X1] : (~para(X0,X3,X1,X2) | midp(X4,X2,X3) | ~para(X0,X2,X1,X3) | ~midp(X4,X0,X1)) )),
  inference(cnf_transformation,[],[f264])).
fof(f264,plain,(
  ! [X0,X1,X2,X3,X4] : (midp(X4,X2,X3) | ~para(X0,X3,X1,X2) | ~para(X0,X2,X1,X3) | ~midp(X4,X0,X1))),
  inference(flattening,[],[f263])).
fof(f263,plain,(
  ! [X0,X1,X2,X3,X4] : (midp(X4,X2,X3) | (~para(X0,X3,X1,X2) | ~para(X0,X2,X1,X3) | ~midp(X4,X0,X1)))),
  inference(ennf_transformation,[],[f143])).
fof(f143,plain,(
  ! [X0,X1,X2,X3,X4] : ((para(X0,X3,X1,X2) & para(X0,X2,X1,X3) & midp(X4,X0,X1)) => midp(X4,X2,X3))),
  inference(rectify,[],[f65])).
fof(f65,axiom,(
  ! [X0,X1,X2,X3,X6] : ((para(X0,X3,X1,X2) & para(X0,X2,X1,X3) & midp(X6,X0,X1)) => midp(X6,X2,X3))),
  file('./7_Book/00EE/03/E042-2.p',ruleD64)).
fof(f548,plain,(
  ( ! [X0,X1] : (para(X0,sK21,X1,sK23) | ~midp(sK25,X0,X1)) )),
  inference(resolution,[],[f422,f476])).
fof(f8664,plain,(
  ( ! [X0] : (midp(X0,sK20,sK21) | ~midp(X0,sK23,sK21)) ) | ~spl28_520),
  inference(avatar_component_clause,[],[f8663])).
fof(f8663,plain,(
  spl28_520 <=> ! [X0] : (midp(X0,sK20,sK21) | ~midp(X0,sK23,sK21))),
  introduced(avatar_definition,[new_symbols(naming,[spl28_520])])).
fof(f13938,plain,(
  ( ! [X6] : (~midp(X6,sK20,sK21) | cong(sK20,X6,sK23,X6)) ) | ~spl28_16),
  inference(resolution,[],[f915,f411])).
fof(f411,plain,(
  ( ! [X2,X0,X3,X1] : (~perp(X0,X1,X1,X2) | ~midp(X3,X0,X2) | cong(X0,X3,X1,X3)) )),
  inference(cnf_transformation,[],[f243])).
fof(f243,plain,(
  ! [X0,X1,X2,X3] : (cong(X0,X3,X1,X3) | ~midp(X3,X0,X2) | ~perp(X0,X1,X1,X2))),
  inference(flattening,[],[f242])).
fof(f242,plain,(
  ! [X0,X1,X2,X3] : (cong(X0,X3,X1,X3) | (~midp(X3,X0,X2) | ~perp(X0,X1,X1,X2)))),
  inference(ennf_transformation,[],[f132])).
fof(f132,plain,(
  ! [X0,X1,X2,X3] : ((midp(X3,X0,X2) & perp(X0,X1,X1,X2)) => cong(X0,X3,X1,X3))),
  inference(rectify,[],[f53])).
fof(f53,axiom,(
  ! [X0,X1,X2,X6] : ((midp(X6,X0,X2) & perp(X0,X1,X1,X2)) => cong(X0,X6,X1,X6))),
  file('./7_Book/00EE/03/E042-2.p',ruleD52)).
fof(f915,plain,(
  perp(sK20,sK23,sK23,sK21) | ~spl28_16),
  inference(avatar_component_clause,[],[f914])).
fof(f914,plain,(
  spl28_16 <=> perp(sK20,sK23,sK23,sK21)),
  introduced(avatar_definition,[new_symbols(naming,[spl28_16])])).
fof(f14164,plain,(
  ( ! [X15,X16] : (perp(sK20,sK23,X15,X16) | ~cong(sK20,X15,sK23,X15)) ) | (~spl28_16 | ~spl28_136 | ~spl28_520)),
  inference(resolution,[],[f13945,f415])).
fof(f415,plain,(
  ( ! [X2,X0,X3,X1] : (~cong(X0,X3,X1,X3) | perp(X0,X1,X2,X3) | ~cong(X0,X2,X1,X2)) )),
  inference(cnf_transformation,[],[f251])).
fof(f251,plain,(
  ! [X0,X1,X2,X3] : (perp(X0,X1,X2,X3) | ~cong(X0,X3,X1,X3) | ~cong(X0,X2,X1,X2))),
  inference(flattening,[],[f250])).
fof(f250,plain,(
  ! [X0,X1,X2,X3] : (perp(X0,X1,X2,X3) | (~cong(X0,X3,X1,X3) | ~cong(X0,X2,X1,X2)))),
  inference(ennf_transformation,[],[f135])).
fof(f135,plain,(
  ! [X0,X1,X2,X3] : ((cong(X0,X3,X1,X3) & cong(X0,X2,X1,X2)) => perp(X0,X1,X2,X3))),
  inference(rectify,[],[f57])).
fof(f57,axiom,(
  ! [X0,X1,X8,X9] : ((cong(X0,X9,X1,X9) & cong(X0,X8,X1,X8)) => perp(X0,X1,X8,X9))),
  file('./7_Book/00EE/03/E042-2.p',ruleD56)).
fof(f15955,plain,(
  ( ! [X28,X26,X27,X25] : (para(X25,X26,X27,X28) | ~perp(X25,X26,sK20,sK23)) ) | (~spl28_16 | ~spl28_136 | ~spl28_520)),
  inference(resolution,[],[f14168,f367])).
fof(f367,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (~perp(X2,X3,X4,X5) | para(X0,X1,X4,X5) | ~perp(X0,X1,X2,X3)) )),
  inference(cnf_transformation,[],[f179])).
fof(f179,plain,(
  ! [X0,X1,X2,X3,X4,X5] : (para(X0,X1,X4,X5) | ~perp(X2,X3,X4,X5) | ~perp(X0,X1,X2,X3))),
  inference(flattening,[],[f178])).
fof(f178,plain,(
  ! [X0,X1,X2,X3,X4,X5] : (para(X0,X1,X4,X5) | (~perp(X2,X3,X4,X5) | ~perp(X0,X1,X2,X3)))),
  inference(ennf_transformation,[],[f9])).
fof(f9,axiom,(
  ! [X0,X1,X2,X3,X4,X5] : ((perp(X2,X3,X4,X5) & perp(X0,X1,X2,X3)) => para(X0,X1,X4,X5))),
  file('./7_Book/00EE/03/E042-2.p',ruleD9)).
fof(f16053,plain,(
  ( ! [X35,X33,X34] : (~para(X33,X35,X34,sK20) | ~coll(sK22,X35,sK20) | cong(sK22,X33,X33,X34)) ) | (~spl28_18 | ~spl28_136)),
  inference(subsumption_resolution,[],[f16050,f6712])).
fof(f16050,plain,(
  ( ! [X35,X33,X34] : (~coll(sK22,X33,X34) | ~para(X33,X35,X34,sK20) | ~coll(sK22,X35,sK20) | cong(sK22,X33,X33,X34)) ) | (~spl28_18 | ~spl28_136)),
  inference(resolution,[],[f14189,f1016])).
fof(f1016,plain,(
  ( ! [X4,X2,X0,X3,X1] : (~cong(X0,X1,X1,X2) | ~coll(X0,X3,X4) | ~para(X3,X1,X4,X2) | ~coll(X0,X1,X2) | cong(X0,X3,X3,X4)) )),
  inference(resolution,[],[f424,f434])).
fof(f434,plain,(
  ( ! [X6,X4,X2,X0,X7,X5,X3,X1] : (~eqratio(X0,X1,X2,X3,X4,X5,X6,X7) | ~cong(X4,X5,X6,X7) | cong(X0,X1,X2,X3)) )),
  inference(cnf_transformation,[],[f283])).
fof(f283,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6,X7] : (cong(X0,X1,X2,X3) | ~cong(X4,X5,X6,X7) | ~eqratio(X0,X1,X2,X3,X4,X5,X6,X7))),
  inference(flattening,[],[f282])).
fof(f282,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6,X7] : (cong(X0,X1,X2,X3) | (~cong(X4,X5,X6,X7) | ~eqratio(X0,X1,X2,X3,X4,X5,X6,X7)))),
  inference(ennf_transformation,[],[f148])).
fof(f148,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6,X7] : ((cong(X4,X5,X6,X7) & eqratio(X0,X1,X2,X3,X4,X5,X6,X7)) => cong(X0,X1,X2,X3))),
  inference(rectify,[],[f76])).
fof(f76,axiom,(
  ! [X0,X1,X2,X3,X8,X9,X10,X11] : ((cong(X8,X9,X10,X11) & eqratio(X0,X1,X2,X3,X8,X9,X10,X11)) => cong(X0,X1,X2,X3))),
  file('./7_Book/00EE/03/E042-2.p',ruleD75)).
fof(f424,plain,(
  ( ! [X4,X2,X0,X3,X1] : (eqratio(X4,X0,X0,X2,X4,X1,X1,X3) | ~coll(X4,X1,X3) | ~coll(X4,X0,X2) | ~para(X0,X1,X2,X3)) )),
  inference(cnf_transformation,[],[f266])).
fof(f266,plain,(
  ! [X0,X1,X2,X3,X4] : (eqratio(X4,X0,X0,X2,X4,X1,X1,X3) | ~coll(X4,X1,X3) | ~coll(X4,X0,X2) | ~para(X0,X1,X2,X3))),
  inference(flattening,[],[f265])).
fof(f265,plain,(
  ! [X0,X1,X2,X3,X4] : (eqratio(X4,X0,X0,X2,X4,X1,X1,X3) | (~coll(X4,X1,X3) | ~coll(X4,X0,X2) | ~para(X0,X1,X2,X3)))),
  inference(ennf_transformation,[],[f144])).
fof(f144,plain,(
  ! [X0,X1,X2,X3,X4] : ((coll(X4,X1,X3) & coll(X4,X0,X2) & para(X0,X1,X2,X3)) => eqratio(X4,X0,X0,X2,X4,X1,X1,X3))),
  inference(rectify,[],[f66])).
fof(f66,axiom,(
  ! [X0,X1,X2,X3,X7] : ((coll(X7,X1,X3) & coll(X7,X0,X2) & para(X0,X1,X2,X3)) => eqratio(X7,X0,X0,X2,X7,X1,X1,X3))),
  file('./7_Book/00EE/03/E042-2.p',ruleD65)).
fof(f14189,plain,(
  ( ! [X0] : (cong(sK22,X0,X0,sK20)) ) | (~spl28_18 | ~spl28_136)),
  inference(resolution,[],[f13956,f381])).
fof(f381,plain,(
  ( ! [X2,X0,X3,X1] : (~cong(X0,X1,X2,X3) | cong(X0,X1,X3,X2)) )),
  inference(cnf_transformation,[],[f198])).
fof(f198,plain,(
  ! [X0,X1,X2,X3] : (cong(X0,X1,X3,X2) | ~cong(X0,X1,X2,X3))),
  inference(ennf_transformation,[],[f23])).
fof(f23,axiom,(
  ! [X0,X1,X2,X3] : (cong(X0,X1,X2,X3) => cong(X0,X1,X3,X2))),
  file('./7_Book/00EE/03/E042-2.p',ruleD23)).
fof(f13956,plain,(
  ( ! [X6] : (cong(sK22,X6,sK20,X6)) ) | (~spl28_18 | ~spl28_136)),
  inference(subsumption_resolution,[],[f13949,f9401])).
fof(f9401,plain,(
  ( ! [X86] : (midp(X86,sK22,sK24)) ) | ~spl28_136),
  inference(subsumption_resolution,[],[f9379,f9319])).
fof(f9379,plain,(
  ( ! [X87,X86] : (midp(X86,sK22,sK24) | ~midp(X86,X87,X87)) ) | ~spl28_136),
  inference(resolution,[],[f9319,f1292])).
fof(f1292,plain,(
  ( ! [X2,X3,X1] : (~midp(sK27,X1,X2) | midp(X3,sK22,sK24) | ~midp(X3,X1,X2)) )),
  inference(subsumption_resolution,[],[f1286,f553])).
fof(f553,plain,(
  ( ! [X10,X11] : (para(X10,sK22,X11,sK24) | ~midp(sK27,X10,X11)) )),
  inference(resolution,[],[f422,f482])).
fof(f482,plain,(
  midp(sK27,sK22,sK24)),
  inference(resolution,[],[f369,f478])).
fof(f478,plain,(
  midp(sK27,sK24,sK22)),
  inference(cnf_transformation,[],[f358])).
fof(f1286,plain,(
  ( ! [X2,X3,X1] : (~midp(sK27,X1,X2) | midp(X3,sK22,sK24) | ~para(X1,sK22,X2,sK24) | ~midp(X3,X1,X2)) )),
  inference(resolution,[],[f552,f423])).
fof(f552,plain,(
  ( ! [X8,X9] : (para(X8,sK24,X9,sK22) | ~midp(sK27,X8,X9)) )),
  inference(resolution,[],[f422,f478])).
fof(f13949,plain,(
  ( ! [X6] : (~midp(X6,sK22,sK24) | cong(sK22,X6,sK20,X6)) ) | ~spl28_18),
  inference(resolution,[],[f925,f411])).
fof(f925,plain,(
  perp(sK22,sK20,sK20,sK24) | ~spl28_18),
  inference(avatar_component_clause,[],[f924])).
fof(f924,plain,(
  spl28_18 <=> perp(sK22,sK20,sK20,sK24)),
  introduced(avatar_definition,[new_symbols(naming,[spl28_18])])).
fof(f19192,plain,(
  ( ! [X24,X25] : (circle(sK22,sK22,X24,X25) | ~cong(sK22,sK22,sK22,X24)) ) | (~spl28_16 | ~spl28_18 | ~spl28_136 | ~spl28_520)),
  inference(resolution,[],[f16055,f370])).
fof(f370,plain,(
  ( ! [X2,X0,X3,X1] : (~cong(X3,X0,X3,X2) | circle(X3,X0,X1,X2) | ~cong(X3,X0,X3,X1)) )),
  inference(cnf_transformation,[],[f184])).
fof(f184,plain,(
  ! [X0,X1,X2,X3] : (circle(X3,X0,X1,X2) | ~cong(X3,X0,X3,X2) | ~cong(X3,X0,X3,X1))),
  inference(flattening,[],[f183])).
fof(f183,plain,(
  ! [X0,X1,X2,X3] : (circle(X3,X0,X1,X2) | (~cong(X3,X0,X3,X2) | ~cong(X3,X0,X3,X1)))),
  inference(ennf_transformation,[],[f98])).
fof(f98,plain,(
  ! [X0,X1,X2,X3] : ((cong(X3,X0,X3,X2) & cong(X3,X0,X3,X1)) => circle(X3,X0,X1,X2))),
  inference(rectify,[],[f12])).
fof(f12,axiom,(
  ! [X0,X1,X2,X7] : ((cong(X7,X0,X7,X2) & cong(X7,X0,X7,X1)) => circle(X7,X0,X1,X2))),
  file('./7_Book/00EE/03/E042-2.p',ruleD12)).
fof(f23437,plain,(
  ( ! [X6,X4,X7,X5] : (midp(X4,X5,X6) | ~midp(X7,X7,X5)) ) | (~spl28_16 | ~spl28_18 | ~spl28_136 | ~spl28_520)),
  inference(subsumption_resolution,[],[f23436,f21777])).
fof(f23436,plain,(
  ( ! [X6,X4,X7,X5] : (midp(X4,X5,X6) | ~midp(X7,X7,X6) | ~midp(X7,X7,X5)) ) | (~spl28_16 | ~spl28_136 | ~spl28_520)),
  inference(subsumption_resolution,[],[f23435,f15970])).
fof(f23435,plain,(
  ( ! [X6,X4,X7,X5] : (midp(X4,X5,X6) | ~para(X7,X6,X7,X4) | ~midp(X7,X7,X6) | ~midp(X7,X7,X5)) ) | ~spl28_136),
  inference(subsumption_resolution,[],[f23420,f6712])).
fof(f23420,plain,(
  ( ! [X6,X4,X7,X5] : (midp(X4,X5,X6) | ~coll(X4,X5,X6) | ~para(X7,X6,X7,X4) | ~midp(X7,X7,X6) | ~midp(X7,X7,X5)) )),
  inference(resolution,[],[f2051,f1945])).
fof(f1945,plain,(
  ( ! [X6,X4,X7,X5] : (circle(X4,X5,X6,X7) | ~midp(X4,X5,X7) | ~midp(X4,X5,X6)) )),
  inference(resolution,[],[f611,f427])).
fof(f427,plain,(
  ( ! [X2,X0,X1] : (cong(X0,X1,X0,X2) | ~midp(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f270])).
fof(f270,plain,(
  ! [X0,X1,X2] : (cong(X0,X1,X0,X2) | ~midp(X0,X1,X2))),
  inference(ennf_transformation,[],[f69])).
fof(f69,axiom,(
  ! [X0,X1,X2] : (midp(X0,X1,X2) => cong(X0,X1,X0,X2))),
  file('./7_Book/00EE/03/E042-2.p',ruleD68)).
fof(f611,plain,(
  ( ! [X2,X0,X3,X1] : (~cong(X0,X1,X0,X2) | circle(X0,X1,X2,X3) | ~midp(X0,X1,X3)) )),
  inference(resolution,[],[f370,f427])).
fof(f2051,plain,(
  ( ! [X78,X76,X77,X75] : (~circle(X75,X75,X78,X76) | midp(X77,X78,X76) | ~coll(X77,X78,X76) | ~para(X75,X76,X75,X77)) )),
  inference(resolution,[],[f665,f410])).
fof(f665,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (eqangle(X0,X1,X2,X3,X0,X1,X4,X5) | ~para(X2,X3,X4,X5)) )),
  inference(resolution,[],[f377,f398])).
fof(f377,plain,(
  ( ! [X6,X4,X2,X0,X7,X5,X3,X1] : (~eqangle(X0,X1,X2,X3,X4,X5,X6,X7) | eqangle(X2,X3,X0,X1,X6,X7,X4,X5)) )),
  inference(cnf_transformation,[],[f193])).
fof(f193,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6,X7] : (eqangle(X2,X3,X0,X1,X6,X7,X4,X5) | ~eqangle(X0,X1,X2,X3,X4,X5,X6,X7))),
  inference(ennf_transformation,[],[f101])).
fof(f101,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6,X7] : (eqangle(X0,X1,X2,X3,X4,X5,X6,X7) => eqangle(X2,X3,X0,X1,X6,X7,X4,X5))),
  inference(rectify,[],[f19])).
fof(f19,axiom,(
  ! [X0,X1,X2,X3,X8,X9,X10,X11] : (eqangle(X0,X1,X2,X3,X8,X9,X10,X11) => eqangle(X2,X3,X0,X1,X10,X11,X8,X9))),
  file('./7_Book/00EE/03/E042-2.p',ruleD19)).
fof(f13851,plain,(
  spl28_16 | ~spl28_12),
  inference(avatar_split_clause,[],[f13850,f904,f914])).
fof(f904,plain,(
  spl28_12 <=> perp(sK21,sK20,sK20,sK23)),
  introduced(avatar_definition,[new_symbols(naming,[spl28_12])])).
fof(f13850,plain,(
  perp(sK20,sK23,sK23,sK21) | ~spl28_12),
  inference(subsumption_resolution,[],[f1486,f905])).
fof(f905,plain,(
  perp(sK21,sK20,sK20,sK23) | ~spl28_12),
  inference(avatar_component_clause,[],[f904])).
fof(f1486,plain,(
  ~perp(sK21,sK20,sK20,sK23) | perp(sK20,sK23,sK23,sK21)),
  inference(resolution,[],[f673,f433])).
fof(f433,plain,(
  ( ! [X6,X4,X2,X0,X7,X5,X3,X1] : (~eqangle(X0,X1,X2,X3,X4,X5,X6,X7) | ~perp(X4,X5,X6,X7) | perp(X0,X1,X2,X3)) )),
  inference(cnf_transformation,[],[f281])).
fof(f281,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6,X7] : (perp(X0,X1,X2,X3) | ~perp(X4,X5,X6,X7) | ~eqangle(X0,X1,X2,X3,X4,X5,X6,X7))),
  inference(flattening,[],[f280])).
fof(f280,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6,X7] : (perp(X0,X1,X2,X3) | (~perp(X4,X5,X6,X7) | ~eqangle(X0,X1,X2,X3,X4,X5,X6,X7)))),
  inference(ennf_transformation,[],[f147])).
fof(f147,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6,X7] : ((perp(X4,X5,X6,X7) & eqangle(X0,X1,X2,X3,X4,X5,X6,X7)) => perp(X0,X1,X2,X3))),
  inference(rectify,[],[f75])).
fof(f75,axiom,(
  ! [X0,X1,X2,X3,X8,X9,X10,X11] : ((perp(X8,X9,X10,X11) & eqangle(X0,X1,X2,X3,X8,X9,X10,X11)) => perp(X0,X1,X2,X3))),
  file('./7_Book/00EE/03/E042-2.p',ruleD74)).
fof(f673,plain,(
  eqangle(sK20,sK23,sK23,sK21,sK21,sK20,sK20,sK23)),
  inference(resolution,[],[f378,f473])).
fof(f473,plain,(
  eqangle(sK21,sK20,sK20,sK23,sK20,sK23,sK23,sK21)),
  inference(cnf_transformation,[],[f358])).
fof(f378,plain,(
  ( ! [X6,X4,X2,X0,X7,X5,X3,X1] : (~eqangle(X0,X1,X2,X3,X4,X5,X6,X7) | eqangle(X4,X5,X6,X7,X0,X1,X2,X3)) )),
  inference(cnf_transformation,[],[f194])).
fof(f194,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6,X7] : (eqangle(X4,X5,X6,X7,X0,X1,X2,X3) | ~eqangle(X0,X1,X2,X3,X4,X5,X6,X7))),
  inference(ennf_transformation,[],[f102])).
fof(f102,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6,X7] : (eqangle(X0,X1,X2,X3,X4,X5,X6,X7) => eqangle(X4,X5,X6,X7,X0,X1,X2,X3))),
  inference(rectify,[],[f20])).
fof(f20,axiom,(
  ! [X0,X1,X2,X3,X8,X9,X10,X11] : (eqangle(X0,X1,X2,X3,X8,X9,X10,X11) => eqangle(X8,X9,X10,X11,X0,X1,X2,X3))),
  file('./7_Book/00EE/03/E042-2.p',ruleD20)).
fof(f13774,plain,(
  ~spl28_67 | spl28_18 | ~spl28_156),
  inference(avatar_split_clause,[],[f10017,f2680,f924,f1588])).
fof(f1588,plain,(
  spl28_67 <=> ~perp(sK20,sK24,sK22,sK24)),
  introduced(avatar_definition,[new_symbols(naming,[spl28_67])])).
fof(f2680,plain,(
  spl28_156 <=> cong(sK22,sK20,sK22,sK24)),
  introduced(avatar_definition,[new_symbols(naming,[spl28_156])])).
fof(f10017,plain,(
  perp(sK22,sK20,sK20,sK24) | ~perp(sK20,sK24,sK22,sK24) | ~spl28_156),
  inference(resolution,[],[f2681,f898])).
fof(f898,plain,(
  ( ! [X21,X19,X20] : (~cong(X21,X19,X21,X20) | perp(X21,X19,X19,X20) | ~perp(X19,X20,X21,X20)) )),
  inference(resolution,[],[f433,f405])).
fof(f405,plain,(
  ( ! [X2,X0,X1] : (eqangle(X2,X0,X0,X1,X0,X1,X2,X1) | ~cong(X2,X0,X2,X1)) )),
  inference(cnf_transformation,[],[f231])).
fof(f231,plain,(
  ! [X0,X1,X2] : (eqangle(X2,X0,X0,X1,X0,X1,X2,X1) | ~cong(X2,X0,X2,X1))),
  inference(ennf_transformation,[],[f126])).
fof(f126,plain,(
  ! [X0,X1,X2] : (cong(X2,X0,X2,X1) => eqangle(X2,X0,X0,X1,X0,X1,X2,X1))),
  inference(rectify,[],[f47])).
fof(f47,axiom,(
  ! [X0,X1,X7] : (cong(X7,X0,X7,X1) => eqangle(X7,X0,X0,X1,X0,X1,X7,X1))),
  file('./7_Book/00EE/03/E042-2.p',ruleD46)).
fof(f2681,plain,(
  cong(sK22,sK20,sK22,sK24) | ~spl28_156),
  inference(avatar_component_clause,[],[f2680])).
fof(f13706,plain,(
  ~spl28_59 | spl28_12 | ~spl28_152),
  inference(avatar_split_clause,[],[f9801,f2643,f904,f1552])).
fof(f1552,plain,(
  spl28_59 <=> ~perp(sK20,sK23,sK21,sK23)),
  introduced(avatar_definition,[new_symbols(naming,[spl28_59])])).
fof(f2643,plain,(
  spl28_152 <=> cong(sK21,sK20,sK21,sK23)),
  introduced(avatar_definition,[new_symbols(naming,[spl28_152])])).
fof(f9801,plain,(
  perp(sK21,sK20,sK20,sK23) | ~perp(sK20,sK23,sK21,sK23) | ~spl28_152),
  inference(resolution,[],[f2644,f898])).
fof(f2644,plain,(
  cong(sK21,sK20,sK21,sK23) | ~spl28_152),
  inference(avatar_component_clause,[],[f2643])).
fof(f13705,plain,(
  spl28_148 | ~spl28_24),
  inference(avatar_split_clause,[],[f5421,f946,f2579])).
fof(f2579,plain,(
  spl28_148 <=> ! [X13,X15,X14] : (~midp(X13,X14,X15) | cong(X14,X15,X13,X14))),
  introduced(avatar_definition,[new_symbols(naming,[spl28_148])])).
fof(f5421,plain,(
  ( ! [X10,X11,X9] : (~midp(X9,X10,X11) | cong(X10,X11,X9,X10)) ) | ~spl28_24),
  inference(resolution,[],[f947,f382])).
fof(f382,plain,(
  ( ! [X2,X0,X3,X1] : (~cong(X0,X1,X2,X3) | cong(X2,X3,X0,X1)) )),
  inference(cnf_transformation,[],[f199])).
fof(f199,plain,(
  ! [X0,X1,X2,X3] : (cong(X2,X3,X0,X1) | ~cong(X0,X1,X2,X3))),
  inference(ennf_transformation,[],[f24])).
fof(f24,axiom,(
  ! [X0,X1,X2,X3] : (cong(X0,X1,X2,X3) => cong(X2,X3,X0,X1))),
  file('./7_Book/00EE/03/E042-2.p',ruleD24)).
fof(f13660,plain,(
  spl28_66 | ~spl28_24 | ~spl28_136 | ~spl28_148),
  inference(avatar_split_clause,[],[f13659,f2579,f2059,f946,f1585])).
fof(f1585,plain,(
  spl28_66 <=> perp(sK20,sK24,sK22,sK24)),
  introduced(avatar_definition,[new_symbols(naming,[spl28_66])])).
fof(f13659,plain,(
  perp(sK20,sK24,sK22,sK24) | (~spl28_24 | ~spl28_136 | ~spl28_148)),
  inference(subsumption_resolution,[],[f7705,f13579])).
fof(f13579,plain,(
  ( ! [X39,X40] : (perp(X39,X40,X40,X39)) ) | (~spl28_24 | ~spl28_136 | ~spl28_148)),
  inference(resolution,[],[f13559,f365])).
fof(f365,plain,(
  ( ! [X2,X0,X3,X1] : (~perp(X0,X1,X2,X3) | perp(X0,X1,X3,X2)) )),
  inference(cnf_transformation,[],[f176])).
fof(f176,plain,(
  ! [X0,X1,X2,X3] : (perp(X0,X1,X3,X2) | ~perp(X0,X1,X2,X3))),
  inference(ennf_transformation,[],[f7])).
fof(f7,axiom,(
  ! [X0,X1,X2,X3] : (perp(X0,X1,X2,X3) => perp(X0,X1,X3,X2))),
  file('./7_Book/00EE/03/E042-2.p',ruleD7)).
fof(f13559,plain,(
  ( ! [X14,X15] : (perp(X14,X15,X14,X15)) ) | (~spl28_24 | ~spl28_136 | ~spl28_148)),
  inference(subsumption_resolution,[],[f13558,f9319])).
fof(f13558,plain,(
  ( ! [X14,X15] : (perp(X14,X15,X14,X15) | ~midp(X15,X14,X14)) ) | (~spl28_24 | ~spl28_136 | ~spl28_148)),
  inference(subsumption_resolution,[],[f13526,f9319])).
fof(f13526,plain,(
  ( ! [X14,X15] : (perp(X14,X15,X14,X15) | ~midp(X14,X15,X15) | ~midp(X15,X14,X14)) ) | (~spl28_24 | ~spl28_148)),
  inference(resolution,[],[f5423,f2580])).
fof(f2580,plain,(
  ( ! [X14,X15,X13] : (cong(X14,X15,X13,X14) | ~midp(X13,X14,X15)) ) | ~spl28_148),
  inference(avatar_component_clause,[],[f2579])).
fof(f5423,plain,(
  ( ! [X17,X15,X16] : (~cong(X15,X17,X16,X17) | perp(X15,X16,X17,X16) | ~midp(X15,X16,X16)) ) | ~spl28_24),
  inference(resolution,[],[f947,f415])).
fof(f7705,plain,(
  ~perp(sK22,sK20,sK20,sK22) | perp(sK20,sK24,sK22,sK24)),
  inference(resolution,[],[f1399,f433])).
fof(f1399,plain,(
  eqangle(sK20,sK24,sK22,sK24,sK22,sK20,sK20,sK22)),
  inference(resolution,[],[f663,f379])).
fof(f663,plain,(
  eqangle(sK20,sK24,sK22,sK20,sK22,sK24,sK20,sK22)),
  inference(resolution,[],[f377,f474])).
fof(f474,plain,(
  eqangle(sK22,sK20,sK20,sK24,sK20,sK22,sK22,sK24)),
  inference(cnf_transformation,[],[f358])).
fof(f13644,plain,(
  spl28_58 | ~spl28_24 | ~spl28_136 | ~spl28_148),
  inference(avatar_split_clause,[],[f13643,f2579,f2059,f946,f1549])).
fof(f1549,plain,(
  spl28_58 <=> perp(sK20,sK23,sK21,sK23)),
  introduced(avatar_definition,[new_symbols(naming,[spl28_58])])).
fof(f13643,plain,(
  perp(sK20,sK23,sK21,sK23) | (~spl28_24 | ~spl28_136 | ~spl28_148)),
  inference(subsumption_resolution,[],[f7630,f13579])).
fof(f7630,plain,(
  ~perp(sK21,sK20,sK20,sK21) | perp(sK20,sK23,sK21,sK23)),
  inference(resolution,[],[f1339,f433])).
fof(f1339,plain,(
  eqangle(sK20,sK23,sK21,sK23,sK21,sK20,sK20,sK21)),
  inference(resolution,[],[f661,f379])).
fof(f661,plain,(
  eqangle(sK20,sK23,sK21,sK20,sK21,sK23,sK20,sK21)),
  inference(resolution,[],[f377,f472])).
fof(f472,plain,(
  eqangle(sK21,sK20,sK20,sK23,sK20,sK21,sK21,sK23)),
  inference(cnf_transformation,[],[f358])).
fof(f10715,plain,(
  spl28_518 | ~spl28_412),
  inference(avatar_split_clause,[],[f10706,f7156,f8657])).
fof(f8657,plain,(
  spl28_518 <=> para(sK23,sK20,sK21,sK21)),
  introduced(avatar_definition,[new_symbols(naming,[spl28_518])])).
fof(f7156,plain,(
  spl28_412 <=> para(sK20,sK23,sK21,sK21)),
  introduced(avatar_definition,[new_symbols(naming,[spl28_412])])).
fof(f10706,plain,(
  para(sK23,sK20,sK21,sK21) | ~spl28_412),
  inference(resolution,[],[f7157,f1961])).
fof(f1961,plain,(
  ( ! [X2,X0,X3,X1] : (~para(X0,X1,X2,X3) | para(X1,X0,X2,X3)) )),
  inference(resolution,[],[f653,f397])).
fof(f397,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (~eqangle(X0,X1,X4,X5,X2,X3,X4,X5) | para(X0,X1,X2,X3)) )),
  inference(cnf_transformation,[],[f218])).
fof(f218,plain,(
  ! [X0,X1,X2,X3,X4,X5] : (para(X0,X1,X2,X3) | ~eqangle(X0,X1,X4,X5,X2,X3,X4,X5))),
  inference(ennf_transformation,[],[f118])).
fof(f118,plain,(
  ! [X0,X1,X2,X3,X4,X5] : (eqangle(X0,X1,X4,X5,X2,X3,X4,X5) => para(X0,X1,X2,X3))),
  inference(rectify,[],[f39])).
fof(f39,axiom,(
  ! [X0,X1,X2,X3,X8,X9] : (eqangle(X0,X1,X8,X9,X2,X3,X8,X9) => para(X0,X1,X2,X3))),
  file('./7_Book/00EE/03/E042-2.p',ruleD39)).
fof(f653,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (eqangle(X0,X1,X2,X3,X4,X5,X2,X3) | ~para(X1,X0,X4,X5)) )),
  inference(resolution,[],[f376,f398])).
fof(f376,plain,(
  ( ! [X6,X4,X2,X0,X7,X5,X3,X1] : (~eqangle(X0,X1,X2,X3,X4,X5,X6,X7) | eqangle(X1,X0,X2,X3,X4,X5,X6,X7)) )),
  inference(cnf_transformation,[],[f192])).
fof(f192,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6,X7] : (eqangle(X1,X0,X2,X3,X4,X5,X6,X7) | ~eqangle(X0,X1,X2,X3,X4,X5,X6,X7))),
  inference(ennf_transformation,[],[f100])).
fof(f100,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6,X7] : (eqangle(X0,X1,X2,X3,X4,X5,X6,X7) => eqangle(X1,X0,X2,X3,X4,X5,X6,X7))),
  inference(rectify,[],[f18])).
fof(f18,axiom,(
  ! [X0,X1,X2,X3,X8,X9,X10,X11] : (eqangle(X0,X1,X2,X3,X8,X9,X10,X11) => eqangle(X1,X0,X2,X3,X8,X9,X10,X11))),
  file('./7_Book/00EE/03/E042-2.p',ruleD18)).
fof(f7157,plain,(
  para(sK20,sK23,sK21,sK21) | ~spl28_412),
  inference(avatar_component_clause,[],[f7156])).
fof(f10702,plain,(
  spl28_412 | ~spl28_426),
  inference(avatar_split_clause,[],[f10695,f7217,f7156])).
fof(f7217,plain,(
  spl28_426 <=> para(sK21,sK21,sK20,sK23)),
  introduced(avatar_definition,[new_symbols(naming,[spl28_426])])).
fof(f10695,plain,(
  para(sK20,sK23,sK21,sK21) | ~spl28_426),
  inference(resolution,[],[f7218,f363])).
fof(f363,plain,(
  ( ! [X2,X0,X3,X1] : (~para(X0,X1,X2,X3) | para(X2,X3,X0,X1)) )),
  inference(cnf_transformation,[],[f173])).
fof(f173,plain,(
  ! [X0,X1,X2,X3] : (para(X2,X3,X0,X1) | ~para(X0,X1,X2,X3))),
  inference(ennf_transformation,[],[f5])).
fof(f5,axiom,(
  ! [X0,X1,X2,X3] : (para(X0,X1,X2,X3) => para(X2,X3,X0,X1))),
  file('./7_Book/00EE/03/E042-2.p',ruleD5)).
fof(f7218,plain,(
  para(sK21,sK21,sK20,sK23) | ~spl28_426),
  inference(avatar_component_clause,[],[f7217])).
fof(f10629,plain,(
  ~spl28_0 | ~spl28_136 | spl28_427 | ~spl28_474),
  inference(avatar_contradiction_clause,[],[f10628])).
fof(f10628,plain,(
  $false | (~spl28_0 | ~spl28_136 | ~spl28_427 | ~spl28_474)),
  inference(subsumption_resolution,[],[f10627,f9740])).
fof(f9740,plain,(
  ~midp(sK21,sK20,sK21) | (~spl28_0 | ~spl28_427)),
  inference(subsumption_resolution,[],[f8301,f7221])).
fof(f7221,plain,(
  ~para(sK21,sK21,sK20,sK23) | ~spl28_427),
  inference(avatar_component_clause,[],[f7220])).
fof(f7220,plain,(
  spl28_427 <=> ~para(sK21,sK21,sK20,sK23)),
  introduced(avatar_definition,[new_symbols(naming,[spl28_427])])).
fof(f8301,plain,(
  para(sK21,sK21,sK20,sK23) | ~midp(sK21,sK20,sK21) | ~spl28_0),
  inference(resolution,[],[f7135,f1637])).
fof(f1637,plain,(
  ( ! [X0,X1] : (para(X0,X0,X0,X1) | ~midp(X0,X1,X0)) )),
  inference(resolution,[],[f584,f517])).
fof(f517,plain,(
  ( ! [X2,X0,X1] : (cong(X0,X1,X0,X2) | ~midp(X0,X2,X1)) )),
  inference(resolution,[],[f382,f427])).
fof(f584,plain,(
  ( ! [X0,X1] : (~cong(X0,X0,X0,X1) | para(X0,X0,X0,X1)) )),
  inference(resolution,[],[f405,f397])).
fof(f7135,plain,(
  ( ! [X2,X3] : (~para(X2,X3,sK21,sK20) | para(X2,X3,sK20,sK23)) ) | ~spl28_0),
  inference(resolution,[],[f852,f364])).
fof(f364,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (~para(X2,X3,X4,X5) | para(X0,X1,X4,X5) | ~para(X0,X1,X2,X3)) )),
  inference(cnf_transformation,[],[f175])).
fof(f175,plain,(
  ! [X0,X1,X2,X3,X4,X5] : (para(X0,X1,X4,X5) | ~para(X2,X3,X4,X5) | ~para(X0,X1,X2,X3))),
  inference(flattening,[],[f174])).
fof(f174,plain,(
  ! [X0,X1,X2,X3,X4,X5] : (para(X0,X1,X4,X5) | (~para(X2,X3,X4,X5) | ~para(X0,X1,X2,X3)))),
  inference(ennf_transformation,[],[f6])).
fof(f6,axiom,(
  ! [X0,X1,X2,X3,X4,X5] : ((para(X2,X3,X4,X5) & para(X0,X1,X2,X3)) => para(X0,X1,X4,X5))),
  file('./7_Book/00EE/03/E042-2.p',ruleD6)).
fof(f852,plain,(
  para(sK21,sK20,sK20,sK23) | ~spl28_0),
  inference(avatar_component_clause,[],[f851])).
fof(f851,plain,(
  spl28_0 <=> para(sK21,sK20,sK20,sK23)),
  introduced(avatar_definition,[new_symbols(naming,[spl28_0])])).
fof(f10627,plain,(
  midp(sK21,sK20,sK21) | (~spl28_136 | ~spl28_474)),
  inference(subsumption_resolution,[],[f10614,f6566])).
fof(f10614,plain,(
  ~coll(sK21,sK20,sK21) | midp(sK21,sK20,sK21) | ~spl28_474),
  inference(resolution,[],[f7634,f426])).
fof(f7634,plain,(
  cong(sK21,sK20,sK21,sK21) | ~spl28_474),
  inference(avatar_component_clause,[],[f7633])).
fof(f7633,plain,(
  spl28_474 <=> cong(sK21,sK20,sK21,sK21)),
  introduced(avatar_definition,[new_symbols(naming,[spl28_474])])).
fof(f9947,plain,(
  ~spl28_136 | ~spl28_144 | spl28_157),
  inference(avatar_contradiction_clause,[],[f9946])).
fof(f9946,plain,(
  $false | (~spl28_136 | ~spl28_144 | ~spl28_157)),
  inference(subsumption_resolution,[],[f8212,f9401])).
fof(f8212,plain,(
  ~midp(sK20,sK22,sK24) | (~spl28_144 | ~spl28_157)),
  inference(resolution,[],[f2560,f2684])).
fof(f2684,plain,(
  ~cong(sK22,sK20,sK22,sK24) | ~spl28_157),
  inference(avatar_component_clause,[],[f2683])).
fof(f2683,plain,(
  spl28_157 <=> ~cong(sK22,sK20,sK22,sK24)),
  introduced(avatar_definition,[new_symbols(naming,[spl28_157])])).
fof(f2560,plain,(
  ( ! [X14,X15,X13] : (cong(X14,X13,X14,X15) | ~midp(X13,X14,X15)) ) | ~spl28_144),
  inference(avatar_component_clause,[],[f2559])).
fof(f2559,plain,(
  spl28_144 <=> ! [X13,X15,X14] : (~midp(X13,X14,X15) | cong(X14,X13,X14,X15))),
  introduced(avatar_definition,[new_symbols(naming,[spl28_144])])).
fof(f9739,plain,(
  ~spl28_136 | ~spl28_144 | spl28_153 | ~spl28_372),
  inference(avatar_contradiction_clause,[],[f9738])).
fof(f9738,plain,(
  $false | (~spl28_136 | ~spl28_144 | ~spl28_153 | ~spl28_372)),
  inference(subsumption_resolution,[],[f9737,f8207])).
fof(f8207,plain,(
  ~midp(sK20,sK21,sK23) | (~spl28_144 | ~spl28_153)),
  inference(resolution,[],[f2560,f2647])).
fof(f2647,plain,(
  ~cong(sK21,sK20,sK21,sK23) | ~spl28_153),
  inference(avatar_component_clause,[],[f2646])).
fof(f2646,plain,(
  spl28_153 <=> ~cong(sK21,sK20,sK21,sK23)),
  introduced(avatar_definition,[new_symbols(naming,[spl28_153])])).
fof(f9737,plain,(
  midp(sK20,sK21,sK23) | (~spl28_136 | ~spl28_372)),
  inference(subsumption_resolution,[],[f9724,f6712])).
fof(f9724,plain,(
  ~coll(sK20,sK21,sK23) | midp(sK20,sK21,sK23) | ~spl28_372),
  inference(resolution,[],[f5269,f426])).
fof(f5269,plain,(
  cong(sK20,sK21,sK20,sK23) | ~spl28_372),
  inference(avatar_component_clause,[],[f5268])).
fof(f5268,plain,(
  spl28_372 <=> cong(sK20,sK21,sK20,sK23)),
  introduced(avatar_definition,[new_symbols(naming,[spl28_372])])).
fof(f9702,plain,(
  spl28_372 | ~spl28_376),
  inference(avatar_split_clause,[],[f9672,f5281,f5268])).
fof(f5281,plain,(
  spl28_376 <=> cong(sK20,sK23,sK20,sK21)),
  introduced(avatar_definition,[new_symbols(naming,[spl28_376])])).
fof(f9672,plain,(
  cong(sK20,sK21,sK20,sK23) | ~spl28_376),
  inference(resolution,[],[f5282,f382])).
fof(f5282,plain,(
  cong(sK20,sK23,sK20,sK21) | ~spl28_376),
  inference(avatar_component_clause,[],[f5281])).
fof(f9660,plain,(
  ~spl28_0 | ~spl28_4 | ~spl28_136 | ~spl28_144 | spl28_377),
  inference(avatar_contradiction_clause,[],[f9659])).
fof(f9659,plain,(
  $false | (~spl28_0 | ~spl28_4 | ~spl28_136 | ~spl28_144 | ~spl28_377)),
  inference(subsumption_resolution,[],[f9647,f8202])).
fof(f8202,plain,(
  ~midp(sK23,sK20,sK21) | (~spl28_144 | ~spl28_377)),
  inference(resolution,[],[f2560,f5285])).
fof(f5285,plain,(
  ~cong(sK20,sK23,sK20,sK21) | ~spl28_377),
  inference(avatar_component_clause,[],[f5284])).
fof(f5284,plain,(
  spl28_377 <=> ~cong(sK20,sK23,sK20,sK21)),
  introduced(avatar_definition,[new_symbols(naming,[spl28_377])])).
fof(f9647,plain,(
  midp(sK23,sK20,sK21) | (~spl28_0 | ~spl28_4 | ~spl28_136)),
  inference(resolution,[],[f9371,f7191])).
fof(f7191,plain,(
  ( ! [X1] : (~midp(sK20,X1,sK23) | midp(sK23,X1,sK21)) ) | (~spl28_4 | ~spl28_136)),
  inference(subsumption_resolution,[],[f7177,f6712])).
fof(f7177,plain,(
  ( ! [X1] : (~coll(sK23,X1,sK21) | midp(sK23,X1,sK21) | ~midp(sK20,X1,sK23)) ) | ~spl28_4),
  inference(resolution,[],[f862,f404])).
fof(f404,plain,(
  ( ! [X4,X2,X0,X3,X1] : (~para(X3,X4,X1,X2) | ~coll(X4,X0,X2) | midp(X4,X0,X2) | ~midp(X3,X0,X1)) )),
  inference(cnf_transformation,[],[f230])).
fof(f230,plain,(
  ! [X0,X1,X2,X3,X4] : (midp(X4,X0,X2) | ~coll(X4,X0,X2) | ~para(X3,X4,X1,X2) | ~midp(X3,X0,X1))),
  inference(flattening,[],[f229])).
fof(f229,plain,(
  ! [X0,X1,X2,X3,X4] : (midp(X4,X0,X2) | (~coll(X4,X0,X2) | ~para(X3,X4,X1,X2) | ~midp(X3,X0,X1)))),
  inference(ennf_transformation,[],[f125])).
fof(f125,plain,(
  ! [X0,X1,X2,X3,X4] : ((coll(X4,X0,X2) & para(X3,X4,X1,X2) & midp(X3,X0,X1)) => midp(X4,X0,X2))),
  inference(rectify,[],[f46])).
fof(f46,axiom,(
  ! [X0,X1,X2,X4,X5] : ((coll(X5,X0,X2) & para(X4,X5,X1,X2) & midp(X4,X0,X1)) => midp(X5,X0,X2))),
  file('./7_Book/00EE/03/E042-2.p',ruleD45)).
fof(f862,plain,(
  para(sK20,sK23,sK23,sK21) | ~spl28_4),
  inference(avatar_component_clause,[],[f861])).
fof(f861,plain,(
  spl28_4 <=> para(sK20,sK23,sK23,sK21)),
  introduced(avatar_definition,[new_symbols(naming,[spl28_4])])).
fof(f9371,plain,(
  midp(sK20,sK20,sK23) | (~spl28_0 | ~spl28_136)),
  inference(resolution,[],[f9319,f7148])).
fof(f7148,plain,(
  ( ! [X1] : (~midp(sK21,X1,sK20) | midp(sK20,X1,sK23)) ) | (~spl28_0 | ~spl28_136)),
  inference(subsumption_resolution,[],[f7134,f6712])).
fof(f7134,plain,(
  ( ! [X1] : (~coll(sK20,X1,sK23) | midp(sK20,X1,sK23) | ~midp(sK21,X1,sK20)) ) | ~spl28_0),
  inference(resolution,[],[f852,f404])).
fof(f9382,plain,(
  ~spl28_136 | ~spl28_144 | spl28_475),
  inference(avatar_contradiction_clause,[],[f9369])).
fof(f9369,plain,(
  $false | (~spl28_136 | ~spl28_144 | ~spl28_475)),
  inference(resolution,[],[f9319,f8208])).
fof(f8208,plain,(
  ~midp(sK20,sK21,sK21) | (~spl28_144 | ~spl28_475)),
  inference(resolution,[],[f2560,f7637])).
fof(f7637,plain,(
  ~cong(sK21,sK20,sK21,sK21) | ~spl28_475),
  inference(avatar_component_clause,[],[f7636])).
fof(f7636,plain,(
  spl28_475 <=> ~cong(sK21,sK20,sK21,sK21)),
  introduced(avatar_definition,[new_symbols(naming,[spl28_475])])).
fof(f8665,plain,(
  ~spl28_519 | spl28_520 | ~spl28_2),
  inference(avatar_split_clause,[],[f8651,f854,f8663,f8660])).
fof(f8660,plain,(
  spl28_519 <=> ~para(sK23,sK20,sK21,sK21)),
  introduced(avatar_definition,[new_symbols(naming,[spl28_519])])).
fof(f854,plain,(
  spl28_2 <=> para(sK20,sK21,sK21,sK23)),
  introduced(avatar_definition,[new_symbols(naming,[spl28_2])])).
fof(f8651,plain,(
  ( ! [X0] : (midp(X0,sK20,sK21) | ~para(sK23,sK20,sK21,sK21) | ~midp(X0,sK23,sK21)) ) | ~spl28_2),
  inference(resolution,[],[f7601,f423])).
fof(f7601,plain,(
  para(sK23,sK21,sK21,sK20) | ~spl28_2),
  inference(resolution,[],[f7262,f362])).
fof(f362,plain,(
  ( ! [X2,X0,X3,X1] : (~para(X0,X1,X2,X3) | para(X0,X1,X3,X2)) )),
  inference(cnf_transformation,[],[f172])).
fof(f172,plain,(
  ! [X0,X1,X2,X3] : (para(X0,X1,X3,X2) | ~para(X0,X1,X2,X3))),
  inference(ennf_transformation,[],[f4])).
fof(f4,axiom,(
  ! [X0,X1,X2,X3] : (para(X0,X1,X2,X3) => para(X0,X1,X3,X2))),
  file('./7_Book/00EE/03/E042-2.p',ruleD4)).
fof(f7262,plain,(
  para(sK23,sK21,sK20,sK21) | ~spl28_2),
  inference(resolution,[],[f7154,f363])).
fof(f7154,plain,(
  para(sK20,sK21,sK23,sK21) | ~spl28_2),
  inference(resolution,[],[f855,f362])).
fof(f855,plain,(
  para(sK20,sK21,sK21,sK23) | ~spl28_2),
  inference(avatar_component_clause,[],[f854])).
fof(f7131,plain,(
  spl28_4 | ~spl28_54),
  inference(avatar_split_clause,[],[f7117,f1536,f861])).
fof(f1536,plain,(
  spl28_54 <=> para(sK20,sK23,sK21,sK23)),
  introduced(avatar_definition,[new_symbols(naming,[spl28_54])])).
fof(f7117,plain,(
  para(sK20,sK23,sK23,sK21) | ~spl28_54),
  inference(resolution,[],[f1537,f362])).
fof(f1537,plain,(
  para(sK20,sK23,sK21,sK23) | ~spl28_54),
  inference(avatar_component_clause,[],[f1536])).
fof(f7123,plain,(
  spl28_2 | ~spl28_1),
  inference(avatar_split_clause,[],[f1463,f848,f854])).
fof(f848,plain,(
  spl28_1 <=> ~para(sK21,sK20,sK20,sK23)),
  introduced(avatar_definition,[new_symbols(naming,[spl28_1])])).
fof(f1463,plain,(
  ~para(sK21,sK20,sK20,sK23) | para(sK20,sK21,sK21,sK23)),
  inference(resolution,[],[f672,f432])).
fof(f432,plain,(
  ( ! [X6,X4,X2,X0,X7,X5,X3,X1] : (~eqangle(X0,X1,X2,X3,X4,X5,X6,X7) | ~para(X4,X5,X6,X7) | para(X0,X1,X2,X3)) )),
  inference(cnf_transformation,[],[f279])).
fof(f279,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6,X7] : (para(X0,X1,X2,X3) | ~para(X4,X5,X6,X7) | ~eqangle(X0,X1,X2,X3,X4,X5,X6,X7))),
  inference(flattening,[],[f278])).
fof(f278,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6,X7] : (para(X0,X1,X2,X3) | (~para(X4,X5,X6,X7) | ~eqangle(X0,X1,X2,X3,X4,X5,X6,X7)))),
  inference(ennf_transformation,[],[f146])).
fof(f146,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6,X7] : ((para(X4,X5,X6,X7) & eqangle(X0,X1,X2,X3,X4,X5,X6,X7)) => para(X0,X1,X2,X3))),
  inference(rectify,[],[f74])).
fof(f74,axiom,(
  ! [X0,X1,X2,X3,X8,X9,X10,X11] : ((para(X8,X9,X10,X11) & eqangle(X0,X1,X2,X3,X8,X9,X10,X11)) => para(X0,X1,X2,X3))),
  file('./7_Book/00EE/03/E042-2.p',ruleD73)).
fof(f672,plain,(
  eqangle(sK20,sK21,sK21,sK23,sK21,sK20,sK20,sK23)),
  inference(resolution,[],[f378,f472])).
fof(f7107,plain,(
  spl28_54),
  inference(avatar_split_clause,[],[f7096,f1536])).
fof(f7096,plain,(
  para(sK20,sK23,sK21,sK23)),
  inference(resolution,[],[f1294,f397])).
fof(f1294,plain,(
  eqangle(sK20,sK23,sK20,sK21,sK21,sK23,sK20,sK21)),
  inference(resolution,[],[f649,f377])).
fof(f649,plain,(
  eqangle(sK20,sK21,sK20,sK23,sK20,sK21,sK21,sK23)),
  inference(resolution,[],[f376,f472])).
fof(f5418,plain,(
  ~spl28_146),
  inference(avatar_contradiction_clause,[],[f5407])).
fof(f5407,plain,(
  $false | ~spl28_146),
  inference(resolution,[],[f2563,f476])).
fof(f2563,plain,(
  ( ! [X12,X10,X11] : (~midp(X10,X11,X12)) ) | ~spl28_146),
  inference(avatar_component_clause,[],[f2562])).
fof(f2562,plain,(
  spl28_146 <=> ! [X11,X10,X12] : ~midp(X10,X11,X12)),
  introduced(avatar_definition,[new_symbols(naming,[spl28_146])])).
fof(f5369,plain,(
  spl28_146 | spl28_146 | ~spl28_138),
  inference(avatar_split_clause,[],[f5368,f2062,f2562,f2562])).
fof(f2062,plain,(
  spl28_138 <=> ! [X1,X0] : ~para(X0,X1,X0,X1)),
  introduced(avatar_definition,[new_symbols(naming,[spl28_138])])).
fof(f5368,plain,(
  ( ! [X45,X43,X41,X44,X42,X40] : (~midp(X40,X41,X42) | ~midp(X44,X45,X43)) ) | ~spl28_138),
  inference(subsumption_resolution,[],[f5367,f2063])).
fof(f2063,plain,(
  ( ! [X0,X1] : (~para(X0,X1,X0,X1)) ) | ~spl28_138),
  inference(avatar_component_clause,[],[f2062])).
fof(f5367,plain,(
  ( ! [X45,X43,X41,X44,X42,X40] : (~midp(X44,X45,X43) | para(X42,X43,X42,X43) | ~midp(X40,X41,X42)) )),
  inference(duplicate_literal_removal,[],[f5356])).
fof(f5356,plain,(
  ( ! [X45,X43,X41,X44,X42,X40] : (~midp(X40,X41,X42) | para(X42,X43,X42,X43) | ~midp(X44,X45,X43) | ~midp(X40,X41,X42) | ~midp(X44,X45,X43)) )),
  inference(resolution,[],[f954,f955])).
fof(f955,plain,(
  ( ! [X26,X24,X23,X27,X25,X22] : (para(X27,X24,sK5(X26,X27,X23,X24,X25,X22),X25) | ~midp(X25,X26,X27) | ~midp(X22,X23,X24)) )),
  inference(resolution,[],[f446,f363])).
fof(f446,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (para(sK5(X0,X1,X2,X3,X4,X5),X4,X1,X3) | ~midp(X5,X2,X3) | ~midp(X4,X0,X1)) )),
  inference(cnf_transformation,[],[f332])).
fof(f332,plain,(
  ! [X0,X1,X2,X3,X4,X5] : ((para(sK5(X0,X1,X2,X3,X4,X5),X5,X0,X2) & para(sK5(X0,X1,X2,X3,X4,X5),X4,X1,X3) & midp(sK5(X0,X1,X2,X3,X4,X5),X0,X3)) | ~midp(X5,X2,X3) | ~midp(X4,X0,X1))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK5])],[f295,f331])).
fof(f331,plain,(
  ! [X5,X4,X3,X2,X1,X0] : (? [X6] : ((para(X6,X5,X0,X2) & para(X6,X4,X1,X3) & midp(X6,X0,X3)) | ~midp(X5,X2,X3) | ~midp(X4,X0,X1)) => ((para(sK5(X0,X1,X2,X3,X4,X5),X5,X0,X2) & para(sK5(X0,X1,X2,X3,X4,X5),X4,X1,X3) & midp(sK5(X0,X1,X2,X3,X4,X5),X0,X3)) | ~midp(X5,X2,X3) | ~midp(X4,X0,X1)))),
  introduced(choice_axiom,[])).
fof(f295,plain,(
  ! [X0,X1,X2,X3,X4,X5] : ? [X6] : ((para(X6,X5,X0,X2) & para(X6,X4,X1,X3) & midp(X6,X0,X3)) | ~midp(X5,X2,X3) | ~midp(X4,X0,X1))),
  inference(flattening,[],[f294])).
fof(f294,plain,(
  ! [X0,X1,X2,X3,X4,X5] : ? [X6] : ((para(X6,X5,X0,X2) & para(X6,X4,X1,X3) & midp(X6,X0,X3)) | (~midp(X5,X2,X3) | ~midp(X4,X0,X1)))),
  inference(ennf_transformation,[],[f154])).
fof(f154,plain,(
  ! [X0,X1,X2,X3,X4,X5] : ? [X6] : ((midp(X5,X2,X3) & midp(X4,X0,X1)) => (para(X6,X5,X0,X2) & para(X6,X4,X1,X3) & midp(X6,X0,X3)))),
  inference(rectify,[],[f82])).
fof(f82,axiom,(
  ! [X0,X1,X2,X3,X6,X16] : ? [X8] : ((midp(X16,X2,X3) & midp(X6,X0,X1)) => (para(X8,X16,X0,X2) & para(X8,X6,X1,X3) & midp(X8,X0,X3)))),
  file('./7_Book/00EE/03/E042-2.p',ruleX6)).
fof(f954,plain,(
  ( ! [X14,X21,X19,X17,X15,X20,X18,X16] : (~para(X20,X21,sK5(X18,X19,X15,X16,X17,X14),X17) | ~midp(X17,X18,X19) | para(X20,X21,X19,X16) | ~midp(X14,X15,X16)) )),
  inference(resolution,[],[f446,f364])).
fof(f3950,plain,(
  ~spl28_26 | ~spl28_136),
  inference(avatar_contradiction_clause,[],[f3949])).
fof(f3949,plain,(
  $false | (~spl28_26 | ~spl28_136)),
  inference(subsumption_resolution,[],[f3948,f2570])).
fof(f2570,plain,(
  ( ! [X2,X3] : (~midp(X2,X2,X3)) ) | ~spl28_26),
  inference(duplicate_literal_removal,[],[f2567])).
fof(f2567,plain,(
  ( ! [X2,X3] : (~midp(X2,X2,X3) | ~midp(X2,X2,X3)) ) | ~spl28_26),
  inference(resolution,[],[f950,f427])).
fof(f950,plain,(
  ( ! [X6,X8,X7] : (~cong(X6,X7,X7,X8) | ~midp(X6,X7,X8)) ) | ~spl28_26),
  inference(avatar_component_clause,[],[f949])).
fof(f949,plain,(
  spl28_26 <=> ! [X7,X8,X6] : (~cong(X6,X7,X7,X8) | ~midp(X6,X7,X8))),
  introduced(avatar_definition,[new_symbols(naming,[spl28_26])])).
fof(f3948,plain,(
  ( ! [X14] : (midp(X14,X14,X14)) ) | ~spl28_136),
  inference(subsumption_resolution,[],[f3926,f2254])).
fof(f2254,plain,(
  ( ! [X0,X1] : (coll(X0,X1,X1)) ) | ~spl28_136),
  inference(resolution,[],[f2060,f425])).
fof(f3926,plain,(
  ( ! [X14] : (~coll(X14,X14,X14) | midp(X14,X14,X14)) ) | ~spl28_136),
  inference(resolution,[],[f3916,f426])).
fof(f3916,plain,(
  ( ! [X0] : (cong(X0,X0,X0,X0)) ) | ~spl28_136),
  inference(subsumption_resolution,[],[f3915,f2223])).
fof(f2223,plain,(
  ( ! [X70,X71,X69] : (cyclic(X70,X70,X69,X71)) ) | ~spl28_136),
  inference(subsumption_resolution,[],[f2222,f2221])).
fof(f2221,plain,(
  ( ! [X68,X66,X67] : (~coll(X66,X68,X67) | cyclic(X67,X67,X66,X68)) ) | ~spl28_136),
  inference(subsumption_resolution,[],[f2202,f2060])).
fof(f2222,plain,(
  ( ! [X70,X71,X69] : (coll(X69,X71,X70) | cyclic(X70,X70,X69,X71)) ) | ~spl28_136),
  inference(subsumption_resolution,[],[f2203,f2060])).
fof(f3915,plain,(
  ( ! [X0,X1] : (cong(X0,X0,X0,X0) | ~cyclic(X0,X0,X1,X0)) ) | ~spl28_136),
  inference(duplicate_literal_removal,[],[f3910])).
fof(f3910,plain,(
  ( ! [X0,X1] : (cong(X0,X0,X0,X0) | ~cyclic(X0,X0,X1,X0) | ~cyclic(X0,X0,X1,X0)) ) | ~spl28_136),
  inference(resolution,[],[f1199,f2223])).
fof(f2565,plain,(
  spl28_144 | spl28_26),
  inference(avatar_split_clause,[],[f2552,f949,f2559])).
fof(f2552,plain,(
  ( ! [X14,X12,X10,X15,X13,X11] : (~midp(X10,X11,X12) | ~midp(X13,X14,X15) | ~cong(X10,X11,X11,X12) | cong(X14,X13,X14,X15)) )),
  inference(resolution,[],[f696,f434])).
fof(f696,plain,(
  ( ! [X6,X10,X8,X7,X11,X9] : (eqratio(X6,X7,X6,X8,X9,X10,X10,X11) | ~midp(X9,X10,X11) | ~midp(X7,X6,X8)) )),
  inference(resolution,[],[f384,f429])).
fof(f429,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (eqratio(X4,X0,X0,X1,X5,X2,X2,X3) | ~midp(X5,X2,X3) | ~midp(X4,X0,X1)) )),
  inference(cnf_transformation,[],[f273])).
fof(f273,plain,(
  ! [X0,X1,X2,X3,X4,X5] : (eqratio(X4,X0,X0,X1,X5,X2,X2,X3) | ~midp(X5,X2,X3) | ~midp(X4,X0,X1))),
  inference(flattening,[],[f272])).
fof(f272,plain,(
  ! [X0,X1,X2,X3,X4,X5] : (eqratio(X4,X0,X0,X1,X5,X2,X2,X3) | (~midp(X5,X2,X3) | ~midp(X4,X0,X1)))),
  inference(ennf_transformation,[],[f145])).
fof(f145,plain,(
  ! [X0,X1,X2,X3,X4,X5] : ((midp(X5,X2,X3) & midp(X4,X0,X1)) => eqratio(X4,X0,X0,X1,X5,X2,X2,X3))),
  inference(rectify,[],[f71])).
fof(f71,axiom,(
  ! [X0,X1,X2,X3,X6,X16] : ((midp(X16,X2,X3) & midp(X6,X0,X1)) => eqratio(X6,X0,X0,X1,X16,X2,X2,X3))),
  file('./7_Book/00EE/03/E042-2.p',ruleD70)).
fof(f384,plain,(
  ( ! [X6,X4,X2,X0,X7,X5,X3,X1] : (~eqratio(X0,X1,X2,X3,X4,X5,X6,X7) | eqratio(X1,X0,X2,X3,X4,X5,X6,X7)) )),
  inference(cnf_transformation,[],[f202])).
fof(f202,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6,X7] : (eqratio(X1,X0,X2,X3,X4,X5,X6,X7) | ~eqratio(X0,X1,X2,X3,X4,X5,X6,X7))),
  inference(ennf_transformation,[],[f105])).
fof(f105,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6,X7] : (eqratio(X0,X1,X2,X3,X4,X5,X6,X7) => eqratio(X1,X0,X2,X3,X4,X5,X6,X7))),
  inference(rectify,[],[f26])).
fof(f26,axiom,(
  ! [X0,X1,X2,X3,X8,X9,X10,X11] : (eqratio(X0,X1,X2,X3,X8,X9,X10,X11) => eqratio(X1,X0,X2,X3,X8,X9,X10,X11))),
  file('./7_Book/00EE/03/E042-2.p',ruleD26)).
fof(f2064,plain,(
  spl28_136 | spl28_138),
  inference(avatar_split_clause,[],[f2035,f2062,f2059])).
fof(f2035,plain,(
  ( ! [X2,X0,X3,X1] : (~para(X0,X1,X0,X1) | para(X2,X3,X2,X3)) )),
  inference(resolution,[],[f665,f397])).
fof(f951,plain,(
  spl28_24 | spl28_26),
  inference(avatar_split_clause,[],[f944,f949,f946])).
fof(f944,plain,(
  ( ! [X6,X10,X8,X7,X11,X9] : (~cong(X6,X7,X7,X8) | cong(X9,X10,X10,X11) | ~midp(X6,X7,X8) | ~midp(X9,X10,X11)) )),
  inference(resolution,[],[f434,f429])).
fof(f866,plain,(
  spl28_0 | ~spl28_5),
  inference(avatar_split_clause,[],[f838,f864,f851])).
fof(f864,plain,(
  spl28_5 <=> ~para(sK20,sK23,sK23,sK21)),
  introduced(avatar_definition,[new_symbols(naming,[spl28_5])])).
fof(f838,plain,(
  ~para(sK20,sK23,sK23,sK21) | para(sK21,sK20,sK20,sK23)),
  inference(resolution,[],[f432,f473])).
% SZS output end Proof for E042-2
% ------------------------------
% Version: Vampire 4.4.0 (commit 7916d27 on 2019-08-23 08:50:16 +0100)
% Termination reason: Refutation

% Memory used [KB]: 17654
% Time elapsed: 0.463 s
% ------------------------------
% ------------------------------
