% Refutation found. Thanks to Tanya!
% SZS status Theorem for 108
% SZS output start Proof for 108
fof(f549,plain,(
  $false),
  inference(resolution,[],[f548,f484])).
fof(f484,plain,(
  ~perp(sK23,sK22,sK21,sK20)),
  inference(cnf_transformation,[],[f357])).
fof(f357,plain,(
  ~perp(sK23,sK22,sK21,sK20) & ! [X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20] : (coll(X12,X7,X8) & coll(X12,X4,X9) & coll(X11,X5,X8) & coll(X11,X4,X7) & coll(X10,X4,X9) & coll(X10,X5,X8) & circle(X7,X5,X9,X20) & circle(X6,X4,X9,X19) & circle(X7,X5,X17,X18) & circle(X6,X4,X8,X16) & circle(X6,X4,X7,X15) & circle(X6,X4,X13,X14) & midp(X6,X5,X4))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK20,sK21,sK22,sK23])],[f319,f356])).
fof(f356,plain,(
  ? [X0,X1,X2,X3] : (~perp(X3,X2,X1,X0) & ! [X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20] : (coll(X12,X7,X8) & coll(X12,X4,X9) & coll(X11,X5,X8) & coll(X11,X4,X7) & coll(X10,X4,X9) & coll(X10,X5,X8) & circle(X7,X5,X9,X20) & circle(X6,X4,X9,X19) & circle(X7,X5,X17,X18) & circle(X6,X4,X8,X16) & circle(X6,X4,X7,X15) & circle(X6,X4,X13,X14) & midp(X6,X5,X4))) => (~perp(sK23,sK22,sK21,sK20) & ! [X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20] : (coll(X12,X7,X8) & coll(X12,X4,X9) & coll(X11,X5,X8) & coll(X11,X4,X7) & coll(X10,X4,X9) & coll(X10,X5,X8) & circle(X7,X5,X9,X20) & circle(X6,X4,X9,X19) & circle(X7,X5,X17,X18) & circle(X6,X4,X8,X16) & circle(X6,X4,X7,X15) & circle(X6,X4,X13,X14) & midp(X6,X5,X4)))),
  introduced(choice_axiom,[])).
fof(f319,plain,(
  ? [X0,X1,X2,X3] : (~perp(X3,X2,X1,X0) & ! [X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20] : (coll(X12,X7,X8) & coll(X12,X4,X9) & coll(X11,X5,X8) & coll(X11,X4,X7) & coll(X10,X4,X9) & coll(X10,X5,X8) & circle(X7,X5,X9,X20) & circle(X6,X4,X9,X19) & circle(X7,X5,X17,X18) & circle(X6,X4,X8,X16) & circle(X6,X4,X7,X15) & circle(X6,X4,X13,X14) & midp(X6,X5,X4)))),
  inference(ennf_transformation,[],[f167])).
fof(f167,plain,(
  ~! [X0,X1,X2,X3] : (! [X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20] : (coll(X12,X7,X8) & coll(X12,X4,X9) & coll(X11,X5,X8) & coll(X11,X4,X7) & coll(X10,X4,X9) & coll(X10,X5,X8) & circle(X7,X5,X9,X20) & circle(X6,X4,X9,X19) & circle(X7,X5,X17,X18) & circle(X6,X4,X8,X16) & circle(X6,X4,X7,X15) & circle(X6,X4,X13,X14) & midp(X6,X5,X4)) => perp(X3,X2,X1,X0))),
  inference(rectify,[],[f96])).
fof(f96,negated_conjecture,(
  ~! [X3,X0,X5,X4] : (! [X0,X1,X7,X2,X8,X3,X12,X4,X5,X18,X19,X20,X21,X22,X23,X24,X25] : (coll(X5,X2,X8) & coll(X5,X0,X3) & coll(X4,X1,X8) & coll(X4,X0,X2) & coll(X12,X0,X3) & coll(X12,X1,X8) & circle(X2,X1,X3,X25) & circle(X7,X0,X3,X24) & circle(X2,X1,X22,X23) & circle(X7,X0,X8,X21) & circle(X7,X0,X2,X20) & circle(X7,X0,X18,X19) & midp(X7,X1,X0)) => perp(X4,X5,X0,X3))),
  inference(negated_conjecture,[],[f95])).
fof(f95,conjecture,(
  ( ! [X4,X0,X5,X3] : (! [X0,X1,X7,X2,X8,X3,X12,X4,X5,X18,X19,X20,X21,X22,X23,X24,X25] : (coll(X5,X2,X8) & coll(X5,X0,X3) & coll(X4,X1,X8) & coll(X4,X0,X2) & coll(X12,X0,X3) & coll(X12,X1,X8) & circle(X2,X1,X3,X25) & circle(X7,X0,X3,X24) & circle(X2,X1,X22,X23) & circle(X7,X0,X8,X21) & circle(X7,X0,X2,X20) & circle(X7,X0,X18,X19) & midp(X7,X1,X0)) => perp(X4,X5,X0,X3)) )),
  file('./6_GDD_FULL/81-109/108.p',exemplo6GDDFULL81109108)).
fof(f548,plain,(
  ( ! [X2,X0,X3,X1] : (perp(X0,X1,X2,X3)) )),
  inference(subsumption_resolution,[],[f543,f502])).
fof(f502,plain,(
  ( ! [X2,X0,X3,X1] : (para(X0,X1,X2,X3)) )),
  inference(subsumption_resolution,[],[f501,f471])).
fof(f471,plain,(
  ( ! [X6,X4,X5] : (midp(X6,X5,X4)) )),
  inference(cnf_transformation,[],[f357])).
fof(f501,plain,(
  ( ! [X4,X2,X0,X3,X1] : (para(X0,X1,X2,X3) | ~midp(X0,X4,X2)) )),
  inference(resolution,[],[f402,f471])).
fof(f402,plain,(
  ( ! [X4,X2,X0,X3,X1] : (~midp(X4,X0,X2) | para(X3,X4,X1,X2) | ~midp(X3,X0,X1)) )),
  inference(cnf_transformation,[],[f228])).
fof(f228,plain,(
  ! [X0,X1,X2,X3,X4] : (para(X3,X4,X1,X2) | ~midp(X4,X0,X2) | ~midp(X3,X0,X1))),
  inference(flattening,[],[f227])).
fof(f227,plain,(
  ! [X0,X1,X2,X3,X4] : (para(X3,X4,X1,X2) | (~midp(X4,X0,X2) | ~midp(X3,X0,X1)))),
  inference(ennf_transformation,[],[f124])).
fof(f124,plain,(
  ! [X0,X1,X2,X3,X4] : ((midp(X4,X0,X2) & midp(X3,X0,X1)) => para(X3,X4,X1,X2))),
  inference(rectify,[],[f45])).
fof(f45,axiom,(
  ! [X0,X1,X2,X4,X5] : ((midp(X5,X0,X2) & midp(X4,X0,X1)) => para(X4,X5,X1,X2))),
  file('./6_GDD_FULL/81-109/108.p',ruleD44)).
fof(f543,plain,(
  ( ! [X4,X2,X0,X3,X1] : (perp(X0,X1,X2,X3) | ~para(X0,X1,X4,X2)) )),
  inference(resolution,[],[f367,f516])).
fof(f516,plain,(
  ( ! [X2,X3,X1] : (perp(X1,X3,X3,X2)) )),
  inference(subsumption_resolution,[],[f515,f483])).
fof(f483,plain,(
  ( ! [X12,X8,X7] : (coll(X12,X7,X8)) )),
  inference(cnf_transformation,[],[f357])).
fof(f515,plain,(
  ( ! [X2,X0,X3,X1] : (~coll(X0,X1,X2) | perp(X1,X3,X3,X2)) )),
  inference(resolution,[],[f411,f477])).
fof(f477,plain,(
  ( ! [X7,X5,X9,X20] : (circle(X7,X5,X9,X20)) )),
  inference(cnf_transformation,[],[f357])).
fof(f411,plain,(
  ( ! [X2,X0,X3,X1] : (~circle(X3,X0,X1,X2) | ~coll(X3,X0,X2) | perp(X0,X1,X1,X2)) )),
  inference(cnf_transformation,[],[f245])).
fof(f245,plain,(
  ! [X0,X1,X2,X3] : (perp(X0,X1,X1,X2) | ~coll(X3,X0,X2) | ~circle(X3,X0,X1,X2))),
  inference(flattening,[],[f244])).
fof(f244,plain,(
  ! [X0,X1,X2,X3] : (perp(X0,X1,X1,X2) | (~coll(X3,X0,X2) | ~circle(X3,X0,X1,X2)))),
  inference(ennf_transformation,[],[f133])).
fof(f133,plain,(
  ! [X0,X1,X2,X3] : ((coll(X3,X0,X2) & circle(X3,X0,X1,X2)) => perp(X0,X1,X1,X2))),
  inference(rectify,[],[f54])).
fof(f54,axiom,(
  ! [X0,X1,X2,X7] : ((coll(X7,X0,X2) & circle(X7,X0,X1,X2)) => perp(X0,X1,X1,X2))),
  file('./6_GDD_FULL/81-109/108.p',ruleD53)).
fof(f367,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (~perp(X2,X3,X4,X5) | perp(X0,X1,X4,X5) | ~para(X0,X1,X2,X3)) )),
  inference(cnf_transformation,[],[f181])).
fof(f181,plain,(
  ! [X0,X1,X2,X3,X4,X5] : (perp(X0,X1,X4,X5) | ~perp(X2,X3,X4,X5) | ~para(X0,X1,X2,X3))),
  inference(flattening,[],[f180])).
fof(f180,plain,(
  ! [X0,X1,X2,X3,X4,X5] : (perp(X0,X1,X4,X5) | (~perp(X2,X3,X4,X5) | ~para(X0,X1,X2,X3)))),
  inference(ennf_transformation,[],[f10])).
fof(f10,axiom,(
  ! [X0,X1,X2,X3,X4,X5] : ((perp(X2,X3,X4,X5) & para(X0,X1,X2,X3)) => perp(X0,X1,X4,X5))),
  file('./6_GDD_FULL/81-109/108.p',ruleD10)).
% SZS output end Proof for 108
% ------------------------------
% Version: Vampire 4.4.0 (commit 7916d27 on 2019-08-23 08:50:16 +0100)
% Termination reason: Refutation

% Memory used [KB]: 5117
% Time elapsed: 0.028 s
% ------------------------------
% ------------------------------
