% Refutation found. Thanks to Tanya!
% SZS status Theorem for 91
% SZS output start Proof for 91
9. ! [X0,X1,X2,X3,X4,X5] : ((perp(X2,X3,X4,X5) & perp(X0,X1,X2,X3)) => para(X0,X1,X4,X5)) [input]
13. ! [X0,X1,X2,X3,X7] : ((cong(X7,X0,X7,X3) & cong(X7,X0,X7,X2) & cong(X7,X0,X7,X1)) => cyclic(X0,X1,X2,X3)) [input]
18. ! [X0,X1,X2,X3,X8,X9,X10,X11] : (eqangle(X0,X1,X2,X3,X8,X9,X10,X11) => eqangle(X1,X0,X2,X3,X8,X9,X10,X11)) [input]
19. ! [X0,X1,X2,X3,X8,X9,X10,X11] : (eqangle(X0,X1,X2,X3,X8,X9,X10,X11) => eqangle(X2,X3,X0,X1,X10,X11,X8,X9)) [input]
20. ! [X0,X1,X2,X3,X8,X9,X10,X11] : (eqangle(X0,X1,X2,X3,X8,X9,X10,X11) => eqangle(X8,X9,X10,X11,X0,X1,X2,X3)) [input]
21. ! [X0,X1,X2,X3,X8,X9,X10,X11] : (eqangle(X0,X1,X2,X3,X8,X9,X10,X11) => eqangle(X0,X1,X8,X9,X2,X3,X10,X11)) [input]
23. ! [X0,X1,X2,X3] : (cong(X0,X1,X2,X3) => cong(X0,X1,X3,X2)) [input]
24. ! [X0,X1,X2,X3] : (cong(X0,X1,X2,X3) => cong(X2,X3,X0,X1)) [input]
25. ! [X0,X1,X2,X3,X4,X5] : ((cong(X2,X3,X4,X5) & cong(X0,X1,X2,X3)) => cong(X0,X1,X4,X5)) [input]
40. ! [X0,X1,X2,X3,X8,X9] : (para(X0,X1,X2,X3) => eqangle(X0,X1,X8,X9,X2,X3,X8,X9)) [input]
45. ! [X0,X1,X2,X4,X5] : ((coll(X5,X0,X2) & para(X4,X5,X1,X2) & midp(X4,X0,X1)) => midp(X5,X0,X2)) [input]
48. ! [X0,X1,X2,X7,X15] : ((perp(X7,X0,X0,X15) & circle(X7,X0,X1,X2)) => eqangle(X0,X15,X0,X1,X2,X0,X2,X1)) [input]
51. ! [X0,X1,X2,X7,X6] : ((eqangle(X0,X1,X0,X2,X7,X1,X7,X6) & coll(X6,X1,X2) & circle(X7,X0,X1,X2)) => midp(X6,X1,X2)) [input]
55. ! [X0,X1,X6,X7] : ((perp(X7,X6,X0,X1) & midp(X6,X0,X1)) => cong(X7,X0,X7,X1)) [input]
67. ! [X0,X1,X2] : ((coll(X0,X1,X2) & cong(X0,X1,X0,X2)) => midp(X0,X1,X2)) [input]
68. ! [X0,X1,X2] : (midp(X0,X1,X2) => cong(X0,X1,X0,X2)) [input]
70. ! [X0,X1,X2,X3,X6,X16] : ((midp(X16,X2,X3) & midp(X6,X0,X1)) => eqratio(X6,X0,X0,X1,X16,X2,X2,X3)) [input]
75. ! [X0,X1,X2,X3,X8,X9,X10,X11] : ((cong(X8,X9,X10,X11) & eqratio(X0,X1,X2,X3,X8,X9,X10,X11)) => cong(X0,X1,X2,X3)) [input]
92. ! [X0,X1,X7] : ? [X2] : ((perp(X0,X7,X7,X1) & cong(X7,X0,X7,X1)) => (cong(X7,X0,X7,X2) & coll(X0,X7,X2))) [input]
94. ! [X0,X1,X2,X7,X18,X5,X12,X8,X19,X20,X21,X22,X23] : (perp(X19,X20,X18,X21) & perp(X5,X12,X8,X21) & coll(X20,X0,X2) & perp(X20,X8,X0,X2) & coll(X19,X0,X1) & perp(X19,X8,X0,X1) & circle(X7,X0,X8,X23) & coll(X12,X0,X1) & perp(X12,X18,X0,X1) & coll(X5,X0,X2) & perp(X5,X18,X0,X2) & circle(X7,X0,X18,X22) & circle(X7,X0,X1,X2)) => cyclic(X0,X18,X8,X21) [input]
95. ~! [X21,X8,X18,X0] : (! [X0,X1,X2,X7,X18,X5,X12,X8,X19,X20,X21,X22,X23] : (perp(X19,X20,X18,X21) & perp(X5,X12,X8,X21) & coll(X20,X0,X2) & perp(X20,X8,X0,X2) & coll(X19,X0,X1) & perp(X19,X8,X0,X1) & circle(X7,X0,X8,X23) & coll(X12,X0,X1) & perp(X12,X18,X0,X1) & coll(X5,X0,X2) & perp(X5,X18,X0,X2) & circle(X7,X0,X18,X22) & circle(X7,X0,X1,X2)) => cyclic(X0,X18,X8,X21)) [negated conjecture 94]
98. ! [X0,X1,X2,X3,X4] : ((cong(X4,X0,X4,X3) & cong(X4,X0,X4,X2) & cong(X4,X0,X4,X1)) => cyclic(X0,X1,X2,X3)) [rectify 13]
99. ! [X0,X1,X2,X3,X4,X5,X6,X7] : (eqangle(X0,X1,X2,X3,X4,X5,X6,X7) => eqangle(X1,X0,X2,X3,X4,X5,X6,X7)) [rectify 18]
100. ! [X0,X1,X2,X3,X4,X5,X6,X7] : (eqangle(X0,X1,X2,X3,X4,X5,X6,X7) => eqangle(X2,X3,X0,X1,X6,X7,X4,X5)) [rectify 19]
101. ! [X0,X1,X2,X3,X4,X5,X6,X7] : (eqangle(X0,X1,X2,X3,X4,X5,X6,X7) => eqangle(X4,X5,X6,X7,X0,X1,X2,X3)) [rectify 20]
102. ! [X0,X1,X2,X3,X4,X5,X6,X7] : (eqangle(X0,X1,X2,X3,X4,X5,X6,X7) => eqangle(X0,X1,X4,X5,X2,X3,X6,X7)) [rectify 21]
118. ! [X0,X1,X2,X3,X4,X5] : (para(X0,X1,X2,X3) => eqangle(X0,X1,X4,X5,X2,X3,X4,X5)) [rectify 40]
123. ! [X0,X1,X2,X3,X4] : ((coll(X4,X0,X2) & para(X3,X4,X1,X2) & midp(X3,X0,X1)) => midp(X4,X0,X2)) [rectify 45]
126. ! [X0,X1,X2,X3,X4] : ((perp(X3,X0,X0,X4) & circle(X3,X0,X1,X2)) => eqangle(X0,X4,X0,X1,X2,X0,X2,X1)) [rectify 48]
129. ! [X0,X1,X2,X3,X4] : ((eqangle(X0,X1,X0,X2,X3,X1,X3,X4) & coll(X4,X1,X2) & circle(X3,X0,X1,X2)) => midp(X4,X1,X2)) [rectify 51]
132. ! [X0,X1,X2,X3] : ((perp(X3,X2,X0,X1) & midp(X2,X0,X1)) => cong(X3,X0,X3,X1)) [rectify 55]
143. ! [X0,X1,X2,X3,X4,X5] : ((midp(X5,X2,X3) & midp(X4,X0,X1)) => eqratio(X4,X0,X0,X1,X5,X2,X2,X3)) [rectify 70]
146. ! [X0,X1,X2,X3,X4,X5,X6,X7] : ((cong(X4,X5,X6,X7) & eqratio(X0,X1,X2,X3,X4,X5,X6,X7)) => cong(X0,X1,X2,X3)) [rectify 75]
163. ! [X0,X1,X2] : ? [X3] : ((perp(X0,X2,X2,X1) & cong(X2,X0,X2,X1)) => (cong(X2,X0,X2,X3) & coll(X0,X2,X3))) [rectify 92]
165. ~! [X0,X1,X2,X3] : (! [X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16] : (perp(X12,X13,X8,X14) & perp(X9,X10,X11,X14) & coll(X13,X4,X6) & perp(X13,X11,X4,X6) & coll(X12,X4,X5) & perp(X12,X11,X4,X5) & circle(X7,X4,X11,X16) & coll(X10,X4,X5) & perp(X10,X8,X4,X5) & coll(X9,X4,X6) & perp(X9,X8,X4,X6) & circle(X7,X4,X8,X15) & circle(X7,X4,X5,X6)) => cyclic(X3,X2,X1,X0)) [rectify 95]
176. ! [X0,X1,X2,X3,X4,X5] : (para(X0,X1,X4,X5) | (~perp(X2,X3,X4,X5) | ~perp(X0,X1,X2,X3))) [ennf transformation 9]
177. ! [X0,X1,X2,X3,X4,X5] : (para(X0,X1,X4,X5) | ~perp(X2,X3,X4,X5) | ~perp(X0,X1,X2,X3)) [flattening 176]
183. ! [X0,X1,X2,X3,X4] : (cyclic(X0,X1,X2,X3) | (~cong(X4,X0,X4,X3) | ~cong(X4,X0,X4,X2) | ~cong(X4,X0,X4,X1))) [ennf transformation 98]
184. ! [X0,X1,X2,X3,X4] : (cyclic(X0,X1,X2,X3) | ~cong(X4,X0,X4,X3) | ~cong(X4,X0,X4,X2) | ~cong(X4,X0,X4,X1)) [flattening 183]
190. ! [X0,X1,X2,X3,X4,X5,X6,X7] : (eqangle(X1,X0,X2,X3,X4,X5,X6,X7) | ~eqangle(X0,X1,X2,X3,X4,X5,X6,X7)) [ennf transformation 99]
191. ! [X0,X1,X2,X3,X4,X5,X6,X7] : (eqangle(X2,X3,X0,X1,X6,X7,X4,X5) | ~eqangle(X0,X1,X2,X3,X4,X5,X6,X7)) [ennf transformation 100]
192. ! [X0,X1,X2,X3,X4,X5,X6,X7] : (eqangle(X4,X5,X6,X7,X0,X1,X2,X3) | ~eqangle(X0,X1,X2,X3,X4,X5,X6,X7)) [ennf transformation 101]
193. ! [X0,X1,X2,X3,X4,X5,X6,X7] : (eqangle(X0,X1,X4,X5,X2,X3,X6,X7) | ~eqangle(X0,X1,X2,X3,X4,X5,X6,X7)) [ennf transformation 102]
196. ! [X0,X1,X2,X3] : (cong(X0,X1,X3,X2) | ~cong(X0,X1,X2,X3)) [ennf transformation 23]
197. ! [X0,X1,X2,X3] : (cong(X2,X3,X0,X1) | ~cong(X0,X1,X2,X3)) [ennf transformation 24]
198. ! [X0,X1,X2,X3,X4,X5] : (cong(X0,X1,X4,X5) | (~cong(X2,X3,X4,X5) | ~cong(X0,X1,X2,X3))) [ennf transformation 25]
199. ! [X0,X1,X2,X3,X4,X5] : (cong(X0,X1,X4,X5) | ~cong(X2,X3,X4,X5) | ~cong(X0,X1,X2,X3)) [flattening 198]
217. ! [X0,X1,X2,X3,X4,X5] : (eqangle(X0,X1,X4,X5,X2,X3,X4,X5) | ~para(X0,X1,X2,X3)) [ennf transformation 118]
225. ! [X0,X1,X2,X3,X4] : (midp(X4,X0,X2) | (~coll(X4,X0,X2) | ~para(X3,X4,X1,X2) | ~midp(X3,X0,X1))) [ennf transformation 123]
226. ! [X0,X1,X2,X3,X4] : (midp(X4,X0,X2) | ~coll(X4,X0,X2) | ~para(X3,X4,X1,X2) | ~midp(X3,X0,X1)) [flattening 225]
230. ! [X0,X1,X2,X3,X4] : (eqangle(X0,X4,X0,X1,X2,X0,X2,X1) | (~perp(X3,X0,X0,X4) | ~circle(X3,X0,X1,X2))) [ennf transformation 126]
231. ! [X0,X1,X2,X3,X4] : (eqangle(X0,X4,X0,X1,X2,X0,X2,X1) | ~perp(X3,X0,X0,X4) | ~circle(X3,X0,X1,X2)) [flattening 230]
236. ! [X0,X1,X2,X3,X4] : (midp(X4,X1,X2) | (~eqangle(X0,X1,X0,X2,X3,X1,X3,X4) | ~coll(X4,X1,X2) | ~circle(X3,X0,X1,X2))) [ennf transformation 129]
237. ! [X0,X1,X2,X3,X4] : (midp(X4,X1,X2) | ~eqangle(X0,X1,X0,X2,X3,X1,X3,X4) | ~coll(X4,X1,X2) | ~circle(X3,X0,X1,X2)) [flattening 236]
244. ! [X0,X1,X2,X3] : (cong(X3,X0,X3,X1) | (~perp(X3,X2,X0,X1) | ~midp(X2,X0,X1))) [ennf transformation 132]
245. ! [X0,X1,X2,X3] : (cong(X3,X0,X3,X1) | ~perp(X3,X2,X0,X1) | ~midp(X2,X0,X1)) [flattening 244]
264. ! [X0,X1,X2] : (midp(X0,X1,X2) | (~coll(X0,X1,X2) | ~cong(X0,X1,X0,X2))) [ennf transformation 67]
265. ! [X0,X1,X2] : (midp(X0,X1,X2) | ~coll(X0,X1,X2) | ~cong(X0,X1,X0,X2)) [flattening 264]
266. ! [X0,X1,X2] : (cong(X0,X1,X0,X2) | ~midp(X0,X1,X2)) [ennf transformation 68]
268. ! [X0,X1,X2,X3,X4,X5] : (eqratio(X4,X0,X0,X1,X5,X2,X2,X3) | (~midp(X5,X2,X3) | ~midp(X4,X0,X1))) [ennf transformation 143]
269. ! [X0,X1,X2,X3,X4,X5] : (eqratio(X4,X0,X0,X1,X5,X2,X2,X3) | ~midp(X5,X2,X3) | ~midp(X4,X0,X1)) [flattening 268]
278. ! [X0,X1,X2,X3,X4,X5,X6,X7] : (cong(X0,X1,X2,X3) | (~cong(X4,X5,X6,X7) | ~eqratio(X0,X1,X2,X3,X4,X5,X6,X7))) [ennf transformation 146]
279. ! [X0,X1,X2,X3,X4,X5,X6,X7] : (cong(X0,X1,X2,X3) | ~cong(X4,X5,X6,X7) | ~eqratio(X0,X1,X2,X3,X4,X5,X6,X7)) [flattening 278]
311. ! [X0,X1,X2] : ? [X3] : ((cong(X2,X0,X2,X3) & coll(X0,X2,X3)) | (~perp(X0,X2,X2,X1) | ~cong(X2,X0,X2,X1))) [ennf transformation 163]
312. ! [X0,X1,X2] : ? [X3] : ((cong(X2,X0,X2,X3) & coll(X0,X2,X3)) | ~perp(X0,X2,X2,X1) | ~cong(X2,X0,X2,X1)) [flattening 311]
315. ? [X0,X1,X2,X3] : (~cyclic(X3,X2,X1,X0) & ! [X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16] : (perp(X12,X13,X8,X14) & perp(X9,X10,X11,X14) & coll(X13,X4,X6) & perp(X13,X11,X4,X6) & coll(X12,X4,X5) & perp(X12,X11,X4,X5) & circle(X7,X4,X11,X16) & coll(X10,X4,X5) & perp(X10,X8,X4,X5) & coll(X9,X4,X6) & perp(X9,X8,X4,X6) & circle(X7,X4,X8,X15) & circle(X7,X4,X5,X6))) [ennf transformation 165]
348. ! [X2,X1,X0] : (? [X3] : ((cong(X2,X0,X2,X3) & coll(X0,X2,X3)) | ~perp(X0,X2,X2,X1) | ~cong(X2,X0,X2,X1)) => ((cong(X2,X0,X2,sK18(X0,X1,X2)) & coll(X0,X2,sK18(X0,X1,X2))) | ~perp(X0,X2,X2,X1) | ~cong(X2,X0,X2,X1))) [choice axiom]
349. ! [X0,X1,X2] : ((cong(X2,X0,X2,sK18(X0,X1,X2)) & coll(X0,X2,sK18(X0,X1,X2))) | ~perp(X0,X2,X2,X1) | ~cong(X2,X0,X2,X1)) [skolemisation 312,348]
352. ? [X0,X1,X2,X3] : (~cyclic(X3,X2,X1,X0) & ! [X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16] : (perp(X12,X13,X8,X14) & perp(X9,X10,X11,X14) & coll(X13,X4,X6) & perp(X13,X11,X4,X6) & coll(X12,X4,X5) & perp(X12,X11,X4,X5) & circle(X7,X4,X11,X16) & coll(X10,X4,X5) & perp(X10,X8,X4,X5) & coll(X9,X4,X6) & perp(X9,X8,X4,X6) & circle(X7,X4,X8,X15) & circle(X7,X4,X5,X6))) => (~cyclic(sK23,sK22,sK21,sK20) & ! [X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16] : (perp(X12,X13,X8,X14) & perp(X9,X10,X11,X14) & coll(X13,X4,X6) & perp(X13,X11,X4,X6) & coll(X12,X4,X5) & perp(X12,X11,X4,X5) & circle(X7,X4,X11,X16) & coll(X10,X4,X5) & perp(X10,X8,X4,X5) & coll(X9,X4,X6) & perp(X9,X8,X4,X6) & circle(X7,X4,X8,X15) & circle(X7,X4,X5,X6))) [choice axiom]
353. ~cyclic(sK23,sK22,sK21,sK20) & ! [X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16] : (perp(X12,X13,X8,X14) & perp(X9,X10,X11,X14) & coll(X13,X4,X6) & perp(X13,X11,X4,X6) & coll(X12,X4,X5) & perp(X12,X11,X4,X5) & circle(X7,X4,X11,X16) & coll(X10,X4,X5) & perp(X10,X8,X4,X5) & coll(X9,X4,X6) & perp(X9,X8,X4,X6) & circle(X7,X4,X8,X15) & circle(X7,X4,X5,X6)) [skolemisation 315,352]
362. ~perp(X2,X3,X4,X5) | para(X0,X1,X4,X5) | ~perp(X0,X1,X2,X3) [cnf transformation 177]
366. ~cong(X4,X0,X4,X3) | cyclic(X0,X1,X2,X3) | ~cong(X4,X0,X4,X2) | ~cong(X4,X0,X4,X1) [cnf transformation 184]
371. ~eqangle(X0,X1,X2,X3,X4,X5,X6,X7) | eqangle(X1,X0,X2,X3,X4,X5,X6,X7) [cnf transformation 190]
372. ~eqangle(X0,X1,X2,X3,X4,X5,X6,X7) | eqangle(X2,X3,X0,X1,X6,X7,X4,X5) [cnf transformation 191]
373. ~eqangle(X0,X1,X2,X3,X4,X5,X6,X7) | eqangle(X4,X5,X6,X7,X0,X1,X2,X3) [cnf transformation 192]
374. ~eqangle(X0,X1,X2,X3,X4,X5,X6,X7) | eqangle(X0,X1,X4,X5,X2,X3,X6,X7) [cnf transformation 193]
376. ~cong(X0,X1,X2,X3) | cong(X0,X1,X3,X2) [cnf transformation 196]
377. ~cong(X0,X1,X2,X3) | cong(X2,X3,X0,X1) [cnf transformation 197]
378. ~cong(X2,X3,X4,X5) | cong(X0,X1,X4,X5) | ~cong(X0,X1,X2,X3) [cnf transformation 199]
393. eqangle(X0,X1,X4,X5,X2,X3,X4,X5) | ~para(X0,X1,X2,X3) [cnf transformation 217]
399. ~para(X3,X4,X1,X2) | ~coll(X4,X0,X2) | midp(X4,X0,X2) | ~midp(X3,X0,X1) [cnf transformation 226]
402. ~circle(X3,X0,X1,X2) | ~perp(X3,X0,X0,X4) | eqangle(X0,X4,X0,X1,X2,X0,X2,X1) [cnf transformation 231]
405. ~eqangle(X0,X1,X0,X2,X3,X1,X3,X4) | midp(X4,X1,X2) | ~coll(X4,X1,X2) | ~circle(X3,X0,X1,X2) [cnf transformation 237]
409. ~perp(X3,X2,X0,X1) | cong(X3,X0,X3,X1) | ~midp(X2,X0,X1) [cnf transformation 245]
421. ~cong(X0,X1,X0,X2) | ~coll(X0,X1,X2) | midp(X0,X1,X2) [cnf transformation 265]
422. cong(X0,X1,X0,X2) | ~midp(X0,X1,X2) [cnf transformation 266]
424. eqratio(X4,X0,X0,X1,X5,X2,X2,X3) | ~midp(X5,X2,X3) | ~midp(X4,X0,X1) [cnf transformation 269]
429. ~eqratio(X0,X1,X2,X3,X4,X5,X6,X7) | ~cong(X4,X5,X6,X7) | cong(X0,X1,X2,X3) [cnf transformation 279]
464. cong(X2,X0,X2,sK18(X0,X1,X2)) | ~perp(X0,X2,X2,X1) | ~cong(X2,X0,X2,X1) [cnf transformation 349]
473. circle(X7,X4,X11,X16) [cnf transformation 353]
477. coll(X13,X4,X6) [cnf transformation 353]
479. perp(X12,X13,X8,X14) [cnf transformation 353]
480. ~cyclic(sK23,sK22,sK21,sK20) [cnf transformation 353]
485. cong(X0,X1,X2,X0) | ~midp(X0,X1,X2) [resolution 376,422]
487. cong(X0,X1,X0,X2) | ~midp(X0,X2,X1) [resolution 377,422]
488. cong(X3,X4,X4,X5) | ~midp(X4,X5,X3) [resolution 377,485]
493. cong(X5,X3,X4,X3) | ~midp(X3,X4,X5) [resolution 488,376]
496. cong(X1,X0,X2,X0) | ~midp(X0,X1,X2) [resolution 493,377]
524. ~midp(X3,X1,X2) | cong(X0,X1,X0,X2) [resolution 409,479]
537. para(X0,X1,X2,X3) | ~perp(X0,X1,X4,X5) [resolution 362,479]
539. para(X0,X1,X2,X3) [subsumption resolution 537,479]
567. ~cong(X5,X6,X7,X9) | cong(X5,X6,X7,X8) | ~midp(X7,X9,X8) [resolution 378,422]
589. cong(X12,X13,X14,X15) | ~midp(X14,X12,X15) | ~midp(X12,X13,X14) [resolution 567,485]
599. ~coll(X12,X13,X14) | midp(X12,X13,X14) | ~midp(X15,X13,X16) [resolution 399,539]
606. ~midp(X15,X13,X16) | midp(X12,X13,X14) [subsumption resolution 599,477]
610. eqangle(X0,X1,X2,X3,X4,X5,X2,X3) | ~para(X1,X0,X4,X5) [resolution 371,393]
614. eqangle(X0,X1,X2,X3,X4,X5,X2,X3) [subsumption resolution 610,539]
617. eqangle(X0,X1,X2,X3,X0,X1,X4,X5) [resolution 372,614]
624. eqangle(X8,X9,X10,X11,X9,X8,X12,X13) [resolution 617,371]
638. eqangle(X18,X19,X19,X18,X20,X21,X22,X23) [resolution 374,624]
828. ~perp(X23,X24,X24,X25) | ~cong(X24,X23,X24,X25) | ~coll(X24,X23,sK18(X23,X25,X24)) | midp(X24,X23,sK18(X23,X25,X24)) [resolution 464,421]
834. ~cong(X24,X23,X24,X25) | ~coll(X24,X23,sK18(X23,X25,X24)) | midp(X24,X23,sK18(X23,X25,X24)) [subsumption resolution 828,479]
835. midp(X24,X23,sK18(X23,X25,X24)) | ~cong(X24,X23,X24,X25) [subsumption resolution 834,477]
845. ~cong(X35,X36,X35,X37) | midp(X38,X36,X39) [resolution 835,606]
848. ~perp(X0,X1,X1,X2) | eqangle(X1,X2,X1,X3,X4,X1,X4,X3) [resolution 402,473]
849. eqangle(X1,X2,X1,X3,X4,X1,X4,X3) [subsumption resolution 848,479]
852. eqangle(X6,X7,X6,X8,X7,X9,X7,X8) [resolution 849,373]
854. eqangle(X14,X15,X15,X16,X17,X15,X17,X16) [resolution 849,371]
902. eqangle(X10,X11,X10,X12,X12,X11,X12,X13) [resolution 852,372]
976. ~midp(X3,X4,X1) | midp(X0,X1,X2) [resolution 845,487]
981. ~midp(X24,X26,X26) | midp(X23,X24,X25) [resolution 845,496]
1014. ~cong(X6,X7,X7,X8) | cong(X9,X10,X10,X11) | ~midp(X6,X7,X8) | ~midp(X9,X10,X11) [resolution 429,424]
1016. 0 <=> ! [X9,X11,X10] : (cong(X9,X10,X10,X11) | ~midp(X9,X10,X11)) [avatar definition]
1017. cong(X9,X10,X10,X11) | ~midp(X9,X10,X11) <- (0) [avatar component clause 1016]
1019. 2 <=> ! [X7,X8,X6] : (~cong(X6,X7,X7,X8) | ~midp(X6,X7,X8)) [avatar definition]
1020. ~cong(X6,X7,X7,X8) | ~midp(X6,X7,X8) <- (2) [avatar component clause 1019]
1021. 0 | 2 [avatar split clause 1014,1019,1016]
1036. ~midp(X14,X15,X16) | ~midp(X15,X14,X16) | ~midp(X14,X15,X15) <- (2) [resolution 1020,589]
1042. ~midp(X14,X15,X16) | ~midp(X14,X15,X15) <- (2) [subsumption resolution 1036,981]
1043. ~midp(X14,X15,X15) <- (2) [subsumption resolution 1042,976]
1087. cyclic(X40,X41,X42,X43) | ~cong(X40,X40,X40,X42) | ~cong(X40,X40,X40,X41) | ~midp(X40,X43,X40) [resolution 366,488]
1209. midp(X70,X71,X72) | ~coll(X70,X71,X72) | ~circle(X72,X73,X71,X72) [resolution 405,902]
1214. midp(X86,X87,X86) | ~coll(X86,X87,X86) | ~circle(X88,X87,X87,X86) [resolution 405,854]
1216. midp(X91,X92,X92) | ~coll(X91,X92,X92) | ~circle(X93,X92,X92,X92) [resolution 405,638]
1297. ~coll(X91,X92,X92) | ~circle(X93,X92,X92,X92) <- (2) [subsumption resolution 1216,1043]
1298. ~circle(X93,X92,X92,X92) <- (2) [subsumption resolution 1297,477]
1299. $false <- (2) [subsumption resolution 1298,473]
1300. ~2 [avatar contradiction clause 1299]
1305. midp(X86,X87,X86) | ~circle(X88,X87,X87,X86) [subsumption resolution 1214,477]
1306. midp(X86,X87,X86) [subsumption resolution 1305,473]
1309. midp(X70,X71,X72) | ~circle(X72,X73,X71,X72) [subsumption resolution 1209,477]
1310. midp(X70,X71,X72) [subsumption resolution 1309,473]
1321. ~cong(X40,X40,X40,X42) | cyclic(X40,X41,X42,X43) | ~cong(X40,X40,X40,X41) [subsumption resolution 1087,1306]
1324. cong(X9,X10,X10,X11) <- (0) [subsumption resolution 1017,1310]
1330. cong(X15,X16,X15,X17) [resolution 1306,524]
5282. cyclic(X6,X7,X8,X9) | ~cong(X6,X6,X6,X7) [resolution 1321,1330]
5307. cyclic(X6,X7,X8,X9) <- (0) [subsumption resolution 5282,1324]
5308. $false <- (0) [resolution 5307,480]
5315. ~0 [avatar contradiction clause 5308]
5316. $false [avatar sat refutation 1021,1300,5315]
% SZS output end Proof for 91
% ------------------------------
% Version: Vampire 4.4.0 (commit 7916d27 on 2019-08-23 08:50:16 +0100)
% Termination reason: Refutation

% Memory used [KB]: 6268
% Time elapsed: 0.063 s
% ------------------------------
% ------------------------------
