% Refutation found. Thanks to Tanya!
% SZS status Theorem for 97
% SZS output start Proof for 97
fof(f38540,plain,(
  $false),
  inference(avatar_sat_refutation,[],[f623,f1114,f2081,f2343,f2732,f3206,f3580,f8881,f20792,f21097,f21099,f38530])).
fof(f38530,plain,(
  ~spl32_0 | ~spl32_162 | ~spl32_192 | ~spl32_236 | ~spl32_240 | ~spl32_372),
  inference(avatar_contradiction_clause,[],[f38525])).
fof(f38525,plain,(
  $false | (~spl32_0 | ~spl32_162 | ~spl32_192 | ~spl32_236 | ~spl32_240 | ~spl32_372)),
  inference(resolution,[],[f22061,f484])).
fof(f484,plain,(
  ~perp(sK28,sK29,sK25,sK20)),
  inference(cnf_transformation,[],[f358])).
fof(f358,plain,(
  ~perp(sK28,sK29,sK25,sK20) & coll(sK29,sK20,sK21) & perp(sK29,sK26,sK20,sK21) & coll(sK28,sK21,sK22) & perp(sK28,sK26,sK21,sK22) & coll(sK27,sK20,sK26) & coll(sK27,sK21,sK22) & circle(sK23,sK20,sK26,sK31) & eqangle(sK26,sK20,sK20,sK21,sK22,sK20,sK20,sK24) & coll(sK25,sK20,sK24) & circle(sK23,sK20,sK25,sK30) & midp(sK24,sK22,sK21) & circle(sK23,sK20,sK21,sK22)),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK20,sK21,sK22,sK23,sK24,sK25,sK26,sK27,sK28,sK29,sK30,sK31])],[f320,f357])).
fof(f357,plain,(
  ? [X0,X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11] : (~perp(X8,X9,X5,X0) & coll(X9,X0,X1) & perp(X9,X6,X0,X1) & coll(X8,X1,X2) & perp(X8,X6,X1,X2) & coll(X7,X0,X6) & coll(X7,X1,X2) & circle(X3,X0,X6,X11) & eqangle(X6,X0,X0,X1,X2,X0,X0,X4) & coll(X5,X0,X4) & circle(X3,X0,X5,X10) & midp(X4,X2,X1) & circle(X3,X0,X1,X2)) => (~perp(sK28,sK29,sK25,sK20) & coll(sK29,sK20,sK21) & perp(sK29,sK26,sK20,sK21) & coll(sK28,sK21,sK22) & perp(sK28,sK26,sK21,sK22) & coll(sK27,sK20,sK26) & coll(sK27,sK21,sK22) & circle(sK23,sK20,sK26,sK31) & eqangle(sK26,sK20,sK20,sK21,sK22,sK20,sK20,sK24) & coll(sK25,sK20,sK24) & circle(sK23,sK20,sK25,sK30) & midp(sK24,sK22,sK21) & circle(sK23,sK20,sK21,sK22))),
  introduced(choice_axiom,[])).
fof(f320,plain,(
  ? [X0,X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11] : (~perp(X8,X9,X5,X0) & coll(X9,X0,X1) & perp(X9,X6,X0,X1) & coll(X8,X1,X2) & perp(X8,X6,X1,X2) & coll(X7,X0,X6) & coll(X7,X1,X2) & circle(X3,X0,X6,X11) & eqangle(X6,X0,X0,X1,X2,X0,X0,X4) & coll(X5,X0,X4) & circle(X3,X0,X5,X10) & midp(X4,X2,X1) & circle(X3,X0,X1,X2))),
  inference(flattening,[],[f319])).
fof(f319,plain,(
  ? [X0,X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11] : (~perp(X8,X9,X5,X0) & (coll(X9,X0,X1) & perp(X9,X6,X0,X1) & coll(X8,X1,X2) & perp(X8,X6,X1,X2) & coll(X7,X0,X6) & coll(X7,X1,X2) & circle(X3,X0,X6,X11) & eqangle(X6,X0,X0,X1,X2,X0,X0,X4) & coll(X5,X0,X4) & circle(X3,X0,X5,X10) & midp(X4,X2,X1) & circle(X3,X0,X1,X2)))),
  inference(ennf_transformation,[],[f167])).
fof(f167,plain,(
  ~! [X0,X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11] : ((coll(X9,X0,X1) & perp(X9,X6,X0,X1) & coll(X8,X1,X2) & perp(X8,X6,X1,X2) & coll(X7,X0,X6) & coll(X7,X1,X2) & circle(X3,X0,X6,X11) & eqangle(X6,X0,X0,X1,X2,X0,X0,X4) & coll(X5,X0,X4) & circle(X3,X0,X5,X10) & midp(X4,X2,X1) & circle(X3,X0,X1,X2)) => perp(X8,X9,X5,X0))),
  inference(rectify,[],[f96])).
fof(f96,negated_conjecture,(
  ~! [X0,X1,X2,X7,X18,X8,X9,X19,X3,X12,X20,X21] : ((coll(X12,X0,X1) & perp(X12,X9,X0,X1) & coll(X3,X1,X2) & perp(X3,X9,X1,X2) & coll(X19,X0,X9) & coll(X19,X1,X2) & circle(X7,X0,X9,X21) & eqangle(X9,X0,X0,X1,X2,X0,X0,X18) & coll(X8,X0,X18) & circle(X7,X0,X8,X20) & midp(X18,X2,X1) & circle(X7,X0,X1,X2)) => perp(X3,X12,X8,X0))),
  inference(negated_conjecture,[],[f95])).
fof(f95,conjecture,(
  ! [X0,X1,X2,X7,X18,X8,X9,X19,X3,X12,X20,X21] : ((coll(X12,X0,X1) & perp(X12,X9,X0,X1) & coll(X3,X1,X2) & perp(X3,X9,X1,X2) & coll(X19,X0,X9) & coll(X19,X1,X2) & circle(X7,X0,X9,X21) & eqangle(X9,X0,X0,X1,X2,X0,X0,X18) & coll(X8,X0,X18) & circle(X7,X0,X8,X20) & midp(X18,X2,X1) & circle(X7,X0,X1,X2)) => perp(X3,X12,X8,X0))),
  file('./6_GDD_FULL/81-109/97.p',exemplo6GDDFULL8110997)).
fof(f22061,plain,(
  ( ! [X45,X43,X46,X44] : (perp(X43,X44,X45,X46)) ) | (~spl32_0 | ~spl32_162 | ~spl32_192 | ~spl32_236 | ~spl32_240 | ~spl32_372)),
  inference(subsumption_resolution,[],[f22029,f21308])).
fof(f21308,plain,(
  ( ! [X10,X8,X7,X9] : (para(X7,X8,X9,X10)) ) | (~spl32_0 | ~spl32_162 | ~spl32_192 | ~spl32_236 | ~spl32_372)),
  inference(subsumption_resolution,[],[f21284,f21232])).
fof(f21232,plain,(
  ( ! [X10,X11] : (midp(X10,sK22,X11)) ) | (~spl32_0 | ~spl32_162 | ~spl32_192 | ~spl32_236 | ~spl32_372)),
  inference(subsumption_resolution,[],[f21231,f16756])).
fof(f16756,plain,(
  ( ! [X4] : (midp(X4,sK20,sK21)) ) | (~spl32_192 | ~spl32_372)),
  inference(subsumption_resolution,[],[f16755,f8496])).
fof(f8496,plain,(
  ( ! [X12,X13] : (midp(X12,X13,X13)) ) | ~spl32_192),
  inference(resolution,[],[f8483,f3861])).
fof(f3861,plain,(
  ( ! [X12,X13,X11] : (~midp(X11,X13,X13) | midp(X11,X12,X12)) ) | ~spl32_192),
  inference(subsumption_resolution,[],[f3853,f2077])).
fof(f2077,plain,(
  ( ! [X2,X3] : (para(X2,X3,X2,X3)) ) | ~spl32_192),
  inference(avatar_component_clause,[],[f2076])).
fof(f2076,plain,(
  spl32_192 <=> ! [X3,X2] : para(X2,X3,X2,X3)),
  introduced(avatar_definition,[new_symbols(naming,[spl32_192])])).
fof(f3853,plain,(
  ( ! [X12,X13,X11] : (midp(X11,X12,X12) | ~para(X13,X12,X13,X12) | ~midp(X11,X13,X13)) ) | ~spl32_192),
  inference(resolution,[],[f2077,f423])).
fof(f423,plain,(
  ( ! [X4,X2,X0,X3,X1] : (~para(X0,X3,X1,X2) | midp(X4,X2,X3) | ~para(X0,X2,X1,X3) | ~midp(X4,X0,X1)) )),
  inference(cnf_transformation,[],[f264])).
fof(f264,plain,(
  ! [X0,X1,X2,X3,X4] : (midp(X4,X2,X3) | ~para(X0,X3,X1,X2) | ~para(X0,X2,X1,X3) | ~midp(X4,X0,X1))),
  inference(flattening,[],[f263])).
fof(f263,plain,(
  ! [X0,X1,X2,X3,X4] : (midp(X4,X2,X3) | (~para(X0,X3,X1,X2) | ~para(X0,X2,X1,X3) | ~midp(X4,X0,X1)))),
  inference(ennf_transformation,[],[f143])).
fof(f143,plain,(
  ! [X0,X1,X2,X3,X4] : ((para(X0,X3,X1,X2) & para(X0,X2,X1,X3) & midp(X4,X0,X1)) => midp(X4,X2,X3))),
  inference(rectify,[],[f65])).
fof(f65,axiom,(
  ! [X0,X1,X2,X3,X6] : ((para(X0,X3,X1,X2) & para(X0,X2,X1,X3) & midp(X6,X0,X1)) => midp(X6,X2,X3))),
  file('./6_GDD_FULL/81-109/97.p',ruleD64)).
fof(f8483,plain,(
  ( ! [X14] : (midp(X14,X14,X14)) ) | ~spl32_192),
  inference(subsumption_resolution,[],[f8462,f3850])).
fof(f3850,plain,(
  ( ! [X4,X3] : (coll(X3,X4,X4)) ) | ~spl32_192),
  inference(resolution,[],[f2077,f425])).
fof(f425,plain,(
  ( ! [X2,X0,X1] : (~para(X0,X1,X0,X2) | coll(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f267])).
fof(f267,plain,(
  ! [X0,X1,X2] : (coll(X0,X1,X2) | ~para(X0,X1,X0,X2))),
  inference(ennf_transformation,[],[f67])).
fof(f67,axiom,(
  ! [X0,X1,X2] : (para(X0,X1,X0,X2) => coll(X0,X1,X2))),
  file('./6_GDD_FULL/81-109/97.p',ruleD66)).
fof(f8462,plain,(
  ( ! [X14] : (~coll(X14,X14,X14) | midp(X14,X14,X14)) ) | ~spl32_192),
  inference(resolution,[],[f8376,f426])).
fof(f426,plain,(
  ( ! [X2,X0,X1] : (~cong(X0,X1,X0,X2) | ~coll(X0,X1,X2) | midp(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f269])).
fof(f269,plain,(
  ! [X0,X1,X2] : (midp(X0,X1,X2) | ~coll(X0,X1,X2) | ~cong(X0,X1,X0,X2))),
  inference(flattening,[],[f268])).
fof(f268,plain,(
  ! [X0,X1,X2] : (midp(X0,X1,X2) | (~coll(X0,X1,X2) | ~cong(X0,X1,X0,X2)))),
  inference(ennf_transformation,[],[f68])).
fof(f68,axiom,(
  ! [X0,X1,X2] : ((coll(X0,X1,X2) & cong(X0,X1,X0,X2)) => midp(X0,X1,X2))),
  file('./6_GDD_FULL/81-109/97.p',ruleD67)).
fof(f8376,plain,(
  ( ! [X35] : (cong(X35,X35,X35,X35)) ) | ~spl32_192),
  inference(subsumption_resolution,[],[f8372,f8334])).
fof(f8334,plain,(
  ( ! [X2,X0,X1] : (cyclic(X0,X0,X1,X2)) ) | ~spl32_192),
  inference(resolution,[],[f3221,f2077])).
fof(f3221,plain,(
  ( ! [X68,X66,X67] : (~para(X66,X67,X66,X67) | cyclic(X67,X67,X66,X68)) )),
  inference(subsumption_resolution,[],[f2218,f2219])).
fof(f2219,plain,(
  ( ! [X70,X71,X69] : (~para(X69,X70,X69,X70) | coll(X69,X71,X70) | cyclic(X70,X70,X69,X71)) )),
  inference(resolution,[],[f822,f400])).
fof(f400,plain,(
  ( ! [X2,X0,X3,X1] : (~eqangle(X2,X0,X2,X1,X3,X0,X3,X1) | coll(X2,X3,X0) | cyclic(X0,X1,X2,X3)) )),
  inference(cnf_transformation,[],[f222])).
fof(f222,plain,(
  ! [X0,X1,X2,X3] : (cyclic(X0,X1,X2,X3) | coll(X2,X3,X0) | ~eqangle(X2,X0,X2,X1,X3,X0,X3,X1))),
  inference(flattening,[],[f221])).
fof(f221,plain,(
  ! [X0,X1,X2,X3] : (cyclic(X0,X1,X2,X3) | (coll(X2,X3,X0) | ~eqangle(X2,X0,X2,X1,X3,X0,X3,X1)))),
  inference(ennf_transformation,[],[f121])).
fof(f121,plain,(
  ! [X0,X1,X2,X3] : ((~coll(X2,X3,X0) & eqangle(X2,X0,X2,X1,X3,X0,X3,X1)) => cyclic(X0,X1,X2,X3))),
  inference(rectify,[],[f42])).
fof(f42,axiom,(
  ! [X0,X1,X8,X9] : ((~coll(X8,X9,X0) & eqangle(X8,X0,X8,X1,X9,X0,X9,X1)) => cyclic(X0,X1,X8,X9))),
  file('./6_GDD_FULL/81-109/97.p',ruleD42a)).
fof(f822,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (eqangle(X0,X1,X2,X3,X4,X5,X4,X5) | ~para(X0,X1,X2,X3)) )),
  inference(resolution,[],[f379,f398])).
fof(f398,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (eqangle(X0,X1,X4,X5,X2,X3,X4,X5) | ~para(X0,X1,X2,X3)) )),
  inference(cnf_transformation,[],[f219])).
fof(f219,plain,(
  ! [X0,X1,X2,X3,X4,X5] : (eqangle(X0,X1,X4,X5,X2,X3,X4,X5) | ~para(X0,X1,X2,X3))),
  inference(ennf_transformation,[],[f119])).
fof(f119,plain,(
  ! [X0,X1,X2,X3,X4,X5] : (para(X0,X1,X2,X3) => eqangle(X0,X1,X4,X5,X2,X3,X4,X5))),
  inference(rectify,[],[f40])).
fof(f40,axiom,(
  ! [X0,X1,X2,X3,X8,X9] : (para(X0,X1,X2,X3) => eqangle(X0,X1,X8,X9,X2,X3,X8,X9))),
  file('./6_GDD_FULL/81-109/97.p',ruleD40)).
fof(f379,plain,(
  ( ! [X6,X4,X2,X0,X7,X5,X3,X1] : (~eqangle(X0,X1,X2,X3,X4,X5,X6,X7) | eqangle(X0,X1,X4,X5,X2,X3,X6,X7)) )),
  inference(cnf_transformation,[],[f195])).
fof(f195,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6,X7] : (eqangle(X0,X1,X4,X5,X2,X3,X6,X7) | ~eqangle(X0,X1,X2,X3,X4,X5,X6,X7))),
  inference(ennf_transformation,[],[f103])).
fof(f103,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6,X7] : (eqangle(X0,X1,X2,X3,X4,X5,X6,X7) => eqangle(X0,X1,X4,X5,X2,X3,X6,X7))),
  inference(rectify,[],[f21])).
fof(f21,axiom,(
  ! [X0,X1,X2,X3,X8,X9,X10,X11] : (eqangle(X0,X1,X2,X3,X8,X9,X10,X11) => eqangle(X0,X1,X8,X9,X2,X3,X10,X11))),
  file('./6_GDD_FULL/81-109/97.p',ruleD21)).
fof(f2218,plain,(
  ( ! [X68,X66,X67] : (~para(X66,X67,X66,X67) | ~coll(X66,X68,X67) | cyclic(X67,X67,X66,X68)) )),
  inference(resolution,[],[f822,f401])).
fof(f401,plain,(
  ( ! [X2,X0,X3,X1] : (~eqangle(X2,X0,X2,X1,X3,X0,X3,X1) | ~coll(X2,X3,X1) | cyclic(X0,X1,X2,X3)) )),
  inference(cnf_transformation,[],[f224])).
fof(f224,plain,(
  ! [X0,X1,X2,X3] : (cyclic(X0,X1,X2,X3) | ~coll(X2,X3,X1) | ~eqangle(X2,X0,X2,X1,X3,X0,X3,X1))),
  inference(flattening,[],[f223])).
fof(f223,plain,(
  ! [X0,X1,X2,X3] : (cyclic(X0,X1,X2,X3) | (~coll(X2,X3,X1) | ~eqangle(X2,X0,X2,X1,X3,X0,X3,X1)))),
  inference(ennf_transformation,[],[f122])).
fof(f122,plain,(
  ! [X0,X1,X2,X3] : ((coll(X2,X3,X1) & eqangle(X2,X0,X2,X1,X3,X0,X3,X1)) => cyclic(X0,X1,X2,X3))),
  inference(rectify,[],[f43])).
fof(f43,axiom,(
  ! [X0,X1,X8,X9] : ((coll(X8,X9,X1) & eqangle(X8,X0,X8,X1,X9,X0,X9,X1)) => cyclic(X0,X1,X8,X9))),
  file('./6_GDD_FULL/81-109/97.p',ruleD42b)).
fof(f8372,plain,(
  ( ! [X35,X36] : (cong(X35,X35,X35,X35) | ~cyclic(X35,X35,X36,X35)) ) | ~spl32_192),
  inference(duplicate_literal_removal,[],[f8359])).
fof(f8359,plain,(
  ( ! [X35,X36] : (cong(X35,X35,X35,X35) | ~cyclic(X35,X35,X36,X35) | ~cyclic(X35,X35,X36,X35)) ) | ~spl32_192),
  inference(resolution,[],[f8334,f1352])).
fof(f1352,plain,(
  ( ! [X10,X8,X7,X9] : (~cyclic(X7,X8,X9,X10) | cong(X7,X8,X7,X8) | ~cyclic(X7,X8,X9,X8) | ~cyclic(X7,X8,X9,X7)) )),
  inference(duplicate_literal_removal,[],[f1346])).
fof(f1346,plain,(
  ( ! [X10,X8,X7,X9] : (cong(X7,X8,X7,X8) | ~cyclic(X7,X8,X9,X10) | ~cyclic(X7,X8,X9,X8) | ~cyclic(X7,X8,X9,X7) | ~cyclic(X7,X8,X9,X10)) )),
  inference(resolution,[],[f402,f399])).
fof(f399,plain,(
  ( ! [X2,X0,X3,X1] : (eqangle(X2,X0,X2,X1,X3,X0,X3,X1) | ~cyclic(X0,X1,X2,X3)) )),
  inference(cnf_transformation,[],[f220])).
fof(f220,plain,(
  ! [X0,X1,X2,X3] : (eqangle(X2,X0,X2,X1,X3,X0,X3,X1) | ~cyclic(X0,X1,X2,X3))),
  inference(ennf_transformation,[],[f120])).
fof(f120,plain,(
  ! [X0,X1,X2,X3] : (cyclic(X0,X1,X2,X3) => eqangle(X2,X0,X2,X1,X3,X0,X3,X1))),
  inference(rectify,[],[f41])).
fof(f41,axiom,(
  ! [X0,X1,X8,X9] : (cyclic(X0,X1,X8,X9) => eqangle(X8,X0,X8,X1,X9,X0,X9,X1))),
  file('./6_GDD_FULL/81-109/97.p',ruleD41)).
fof(f402,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (~eqangle(X2,X0,X2,X1,X5,X3,X5,X4) | cong(X0,X1,X3,X4) | ~cyclic(X0,X1,X2,X5) | ~cyclic(X0,X1,X2,X4) | ~cyclic(X0,X1,X2,X3)) )),
  inference(cnf_transformation,[],[f226])).
fof(f226,plain,(
  ! [X0,X1,X2,X3,X4,X5] : (cong(X0,X1,X3,X4) | ~eqangle(X2,X0,X2,X1,X5,X3,X5,X4) | ~cyclic(X0,X1,X2,X5) | ~cyclic(X0,X1,X2,X4) | ~cyclic(X0,X1,X2,X3))),
  inference(flattening,[],[f225])).
fof(f225,plain,(
  ! [X0,X1,X2,X3,X4,X5] : (cong(X0,X1,X3,X4) | (~eqangle(X2,X0,X2,X1,X5,X3,X5,X4) | ~cyclic(X0,X1,X2,X5) | ~cyclic(X0,X1,X2,X4) | ~cyclic(X0,X1,X2,X3)))),
  inference(ennf_transformation,[],[f123])).
fof(f123,plain,(
  ! [X0,X1,X2,X3,X4,X5] : ((eqangle(X2,X0,X2,X1,X5,X3,X5,X4) & cyclic(X0,X1,X2,X5) & cyclic(X0,X1,X2,X4) & cyclic(X0,X1,X2,X3)) => cong(X0,X1,X3,X4))),
  inference(rectify,[],[f44])).
fof(f44,axiom,(
  ! [X0,X1,X2,X8,X9,X14] : ((eqangle(X2,X0,X2,X1,X14,X8,X14,X9) & cyclic(X0,X1,X2,X14) & cyclic(X0,X1,X2,X9) & cyclic(X0,X1,X2,X8)) => cong(X0,X1,X8,X9))),
  file('./6_GDD_FULL/81-109/97.p',ruleD43)).
fof(f16755,plain,(
  ( ! [X4] : (midp(X4,sK20,sK21) | ~midp(X4,sK20,sK20)) ) | ~spl32_372),
  inference(subsumption_resolution,[],[f16748,f9830])).
fof(f9830,plain,(
  para(sK20,sK20,sK20,sK21) | ~spl32_372),
  inference(subsumption_resolution,[],[f9820,f1559])).
fof(f1559,plain,(
  para(sK20,sK21,sK20,sK21)),
  inference(resolution,[],[f683,f518])).
fof(f518,plain,(
  perp(sK20,sK21,sK29,sK26)),
  inference(resolution,[],[f366,f482])).
fof(f482,plain,(
  perp(sK29,sK26,sK20,sK21)),
  inference(cnf_transformation,[],[f358])).
fof(f366,plain,(
  ( ! [X2,X0,X3,X1] : (~perp(X0,X1,X2,X3) | perp(X2,X3,X0,X1)) )),
  inference(cnf_transformation,[],[f177])).
fof(f177,plain,(
  ! [X0,X1,X2,X3] : (perp(X2,X3,X0,X1) | ~perp(X0,X1,X2,X3))),
  inference(ennf_transformation,[],[f8])).
fof(f8,axiom,(
  ! [X0,X1,X2,X3] : (perp(X0,X1,X2,X3) => perp(X2,X3,X0,X1))),
  file('./6_GDD_FULL/81-109/97.p',ruleD8)).
fof(f683,plain,(
  ( ! [X2,X3] : (~perp(X2,X3,sK29,sK26) | para(X2,X3,sK20,sK21)) )),
  inference(resolution,[],[f367,f482])).
fof(f367,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (~perp(X2,X3,X4,X5) | para(X0,X1,X4,X5) | ~perp(X0,X1,X2,X3)) )),
  inference(cnf_transformation,[],[f179])).
fof(f179,plain,(
  ! [X0,X1,X2,X3,X4,X5] : (para(X0,X1,X4,X5) | ~perp(X2,X3,X4,X5) | ~perp(X0,X1,X2,X3))),
  inference(flattening,[],[f178])).
fof(f178,plain,(
  ! [X0,X1,X2,X3,X4,X5] : (para(X0,X1,X4,X5) | (~perp(X2,X3,X4,X5) | ~perp(X0,X1,X2,X3)))),
  inference(ennf_transformation,[],[f9])).
fof(f9,axiom,(
  ! [X0,X1,X2,X3,X4,X5] : ((perp(X2,X3,X4,X5) & perp(X0,X1,X2,X3)) => para(X0,X1,X4,X5))),
  file('./6_GDD_FULL/81-109/97.p',ruleD9)).
fof(f9820,plain,(
  para(sK20,sK20,sK20,sK21) | ~para(sK20,sK21,sK20,sK21) | ~spl32_372),
  inference(resolution,[],[f4824,f1065])).
fof(f1065,plain,(
  ( ! [X21,X19,X20] : (~cong(X21,X19,X21,X20) | para(X21,X19,X19,X20) | ~para(X19,X20,X21,X20)) )),
  inference(resolution,[],[f432,f405])).
fof(f405,plain,(
  ( ! [X2,X0,X1] : (eqangle(X2,X0,X0,X1,X0,X1,X2,X1) | ~cong(X2,X0,X2,X1)) )),
  inference(cnf_transformation,[],[f231])).
fof(f231,plain,(
  ! [X0,X1,X2] : (eqangle(X2,X0,X0,X1,X0,X1,X2,X1) | ~cong(X2,X0,X2,X1))),
  inference(ennf_transformation,[],[f126])).
fof(f126,plain,(
  ! [X0,X1,X2] : (cong(X2,X0,X2,X1) => eqangle(X2,X0,X0,X1,X0,X1,X2,X1))),
  inference(rectify,[],[f47])).
fof(f47,axiom,(
  ! [X0,X1,X7] : (cong(X7,X0,X7,X1) => eqangle(X7,X0,X0,X1,X0,X1,X7,X1))),
  file('./6_GDD_FULL/81-109/97.p',ruleD46)).
fof(f432,plain,(
  ( ! [X6,X4,X2,X0,X7,X5,X3,X1] : (~eqangle(X0,X1,X2,X3,X4,X5,X6,X7) | ~para(X4,X5,X6,X7) | para(X0,X1,X2,X3)) )),
  inference(cnf_transformation,[],[f279])).
fof(f279,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6,X7] : (para(X0,X1,X2,X3) | ~para(X4,X5,X6,X7) | ~eqangle(X0,X1,X2,X3,X4,X5,X6,X7))),
  inference(flattening,[],[f278])).
fof(f278,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6,X7] : (para(X0,X1,X2,X3) | (~para(X4,X5,X6,X7) | ~eqangle(X0,X1,X2,X3,X4,X5,X6,X7)))),
  inference(ennf_transformation,[],[f146])).
fof(f146,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6,X7] : ((para(X4,X5,X6,X7) & eqangle(X0,X1,X2,X3,X4,X5,X6,X7)) => para(X0,X1,X2,X3))),
  inference(rectify,[],[f74])).
fof(f74,axiom,(
  ! [X0,X1,X2,X3,X8,X9,X10,X11] : ((para(X8,X9,X10,X11) & eqangle(X0,X1,X2,X3,X8,X9,X10,X11)) => para(X0,X1,X2,X3))),
  file('./6_GDD_FULL/81-109/97.p',ruleD73)).
fof(f4824,plain,(
  cong(sK20,sK20,sK20,sK21) | ~spl32_372),
  inference(avatar_component_clause,[],[f4823])).
fof(f4823,plain,(
  spl32_372 <=> cong(sK20,sK20,sK20,sK21)),
  introduced(avatar_definition,[new_symbols(naming,[spl32_372])])).
fof(f16748,plain,(
  ( ! [X4] : (midp(X4,sK20,sK21) | ~para(sK20,sK20,sK20,sK21) | ~midp(X4,sK20,sK20)) ) | ~spl32_372),
  inference(resolution,[],[f9827,f423])).
fof(f9827,plain,(
  para(sK20,sK21,sK20,sK20) | ~spl32_372),
  inference(resolution,[],[f4824,f2179])).
fof(f2179,plain,(
  ( ! [X0,X1] : (~cong(X0,X0,X0,X1) | para(X0,X1,X0,X0)) )),
  inference(resolution,[],[f819,f397])).
fof(f397,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (~eqangle(X0,X1,X4,X5,X2,X3,X4,X5) | para(X0,X1,X2,X3)) )),
  inference(cnf_transformation,[],[f218])).
fof(f218,plain,(
  ! [X0,X1,X2,X3,X4,X5] : (para(X0,X1,X2,X3) | ~eqangle(X0,X1,X4,X5,X2,X3,X4,X5))),
  inference(ennf_transformation,[],[f118])).
fof(f118,plain,(
  ! [X0,X1,X2,X3,X4,X5] : (eqangle(X0,X1,X4,X5,X2,X3,X4,X5) => para(X0,X1,X2,X3))),
  inference(rectify,[],[f39])).
fof(f39,axiom,(
  ! [X0,X1,X2,X3,X8,X9] : (eqangle(X0,X1,X8,X9,X2,X3,X8,X9) => para(X0,X1,X2,X3))),
  file('./6_GDD_FULL/81-109/97.p',ruleD39)).
fof(f819,plain,(
  ( ! [X12,X10,X11] : (eqangle(X10,X11,X12,X11,X12,X10,X10,X11) | ~cong(X12,X10,X12,X11)) )),
  inference(resolution,[],[f378,f405])).
fof(f378,plain,(
  ( ! [X6,X4,X2,X0,X7,X5,X3,X1] : (~eqangle(X0,X1,X2,X3,X4,X5,X6,X7) | eqangle(X4,X5,X6,X7,X0,X1,X2,X3)) )),
  inference(cnf_transformation,[],[f194])).
fof(f194,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6,X7] : (eqangle(X4,X5,X6,X7,X0,X1,X2,X3) | ~eqangle(X0,X1,X2,X3,X4,X5,X6,X7))),
  inference(ennf_transformation,[],[f102])).
fof(f102,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6,X7] : (eqangle(X0,X1,X2,X3,X4,X5,X6,X7) => eqangle(X4,X5,X6,X7,X0,X1,X2,X3))),
  inference(rectify,[],[f20])).
fof(f20,axiom,(
  ! [X0,X1,X2,X3,X8,X9,X10,X11] : (eqangle(X0,X1,X2,X3,X8,X9,X10,X11) => eqangle(X8,X9,X10,X11,X0,X1,X2,X3))),
  file('./6_GDD_FULL/81-109/97.p',ruleD20)).
fof(f21231,plain,(
  ( ! [X10,X11] : (midp(X10,sK22,X11) | ~midp(X10,sK20,sK21)) ) | (~spl32_0 | ~spl32_162 | ~spl32_192 | ~spl32_236)),
  inference(subsumption_resolution,[],[f21218,f18781])).
fof(f18781,plain,(
  ( ! [X6,X5] : (para(X5,sK22,sK21,X6)) ) | (~spl32_192 | ~spl32_236)),
  inference(resolution,[],[f11318,f362])).
fof(f362,plain,(
  ( ! [X2,X0,X3,X1] : (~para(X0,X1,X2,X3) | para(X0,X1,X3,X2)) )),
  inference(cnf_transformation,[],[f172])).
fof(f172,plain,(
  ! [X0,X1,X2,X3] : (para(X0,X1,X3,X2) | ~para(X0,X1,X2,X3))),
  inference(ennf_transformation,[],[f4])).
fof(f4,axiom,(
  ! [X0,X1,X2,X3] : (para(X0,X1,X2,X3) => para(X0,X1,X3,X2))),
  file('./6_GDD_FULL/81-109/97.p',ruleD4)).
fof(f11318,plain,(
  ( ! [X45,X44] : (para(X44,sK22,X45,sK21)) ) | (~spl32_192 | ~spl32_236)),
  inference(subsumption_resolution,[],[f11317,f8622])).
fof(f8622,plain,(
  ( ! [X17,X16] : (midp(X16,X17,X16)) ) | (~spl32_192 | ~spl32_236)),
  inference(resolution,[],[f8496,f7726])).
fof(f7726,plain,(
  ( ! [X28,X29,X27] : (~midp(X27,X28,X29) | midp(X28,X27,X29)) ) | (~spl32_192 | ~spl32_236)),
  inference(subsumption_resolution,[],[f7676,f3900])).
fof(f3900,plain,(
  ( ! [X2,X0,X1] : (coll(X0,X1,X2)) ) | ~spl32_192),
  inference(subsumption_resolution,[],[f3897,f3883])).
fof(f3883,plain,(
  ( ! [X6,X5] : (coll(X5,X5,X6)) ) | ~spl32_192),
  inference(resolution,[],[f3863,f359])).
fof(f359,plain,(
  ( ! [X2,X0,X1] : (~coll(X0,X1,X2) | coll(X0,X2,X1)) )),
  inference(cnf_transformation,[],[f168])).
fof(f168,plain,(
  ! [X0,X1,X2] : (coll(X0,X2,X1) | ~coll(X0,X1,X2))),
  inference(ennf_transformation,[],[f1])).
fof(f1,axiom,(
  ! [X0,X1,X2] : (coll(X0,X1,X2) => coll(X0,X2,X1))),
  file('./6_GDD_FULL/81-109/97.p',ruleD1)).
fof(f3863,plain,(
  ( ! [X4,X3] : (coll(X3,X4,X3)) ) | ~spl32_192),
  inference(resolution,[],[f3850,f360])).
fof(f360,plain,(
  ( ! [X2,X0,X1] : (~coll(X0,X1,X2) | coll(X1,X0,X2)) )),
  inference(cnf_transformation,[],[f169])).
fof(f169,plain,(
  ! [X0,X1,X2] : (coll(X1,X0,X2) | ~coll(X0,X1,X2))),
  inference(ennf_transformation,[],[f2])).
fof(f2,axiom,(
  ! [X0,X1,X2] : (coll(X0,X1,X2) => coll(X1,X0,X2))),
  file('./6_GDD_FULL/81-109/97.p',ruleD2)).
fof(f3897,plain,(
  ( ! [X2,X0,X1] : (coll(X0,X1,X2) | ~coll(X2,X2,X0)) ) | ~spl32_192),
  inference(resolution,[],[f3883,f361])).
fof(f361,plain,(
  ( ! [X2,X0,X3,X1] : (~coll(X0,X1,X3) | coll(X2,X3,X0) | ~coll(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f171])).
fof(f171,plain,(
  ! [X0,X1,X2,X3] : (coll(X2,X3,X0) | ~coll(X0,X1,X3) | ~coll(X0,X1,X2))),
  inference(flattening,[],[f170])).
fof(f170,plain,(
  ! [X0,X1,X2,X3] : (coll(X2,X3,X0) | (~coll(X0,X1,X3) | ~coll(X0,X1,X2)))),
  inference(ennf_transformation,[],[f3])).
fof(f3,axiom,(
  ! [X0,X1,X2,X3] : ((coll(X0,X1,X3) & coll(X0,X1,X2)) => coll(X2,X3,X0))),
  file('./6_GDD_FULL/81-109/97.p',ruleD3)).
fof(f7676,plain,(
  ( ! [X28,X29,X27] : (~midp(X27,X28,X29) | ~coll(X28,X27,X29) | midp(X28,X27,X29)) ) | ~spl32_236),
  inference(resolution,[],[f2727,f426])).
fof(f2727,plain,(
  ( ! [X14,X15,X13] : (cong(X14,X13,X14,X15) | ~midp(X13,X14,X15)) ) | ~spl32_236),
  inference(avatar_component_clause,[],[f2726])).
fof(f2726,plain,(
  spl32_236 <=> ! [X13,X15,X14] : (~midp(X13,X14,X15) | cong(X14,X13,X14,X15))),
  introduced(avatar_definition,[new_symbols(naming,[spl32_236])])).
fof(f11317,plain,(
  ( ! [X45,X44] : (para(X44,sK22,X45,sK21) | ~midp(sK24,X44,sK24)) ) | ~spl32_192),
  inference(subsumption_resolution,[],[f11303,f8706])).
fof(f8706,plain,(
  ( ! [X8] : (midp(sK21,sK22,X8)) ) | ~spl32_192),
  inference(resolution,[],[f8510,f3966])).
fof(f3966,plain,(
  ( ! [X12,X10,X11] : (~midp(X12,X11,X10) | midp(X10,X11,X12)) ) | ~spl32_192),
  inference(subsumption_resolution,[],[f3959,f3900])).
fof(f3959,plain,(
  ( ! [X12,X10,X11] : (~coll(X10,X11,X12) | midp(X10,X11,X12) | ~midp(X12,X11,X10)) ) | ~spl32_192),
  inference(resolution,[],[f3857,f404])).
fof(f404,plain,(
  ( ! [X4,X2,X0,X3,X1] : (~para(X3,X4,X1,X2) | ~coll(X4,X0,X2) | midp(X4,X0,X2) | ~midp(X3,X0,X1)) )),
  inference(cnf_transformation,[],[f230])).
fof(f230,plain,(
  ! [X0,X1,X2,X3,X4] : (midp(X4,X0,X2) | ~coll(X4,X0,X2) | ~para(X3,X4,X1,X2) | ~midp(X3,X0,X1))),
  inference(flattening,[],[f229])).
fof(f229,plain,(
  ! [X0,X1,X2,X3,X4] : (midp(X4,X0,X2) | (~coll(X4,X0,X2) | ~para(X3,X4,X1,X2) | ~midp(X3,X0,X1)))),
  inference(ennf_transformation,[],[f125])).
fof(f125,plain,(
  ! [X0,X1,X2,X3,X4] : ((coll(X4,X0,X2) & para(X3,X4,X1,X2) & midp(X3,X0,X1)) => midp(X4,X0,X2))),
  inference(rectify,[],[f46])).
fof(f46,axiom,(
  ! [X0,X1,X2,X4,X5] : ((coll(X5,X0,X2) & para(X4,X5,X1,X2) & midp(X4,X0,X1)) => midp(X5,X0,X2))),
  file('./6_GDD_FULL/81-109/97.p',ruleD45)).
fof(f3857,plain,(
  ( ! [X24,X23] : (para(X23,X24,X24,X23)) ) | ~spl32_192),
  inference(resolution,[],[f2077,f362])).
fof(f8510,plain,(
  ( ! [X0] : (midp(X0,sK22,sK21)) ) | ~spl32_192),
  inference(resolution,[],[f8509,f369])).
fof(f369,plain,(
  ( ! [X2,X0,X1] : (~midp(X2,X1,X0) | midp(X2,X0,X1)) )),
  inference(cnf_transformation,[],[f182])).
fof(f182,plain,(
  ! [X0,X1,X2] : (midp(X2,X0,X1) | ~midp(X2,X1,X0))),
  inference(ennf_transformation,[],[f97])).
fof(f97,plain,(
  ! [X0,X1,X2] : (midp(X2,X1,X0) => midp(X2,X0,X1))),
  inference(rectify,[],[f11])).
fof(f11,axiom,(
  ! [X0,X1,X6] : (midp(X6,X1,X0) => midp(X6,X0,X1))),
  file('./6_GDD_FULL/81-109/97.p',ruleD11)).
fof(f8509,plain,(
  ( ! [X33] : (midp(X33,sK21,sK22)) ) | ~spl32_192),
  inference(subsumption_resolution,[],[f8506,f8496])).
fof(f8506,plain,(
  ( ! [X33] : (midp(X33,sK21,sK22) | ~midp(X33,sK24,sK24)) ) | ~spl32_192),
  inference(resolution,[],[f8483,f1415])).
fof(f1415,plain,(
  ( ! [X2,X3,X1] : (~midp(sK24,X1,X2) | midp(X3,sK21,sK22) | ~midp(X3,X1,X2)) )),
  inference(subsumption_resolution,[],[f1410,f564])).
fof(f564,plain,(
  ( ! [X2,X3] : (para(X2,sK21,X3,sK22) | ~midp(sK24,X2,X3)) )),
  inference(resolution,[],[f422,f497])).
fof(f497,plain,(
  midp(sK24,sK21,sK22)),
  inference(resolution,[],[f369,f473])).
fof(f473,plain,(
  midp(sK24,sK22,sK21)),
  inference(cnf_transformation,[],[f358])).
fof(f422,plain,(
  ( ! [X4,X2,X0,X3,X1] : (~midp(X4,X2,X3) | para(X0,X2,X1,X3) | ~midp(X4,X0,X1)) )),
  inference(cnf_transformation,[],[f262])).
fof(f262,plain,(
  ! [X0,X1,X2,X3,X4] : (para(X0,X2,X1,X3) | ~midp(X4,X2,X3) | ~midp(X4,X0,X1))),
  inference(flattening,[],[f261])).
fof(f261,plain,(
  ! [X0,X1,X2,X3,X4] : (para(X0,X2,X1,X3) | (~midp(X4,X2,X3) | ~midp(X4,X0,X1)))),
  inference(ennf_transformation,[],[f142])).
fof(f142,plain,(
  ! [X0,X1,X2,X3,X4] : ((midp(X4,X2,X3) & midp(X4,X0,X1)) => para(X0,X2,X1,X3))),
  inference(rectify,[],[f64])).
fof(f64,axiom,(
  ! [X0,X1,X2,X3,X6] : ((midp(X6,X2,X3) & midp(X6,X0,X1)) => para(X0,X2,X1,X3))),
  file('./6_GDD_FULL/81-109/97.p',ruleD63)).
fof(f1410,plain,(
  ( ! [X2,X3,X1] : (~midp(sK24,X1,X2) | midp(X3,sK21,sK22) | ~para(X1,sK21,X2,sK22) | ~midp(X3,X1,X2)) )),
  inference(resolution,[],[f563,f423])).
fof(f563,plain,(
  ( ! [X0,X1] : (para(X0,sK22,X1,sK21) | ~midp(sK24,X0,X1)) )),
  inference(resolution,[],[f422,f473])).
fof(f11303,plain,(
  ( ! [X45,X44] : (para(X44,sK22,X45,sK21) | ~midp(sK21,sK22,X45) | ~midp(sK24,X44,sK24)) )),
  inference(resolution,[],[f1402,f1414])).
fof(f1414,plain,(
  ( ! [X14,X13] : (para(X13,sK22,sK21,X14) | ~midp(sK24,X13,X14)) )),
  inference(resolution,[],[f563,f362])).
fof(f1402,plain,(
  ( ! [X10,X8,X7,X9] : (~para(X9,X10,X7,sK24) | para(X9,X10,X8,sK21) | ~midp(X7,sK22,X8)) )),
  inference(resolution,[],[f558,f364])).
fof(f364,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (~para(X2,X3,X4,X5) | para(X0,X1,X4,X5) | ~para(X0,X1,X2,X3)) )),
  inference(cnf_transformation,[],[f175])).
fof(f175,plain,(
  ! [X0,X1,X2,X3,X4,X5] : (para(X0,X1,X4,X5) | ~para(X2,X3,X4,X5) | ~para(X0,X1,X2,X3))),
  inference(flattening,[],[f174])).
fof(f174,plain,(
  ! [X0,X1,X2,X3,X4,X5] : (para(X0,X1,X4,X5) | (~para(X2,X3,X4,X5) | ~para(X0,X1,X2,X3)))),
  inference(ennf_transformation,[],[f6])).
fof(f6,axiom,(
  ! [X0,X1,X2,X3,X4,X5] : ((para(X2,X3,X4,X5) & para(X0,X1,X2,X3)) => para(X0,X1,X4,X5))),
  file('./6_GDD_FULL/81-109/97.p',ruleD6)).
fof(f558,plain,(
  ( ! [X0,X1] : (para(X0,sK24,X1,sK21) | ~midp(X0,sK22,X1)) )),
  inference(resolution,[],[f403,f473])).
fof(f403,plain,(
  ( ! [X4,X2,X0,X3,X1] : (~midp(X4,X0,X2) | para(X3,X4,X1,X2) | ~midp(X3,X0,X1)) )),
  inference(cnf_transformation,[],[f228])).
fof(f228,plain,(
  ! [X0,X1,X2,X3,X4] : (para(X3,X4,X1,X2) | ~midp(X4,X0,X2) | ~midp(X3,X0,X1))),
  inference(flattening,[],[f227])).
fof(f227,plain,(
  ! [X0,X1,X2,X3,X4] : (para(X3,X4,X1,X2) | (~midp(X4,X0,X2) | ~midp(X3,X0,X1)))),
  inference(ennf_transformation,[],[f124])).
fof(f124,plain,(
  ! [X0,X1,X2,X3,X4] : ((midp(X4,X0,X2) & midp(X3,X0,X1)) => para(X3,X4,X1,X2))),
  inference(rectify,[],[f45])).
fof(f45,axiom,(
  ! [X0,X1,X2,X4,X5] : ((midp(X5,X0,X2) & midp(X4,X0,X1)) => para(X4,X5,X1,X2))),
  file('./6_GDD_FULL/81-109/97.p',ruleD44)).
fof(f21218,plain,(
  ( ! [X10,X11] : (midp(X10,sK22,X11) | ~para(sK20,sK22,sK21,X11) | ~midp(X10,sK20,sK21)) ) | (~spl32_0 | ~spl32_162 | ~spl32_192 | ~spl32_236)),
  inference(resolution,[],[f21114,f423])).
fof(f21114,plain,(
  ( ! [X0] : (para(sK20,X0,sK21,sK22)) ) | (~spl32_0 | ~spl32_162 | ~spl32_192 | ~spl32_236)),
  inference(resolution,[],[f21102,f2348])).
fof(f2348,plain,(
  ( ! [X8,X7] : (~perp(X7,X8,sK20,sK21) | para(X7,X8,sK21,sK22)) ) | ~spl32_0),
  inference(resolution,[],[f616,f367])).
fof(f616,plain,(
  perp(sK20,sK21,sK21,sK22) | ~spl32_0),
  inference(avatar_component_clause,[],[f615])).
fof(f615,plain,(
  spl32_0 <=> perp(sK20,sK21,sK21,sK22)),
  introduced(avatar_definition,[new_symbols(naming,[spl32_0])])).
fof(f21102,plain,(
  ( ! [X10] : (perp(sK20,X10,sK20,sK21)) ) | (~spl32_162 | ~spl32_192 | ~spl32_236)),
  inference(subsumption_resolution,[],[f1835,f21028])).
fof(f21028,plain,(
  ( ! [X6,X8,X7] : (perp(X6,X8,X8,X7)) ) | (~spl32_192 | ~spl32_236)),
  inference(subsumption_resolution,[],[f21003,f3900])).
fof(f21003,plain,(
  ( ! [X6,X8,X7,X5] : (~coll(X5,X6,X7) | perp(X6,X8,X8,X7)) ) | (~spl32_192 | ~spl32_236)),
  inference(resolution,[],[f10843,f412])).
fof(f412,plain,(
  ( ! [X2,X0,X3,X1] : (~circle(X3,X0,X1,X2) | ~coll(X3,X0,X2) | perp(X0,X1,X1,X2)) )),
  inference(cnf_transformation,[],[f245])).
fof(f245,plain,(
  ! [X0,X1,X2,X3] : (perp(X0,X1,X1,X2) | ~coll(X3,X0,X2) | ~circle(X3,X0,X1,X2))),
  inference(flattening,[],[f244])).
fof(f244,plain,(
  ! [X0,X1,X2,X3] : (perp(X0,X1,X1,X2) | (~coll(X3,X0,X2) | ~circle(X3,X0,X1,X2)))),
  inference(ennf_transformation,[],[f133])).
fof(f133,plain,(
  ! [X0,X1,X2,X3] : ((coll(X3,X0,X2) & circle(X3,X0,X1,X2)) => perp(X0,X1,X1,X2))),
  inference(rectify,[],[f54])).
fof(f54,axiom,(
  ! [X0,X1,X2,X7] : ((coll(X7,X0,X2) & circle(X7,X0,X1,X2)) => perp(X0,X1,X1,X2))),
  file('./6_GDD_FULL/81-109/97.p',ruleD53)).
fof(f10843,plain,(
  ( ! [X2,X0,X3,X1] : (circle(X0,X2,X3,X1)) ) | (~spl32_192 | ~spl32_236)),
  inference(subsumption_resolution,[],[f10842,f10708])).
fof(f10708,plain,(
  ( ! [X4,X2,X3] : (cong(X2,X3,X2,X4)) ) | ~spl32_192),
  inference(subsumption_resolution,[],[f10661,f8620])).
fof(f8620,plain,(
  ( ! [X12,X13] : (midp(X12,X12,X13)) ) | ~spl32_192),
  inference(resolution,[],[f8496,f3966])).
fof(f10661,plain,(
  ( ! [X4,X2,X3] : (cong(X2,X3,X2,X4) | ~midp(X2,X2,X3)) ) | ~spl32_192),
  inference(resolution,[],[f1967,f8620])).
fof(f1967,plain,(
  ( ! [X2,X0,X3,X1] : (~midp(X0,X3,X2) | cong(X0,X1,X0,X2) | ~midp(X0,X3,X1)) )),
  inference(resolution,[],[f710,f530])).
fof(f530,plain,(
  ( ! [X2,X0,X1] : (cong(X0,X1,X0,X2) | ~midp(X0,X2,X1)) )),
  inference(resolution,[],[f382,f427])).
fof(f427,plain,(
  ( ! [X2,X0,X1] : (cong(X0,X1,X0,X2) | ~midp(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f270])).
fof(f270,plain,(
  ! [X0,X1,X2] : (cong(X0,X1,X0,X2) | ~midp(X0,X1,X2))),
  inference(ennf_transformation,[],[f69])).
fof(f69,axiom,(
  ! [X0,X1,X2] : (midp(X0,X1,X2) => cong(X0,X1,X0,X2))),
  file('./6_GDD_FULL/81-109/97.p',ruleD68)).
fof(f382,plain,(
  ( ! [X2,X0,X3,X1] : (~cong(X0,X1,X2,X3) | cong(X2,X3,X0,X1)) )),
  inference(cnf_transformation,[],[f199])).
fof(f199,plain,(
  ! [X0,X1,X2,X3] : (cong(X2,X3,X0,X1) | ~cong(X0,X1,X2,X3))),
  inference(ennf_transformation,[],[f24])).
fof(f24,axiom,(
  ! [X0,X1,X2,X3] : (cong(X0,X1,X2,X3) => cong(X2,X3,X0,X1))),
  file('./6_GDD_FULL/81-109/97.p',ruleD24)).
fof(f710,plain,(
  ( ! [X4,X2,X0,X3,X1] : (~cong(X0,X1,X2,X4) | cong(X0,X1,X2,X3) | ~midp(X2,X4,X3)) )),
  inference(resolution,[],[f383,f427])).
fof(f383,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (~cong(X2,X3,X4,X5) | cong(X0,X1,X4,X5) | ~cong(X0,X1,X2,X3)) )),
  inference(cnf_transformation,[],[f201])).
fof(f201,plain,(
  ! [X0,X1,X2,X3,X4,X5] : (cong(X0,X1,X4,X5) | ~cong(X2,X3,X4,X5) | ~cong(X0,X1,X2,X3))),
  inference(flattening,[],[f200])).
fof(f200,plain,(
  ! [X0,X1,X2,X3,X4,X5] : (cong(X0,X1,X4,X5) | (~cong(X2,X3,X4,X5) | ~cong(X0,X1,X2,X3)))),
  inference(ennf_transformation,[],[f25])).
fof(f25,axiom,(
  ! [X0,X1,X2,X3,X4,X5] : ((cong(X2,X3,X4,X5) & cong(X0,X1,X2,X3)) => cong(X0,X1,X4,X5))),
  file('./6_GDD_FULL/81-109/97.p',ruleD25)).
fof(f10842,plain,(
  ( ! [X2,X0,X3,X1] : (circle(X0,X2,X3,X1) | ~cong(X0,X2,X0,X3)) ) | (~spl32_192 | ~spl32_236)),
  inference(subsumption_resolution,[],[f10841,f8622])).
fof(f10841,plain,(
  ( ! [X2,X0,X3,X1] : (~midp(X0,X2,X0) | circle(X0,X2,X3,X1) | ~cong(X0,X2,X0,X3)) ) | ~spl32_192),
  inference(subsumption_resolution,[],[f10800,f8620])).
fof(f10800,plain,(
  ( ! [X2,X0,X3,X1] : (~midp(X0,X0,X1) | ~midp(X0,X2,X0) | circle(X0,X2,X3,X1) | ~cong(X0,X2,X0,X3)) )),
  inference(resolution,[],[f1970,f370])).
fof(f370,plain,(
  ( ! [X2,X0,X3,X1] : (~cong(X3,X0,X3,X2) | circle(X3,X0,X1,X2) | ~cong(X3,X0,X3,X1)) )),
  inference(cnf_transformation,[],[f184])).
fof(f184,plain,(
  ! [X0,X1,X2,X3] : (circle(X3,X0,X1,X2) | ~cong(X3,X0,X3,X2) | ~cong(X3,X0,X3,X1))),
  inference(flattening,[],[f183])).
fof(f183,plain,(
  ! [X0,X1,X2,X3] : (circle(X3,X0,X1,X2) | (~cong(X3,X0,X3,X2) | ~cong(X3,X0,X3,X1)))),
  inference(ennf_transformation,[],[f98])).
fof(f98,plain,(
  ! [X0,X1,X2,X3] : ((cong(X3,X0,X3,X2) & cong(X3,X0,X3,X1)) => circle(X3,X0,X1,X2))),
  inference(rectify,[],[f12])).
fof(f12,axiom,(
  ! [X0,X1,X2,X7] : ((cong(X7,X0,X7,X2) & cong(X7,X0,X7,X1)) => circle(X7,X0,X1,X2))),
  file('./6_GDD_FULL/81-109/97.p',ruleD12)).
fof(f1970,plain,(
  ( ! [X14,X12,X15,X13] : (cong(X12,X13,X14,X15) | ~midp(X14,X12,X15) | ~midp(X12,X13,X14)) )),
  inference(resolution,[],[f710,f527])).
fof(f527,plain,(
  ( ! [X2,X0,X1] : (cong(X0,X1,X2,X0) | ~midp(X0,X1,X2)) )),
  inference(resolution,[],[f381,f427])).
fof(f381,plain,(
  ( ! [X2,X0,X3,X1] : (~cong(X0,X1,X2,X3) | cong(X0,X1,X3,X2)) )),
  inference(cnf_transformation,[],[f198])).
fof(f198,plain,(
  ! [X0,X1,X2,X3] : (cong(X0,X1,X3,X2) | ~cong(X0,X1,X2,X3))),
  inference(ennf_transformation,[],[f23])).
fof(f23,axiom,(
  ! [X0,X1,X2,X3] : (cong(X0,X1,X2,X3) => cong(X0,X1,X3,X2))),
  file('./6_GDD_FULL/81-109/97.p',ruleD23)).
fof(f1835,plain,(
  ( ! [X10] : (~perp(sK23,sK20,sK20,X10) | perp(sK20,X10,sK20,sK21)) ) | ~spl32_162),
  inference(avatar_component_clause,[],[f1834])).
fof(f1834,plain,(
  spl32_162 <=> ! [X10] : (~perp(sK23,sK20,sK20,X10) | perp(sK20,X10,sK20,sK21))),
  introduced(avatar_definition,[new_symbols(naming,[spl32_162])])).
fof(f21284,plain,(
  ( ! [X10,X8,X7,X9] : (para(X7,X8,X9,X10) | ~midp(X7,sK22,X9)) ) | (~spl32_0 | ~spl32_162 | ~spl32_192 | ~spl32_236 | ~spl32_372)),
  inference(resolution,[],[f21232,f403])).
fof(f22029,plain,(
  ( ! [X47,X45,X43,X46,X44] : (perp(X43,X44,X45,X46) | ~para(X43,X44,X46,X47)) ) | (~spl32_192 | ~spl32_240)),
  inference(resolution,[],[f21199,f368])).
fof(f368,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (~perp(X2,X3,X4,X5) | perp(X0,X1,X4,X5) | ~para(X0,X1,X2,X3)) )),
  inference(cnf_transformation,[],[f181])).
fof(f181,plain,(
  ! [X0,X1,X2,X3,X4,X5] : (perp(X0,X1,X4,X5) | ~perp(X2,X3,X4,X5) | ~para(X0,X1,X2,X3))),
  inference(flattening,[],[f180])).
fof(f180,plain,(
  ! [X0,X1,X2,X3,X4,X5] : (perp(X0,X1,X4,X5) | (~perp(X2,X3,X4,X5) | ~para(X0,X1,X2,X3)))),
  inference(ennf_transformation,[],[f10])).
fof(f10,axiom,(
  ! [X0,X1,X2,X3,X4,X5] : ((perp(X2,X3,X4,X5) & para(X0,X1,X2,X3)) => perp(X0,X1,X4,X5))),
  file('./6_GDD_FULL/81-109/97.p',ruleD10)).
fof(f21199,plain,(
  ( ! [X116,X114,X115] : (perp(X115,X114,X116,X115)) ) | (~spl32_192 | ~spl32_240)),
  inference(subsumption_resolution,[],[f21198,f21072])).
fof(f21072,plain,(
  ( ! [X19,X17,X18] : (cong(X18,X17,X19,X17)) ) | ~spl32_192),
  inference(subsumption_resolution,[],[f21033,f8496])).
fof(f21033,plain,(
  ( ! [X19,X17,X18] : (~midp(X17,X18,X18) | cong(X18,X17,X19,X17)) ) | ~spl32_192),
  inference(resolution,[],[f13929,f411])).
fof(f411,plain,(
  ( ! [X2,X0,X3,X1] : (~perp(X0,X1,X1,X2) | ~midp(X3,X0,X2) | cong(X0,X3,X1,X3)) )),
  inference(cnf_transformation,[],[f243])).
fof(f243,plain,(
  ! [X0,X1,X2,X3] : (cong(X0,X3,X1,X3) | ~midp(X3,X0,X2) | ~perp(X0,X1,X1,X2))),
  inference(flattening,[],[f242])).
fof(f242,plain,(
  ! [X0,X1,X2,X3] : (cong(X0,X3,X1,X3) | (~midp(X3,X0,X2) | ~perp(X0,X1,X1,X2)))),
  inference(ennf_transformation,[],[f132])).
fof(f132,plain,(
  ! [X0,X1,X2,X3] : ((midp(X3,X0,X2) & perp(X0,X1,X1,X2)) => cong(X0,X3,X1,X3))),
  inference(rectify,[],[f53])).
fof(f53,axiom,(
  ! [X0,X1,X2,X6] : ((midp(X6,X0,X2) & perp(X0,X1,X1,X2)) => cong(X0,X6,X1,X6))),
  file('./6_GDD_FULL/81-109/97.p',ruleD52)).
fof(f13929,plain,(
  ( ! [X8,X7] : (perp(X7,X8,X8,X7)) ) | ~spl32_192),
  inference(subsumption_resolution,[],[f13928,f8620])).
fof(f13928,plain,(
  ( ! [X8,X7] : (perp(X7,X8,X8,X7) | ~midp(X8,X8,X7)) ) | ~spl32_192),
  inference(subsumption_resolution,[],[f13897,f8620])).
fof(f13897,plain,(
  ( ! [X8,X7] : (perp(X7,X8,X8,X7) | ~midp(X7,X7,X8) | ~midp(X8,X8,X7)) )),
  inference(resolution,[],[f1509,f1507])).
fof(f1507,plain,(
  ( ! [X6,X7,X5] : (cong(X7,X5,X5,X6) | ~midp(X5,X6,X7)) )),
  inference(resolution,[],[f527,f382])).
fof(f1509,plain,(
  ( ! [X12,X13,X11] : (~cong(X11,X13,X12,X13) | perp(X11,X12,X13,X11) | ~midp(X11,X11,X12)) )),
  inference(resolution,[],[f527,f415])).
fof(f415,plain,(
  ( ! [X2,X0,X3,X1] : (~cong(X0,X3,X1,X3) | perp(X0,X1,X2,X3) | ~cong(X0,X2,X1,X2)) )),
  inference(cnf_transformation,[],[f251])).
fof(f251,plain,(
  ! [X0,X1,X2,X3] : (perp(X0,X1,X2,X3) | ~cong(X0,X3,X1,X3) | ~cong(X0,X2,X1,X2))),
  inference(flattening,[],[f250])).
fof(f250,plain,(
  ! [X0,X1,X2,X3] : (perp(X0,X1,X2,X3) | (~cong(X0,X3,X1,X3) | ~cong(X0,X2,X1,X2)))),
  inference(ennf_transformation,[],[f135])).
fof(f135,plain,(
  ! [X0,X1,X2,X3] : ((cong(X0,X3,X1,X3) & cong(X0,X2,X1,X2)) => perp(X0,X1,X2,X3))),
  inference(rectify,[],[f57])).
fof(f57,axiom,(
  ! [X0,X1,X8,X9] : ((cong(X0,X9,X1,X9) & cong(X0,X8,X1,X8)) => perp(X0,X1,X8,X9))),
  file('./6_GDD_FULL/81-109/97.p',ruleD56)).
fof(f21198,plain,(
  ( ! [X116,X114,X115] : (perp(X115,X114,X116,X115) | ~cong(X115,X116,X114,X116)) ) | (~spl32_192 | ~spl32_240)),
  inference(subsumption_resolution,[],[f21181,f8496])).
fof(f21181,plain,(
  ( ! [X116,X114,X115] : (~midp(X114,X115,X115) | perp(X115,X114,X116,X115) | ~cong(X115,X116,X114,X116)) ) | ~spl32_240),
  inference(resolution,[],[f2747,f415])).
fof(f2747,plain,(
  ( ! [X14,X15,X13] : (cong(X14,X15,X13,X14) | ~midp(X13,X14,X15)) ) | ~spl32_240),
  inference(avatar_component_clause,[],[f2746])).
fof(f2746,plain,(
  spl32_240 <=> ! [X13,X15,X14] : (~midp(X13,X14,X15) | cong(X14,X15,X13,X14))),
  introduced(avatar_definition,[new_symbols(naming,[spl32_240])])).
fof(f21099,plain,(
  spl32_240 | ~spl32_84),
  inference(avatar_split_clause,[],[f3692,f1109,f2746])).
fof(f1109,plain,(
  spl32_84 <=> ! [X9,X11,X10] : (cong(X9,X10,X10,X11) | ~midp(X9,X10,X11))),
  introduced(avatar_definition,[new_symbols(naming,[spl32_84])])).
fof(f3692,plain,(
  ( ! [X10,X11,X9] : (~midp(X9,X10,X11) | cong(X10,X11,X9,X10)) ) | ~spl32_84),
  inference(resolution,[],[f1110,f382])).
fof(f1110,plain,(
  ( ! [X10,X11,X9] : (cong(X9,X10,X10,X11) | ~midp(X9,X10,X11)) ) | ~spl32_84),
  inference(avatar_component_clause,[],[f1109])).
fof(f21097,plain,(
  ~spl32_192 | ~spl32_240 | ~spl32_860),
  inference(avatar_contradiction_clause,[],[f21096])).
fof(f21096,plain,(
  $false | (~spl32_192 | ~spl32_240 | ~spl32_860)),
  inference(resolution,[],[f13931,f20683])).
fof(f20683,plain,(
  ( ! [X30] : (~perp(X30,sK20,X30,sK21)) ) | ~spl32_860),
  inference(avatar_component_clause,[],[f20682])).
fof(f20682,plain,(
  spl32_860 <=> ! [X30] : ~perp(X30,sK20,X30,sK21)),
  introduced(avatar_definition,[new_symbols(naming,[spl32_860])])).
fof(f13931,plain,(
  ( ! [X12,X11] : (perp(X11,X12,X11,X11)) ) | (~spl32_192 | ~spl32_240)),
  inference(subsumption_resolution,[],[f13930,f8496])).
fof(f13930,plain,(
  ( ! [X12,X11] : (perp(X11,X12,X11,X11) | ~midp(X12,X11,X11)) ) | (~spl32_192 | ~spl32_240)),
  inference(subsumption_resolution,[],[f13899,f8620])).
fof(f13899,plain,(
  ( ! [X12,X11] : (perp(X11,X12,X11,X11) | ~midp(X11,X11,X12) | ~midp(X12,X11,X11)) ) | ~spl32_240),
  inference(resolution,[],[f1509,f2747])).
fof(f20792,plain,(
  spl32_860 | spl32_162 | ~spl32_192),
  inference(avatar_split_clause,[],[f13765,f2076,f1834,f20682])).
fof(f13765,plain,(
  ( ! [X30,X29] : (~perp(sK23,sK20,sK20,X29) | ~perp(X30,sK20,X30,sK21) | perp(sK20,X29,sK20,sK21)) ) | ~spl32_192),
  inference(resolution,[],[f12718,f433])).
fof(f433,plain,(
  ( ! [X6,X4,X2,X0,X7,X5,X3,X1] : (~eqangle(X0,X1,X2,X3,X4,X5,X6,X7) | ~perp(X4,X5,X6,X7) | perp(X0,X1,X2,X3)) )),
  inference(cnf_transformation,[],[f281])).
fof(f281,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6,X7] : (perp(X0,X1,X2,X3) | ~perp(X4,X5,X6,X7) | ~eqangle(X0,X1,X2,X3,X4,X5,X6,X7))),
  inference(flattening,[],[f280])).
fof(f280,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6,X7] : (perp(X0,X1,X2,X3) | (~perp(X4,X5,X6,X7) | ~eqangle(X0,X1,X2,X3,X4,X5,X6,X7)))),
  inference(ennf_transformation,[],[f147])).
fof(f147,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6,X7] : ((perp(X4,X5,X6,X7) & eqangle(X0,X1,X2,X3,X4,X5,X6,X7)) => perp(X0,X1,X2,X3))),
  inference(rectify,[],[f75])).
fof(f75,axiom,(
  ! [X0,X1,X2,X3,X8,X9,X10,X11] : ((perp(X8,X9,X10,X11) & eqangle(X0,X1,X2,X3,X8,X9,X10,X11)) => perp(X0,X1,X2,X3))),
  file('./6_GDD_FULL/81-109/97.p',ruleD74)).
fof(f12718,plain,(
  ( ! [X151,X150] : (eqangle(sK20,X150,sK20,sK21,X151,sK20,X151,sK21) | ~perp(sK23,sK20,sK20,X150)) ) | ~spl32_192),
  inference(subsumption_resolution,[],[f4727,f8373])).
fof(f8373,plain,(
  ( ! [X12,X10,X11,X9] : (cyclic(X9,X10,X11,X12)) ) | ~spl32_192),
  inference(subsumption_resolution,[],[f8351,f8334])).
fof(f8351,plain,(
  ( ! [X12,X10,X11,X9] : (cyclic(X9,X10,X11,X12) | ~cyclic(X9,X9,X10,X11)) ) | ~spl32_192),
  inference(resolution,[],[f8334,f375])).
fof(f375,plain,(
  ( ! [X4,X2,X0,X3,X1] : (~cyclic(X0,X1,X2,X4) | cyclic(X1,X2,X3,X4) | ~cyclic(X0,X1,X2,X3)) )),
  inference(cnf_transformation,[],[f191])).
fof(f191,plain,(
  ! [X0,X1,X2,X3,X4] : (cyclic(X1,X2,X3,X4) | ~cyclic(X0,X1,X2,X4) | ~cyclic(X0,X1,X2,X3))),
  inference(flattening,[],[f190])).
fof(f190,plain,(
  ! [X0,X1,X2,X3,X4] : (cyclic(X1,X2,X3,X4) | (~cyclic(X0,X1,X2,X4) | ~cyclic(X0,X1,X2,X3)))),
  inference(ennf_transformation,[],[f17])).
fof(f17,axiom,(
  ! [X0,X1,X2,X3,X4] : ((cyclic(X0,X1,X2,X4) & cyclic(X0,X1,X2,X3)) => cyclic(X1,X2,X3,X4))),
  file('./6_GDD_FULL/81-109/97.p',ruleD17)).
fof(f4727,plain,(
  ( ! [X151,X150] : (eqangle(sK20,X150,sK20,sK21,X151,sK20,X151,sK21) | ~cyclic(sK20,sK21,sK22,X151) | ~perp(sK23,sK20,sK20,X150)) )),
  inference(resolution,[],[f1276,f1023])).
fof(f1023,plain,(
  ( ! [X0] : (eqangle(sK20,X0,sK20,sK21,sK22,sK20,sK22,sK21) | ~perp(sK23,sK20,sK20,X0)) )),
  inference(resolution,[],[f407,f472])).
fof(f472,plain,(
  circle(sK23,sK20,sK21,sK22)),
  inference(cnf_transformation,[],[f358])).
fof(f407,plain,(
  ( ! [X4,X2,X0,X3,X1] : (~circle(X3,X0,X1,X2) | ~perp(X3,X0,X0,X4) | eqangle(X0,X4,X0,X1,X2,X0,X2,X1)) )),
  inference(cnf_transformation,[],[f235])).
fof(f235,plain,(
  ! [X0,X1,X2,X3,X4] : (eqangle(X0,X4,X0,X1,X2,X0,X2,X1) | ~perp(X3,X0,X0,X4) | ~circle(X3,X0,X1,X2))),
  inference(flattening,[],[f234])).
fof(f234,plain,(
  ! [X0,X1,X2,X3,X4] : (eqangle(X0,X4,X0,X1,X2,X0,X2,X1) | (~perp(X3,X0,X0,X4) | ~circle(X3,X0,X1,X2)))),
  inference(ennf_transformation,[],[f128])).
fof(f128,plain,(
  ! [X0,X1,X2,X3,X4] : ((perp(X3,X0,X0,X4) & circle(X3,X0,X1,X2)) => eqangle(X0,X4,X0,X1,X2,X0,X2,X1))),
  inference(rectify,[],[f49])).
fof(f49,axiom,(
  ! [X0,X1,X2,X7,X15] : ((perp(X7,X0,X0,X15) & circle(X7,X0,X1,X2)) => eqangle(X0,X15,X0,X1,X2,X0,X2,X1))),
  file('./6_GDD_FULL/81-109/97.p',ruleD48)).
fof(f1276,plain,(
  ( ! [X28,X26,X24,X23,X29,X27,X25,X22] : (~eqangle(X22,X23,X24,X25,X29,X27,X29,X28) | eqangle(X22,X23,X24,X25,X26,X27,X26,X28) | ~cyclic(X27,X28,X29,X26)) )),
  inference(resolution,[],[f380,f399])).
fof(f380,plain,(
  ( ! [X6,X4,X2,X0,X10,X8,X7,X5,X3,X1,X11,X9] : (~eqangle(X4,X5,X6,X7,X8,X9,X10,X11) | eqangle(X0,X1,X2,X3,X8,X9,X10,X11) | ~eqangle(X0,X1,X2,X3,X4,X5,X6,X7)) )),
  inference(cnf_transformation,[],[f197])).
fof(f197,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11] : (eqangle(X0,X1,X2,X3,X8,X9,X10,X11) | ~eqangle(X4,X5,X6,X7,X8,X9,X10,X11) | ~eqangle(X0,X1,X2,X3,X4,X5,X6,X7))),
  inference(flattening,[],[f196])).
fof(f196,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11] : (eqangle(X0,X1,X2,X3,X8,X9,X10,X11) | (~eqangle(X4,X5,X6,X7,X8,X9,X10,X11) | ~eqangle(X0,X1,X2,X3,X4,X5,X6,X7)))),
  inference(ennf_transformation,[],[f104])).
fof(f104,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11] : ((eqangle(X4,X5,X6,X7,X8,X9,X10,X11) & eqangle(X0,X1,X2,X3,X4,X5,X6,X7)) => eqangle(X0,X1,X2,X3,X8,X9,X10,X11))),
  inference(rectify,[],[f22])).
fof(f22,axiom,(
  ! [X0,X1,X2,X3,X8,X9,X10,X11,X4,X5,X12,X13] : ((eqangle(X8,X9,X10,X11,X4,X5,X12,X13) & eqangle(X0,X1,X2,X3,X8,X9,X10,X11)) => eqangle(X0,X1,X2,X3,X4,X5,X12,X13))),
  file('./6_GDD_FULL/81-109/97.p',ruleD22)).
fof(f8881,plain,(
  ~spl32_192 | spl32_373),
  inference(avatar_contradiction_clause,[],[f8838])).
fof(f8838,plain,(
  $false | (~spl32_192 | ~spl32_373)),
  inference(resolution,[],[f8620,f4858])).
fof(f4858,plain,(
  ~midp(sK20,sK20,sK21) | ~spl32_373),
  inference(resolution,[],[f4827,f427])).
fof(f4827,plain,(
  ~cong(sK20,sK20,sK20,sK21) | ~spl32_373),
  inference(avatar_component_clause,[],[f4826])).
fof(f4826,plain,(
  spl32_373 <=> ~cong(sK20,sK20,sK20,sK21)),
  introduced(avatar_definition,[new_symbols(naming,[spl32_373])])).
fof(f3580,plain,(
  ~spl32_194),
  inference(avatar_contradiction_clause,[],[f3575])).
fof(f3575,plain,(
  $false | ~spl32_194),
  inference(resolution,[],[f2080,f1559])).
fof(f2080,plain,(
  ( ! [X0,X1] : (~para(X0,X1,X0,X1)) ) | ~spl32_194),
  inference(avatar_component_clause,[],[f2079])).
fof(f2079,plain,(
  spl32_194 <=> ! [X1,X0] : ~para(X0,X1,X0,X1)),
  introduced(avatar_definition,[new_symbols(naming,[spl32_194])])).
fof(f3206,plain,(
  ~spl32_86 | ~spl32_192),
  inference(avatar_contradiction_clause,[],[f3199])).
fof(f3199,plain,(
  $false | (~spl32_86 | ~spl32_192)),
  inference(resolution,[],[f3132,f472])).
fof(f3132,plain,(
  ( ! [X12,X10,X11,X9] : (~circle(X11,X12,X10,X9)) ) | (~spl32_86 | ~spl32_192)),
  inference(subsumption_resolution,[],[f3131,f2738])).
fof(f2738,plain,(
  ( ! [X0,X1] : (~midp(X0,X1,X0)) ) | ~spl32_86),
  inference(subsumption_resolution,[],[f2733,f369])).
fof(f2733,plain,(
  ( ! [X0,X1] : (~midp(X0,X0,X1) | ~midp(X0,X1,X0)) ) | ~spl32_86),
  inference(resolution,[],[f1113,f530])).
fof(f1113,plain,(
  ( ! [X6,X8,X7] : (~cong(X6,X7,X7,X8) | ~midp(X6,X7,X8)) ) | ~spl32_86),
  inference(avatar_component_clause,[],[f1112])).
fof(f1112,plain,(
  spl32_86 <=> ! [X7,X8,X6] : (~cong(X6,X7,X7,X8) | ~midp(X6,X7,X8))),
  introduced(avatar_definition,[new_symbols(naming,[spl32_86])])).
fof(f3131,plain,(
  ( ! [X12,X10,X11,X9] : (~circle(X11,X12,X10,X9) | midp(X9,X10,X9)) ) | ~spl32_192),
  inference(subsumption_resolution,[],[f3124,f2279])).
fof(f2279,plain,(
  ( ! [X4,X3] : (coll(X3,X4,X3)) ) | ~spl32_192),
  inference(resolution,[],[f2266,f360])).
fof(f2266,plain,(
  ( ! [X0,X1] : (coll(X0,X1,X1)) ) | ~spl32_192),
  inference(resolution,[],[f2077,f425])).
fof(f3124,plain,(
  ( ! [X12,X10,X11,X9] : (~coll(X9,X10,X9) | ~circle(X11,X12,X10,X9) | midp(X9,X10,X9)) ) | ~spl32_192),
  inference(resolution,[],[f1181,f2656])).
fof(f2656,plain,(
  ( ! [X12,X10,X11,X9] : (cyclic(X9,X10,X11,X12)) ) | ~spl32_192),
  inference(subsumption_resolution,[],[f2649,f2239])).
fof(f2239,plain,(
  ( ! [X70,X71,X69] : (cyclic(X70,X70,X69,X71)) ) | ~spl32_192),
  inference(subsumption_resolution,[],[f2238,f2237])).
fof(f2237,plain,(
  ( ! [X68,X66,X67] : (~coll(X66,X68,X67) | cyclic(X67,X67,X66,X68)) ) | ~spl32_192),
  inference(subsumption_resolution,[],[f2218,f2077])).
fof(f2238,plain,(
  ( ! [X70,X71,X69] : (coll(X69,X71,X70) | cyclic(X70,X70,X69,X71)) ) | ~spl32_192),
  inference(subsumption_resolution,[],[f2219,f2077])).
fof(f2649,plain,(
  ( ! [X12,X10,X11,X9] : (cyclic(X9,X10,X11,X12) | ~cyclic(X9,X9,X10,X11)) ) | ~spl32_192),
  inference(resolution,[],[f2239,f375])).
fof(f1181,plain,(
  ( ! [X6,X8,X7,X9] : (~cyclic(X7,X6,X9,X8) | ~coll(X6,X7,X6) | ~circle(X8,X9,X7,X6) | midp(X6,X7,X6)) )),
  inference(resolution,[],[f410,f399])).
fof(f410,plain,(
  ( ! [X4,X2,X0,X3,X1] : (~eqangle(X0,X1,X0,X2,X3,X1,X3,X4) | midp(X4,X1,X2) | ~coll(X4,X1,X2) | ~circle(X3,X0,X1,X2)) )),
  inference(cnf_transformation,[],[f241])).
fof(f241,plain,(
  ! [X0,X1,X2,X3,X4] : (midp(X4,X1,X2) | ~eqangle(X0,X1,X0,X2,X3,X1,X3,X4) | ~coll(X4,X1,X2) | ~circle(X3,X0,X1,X2))),
  inference(flattening,[],[f240])).
fof(f240,plain,(
  ! [X0,X1,X2,X3,X4] : (midp(X4,X1,X2) | (~eqangle(X0,X1,X0,X2,X3,X1,X3,X4) | ~coll(X4,X1,X2) | ~circle(X3,X0,X1,X2)))),
  inference(ennf_transformation,[],[f131])).
fof(f131,plain,(
  ! [X0,X1,X2,X3,X4] : ((eqangle(X0,X1,X0,X2,X3,X1,X3,X4) & coll(X4,X1,X2) & circle(X3,X0,X1,X2)) => midp(X4,X1,X2))),
  inference(rectify,[],[f52])).
fof(f52,axiom,(
  ! [X0,X1,X2,X7,X6] : ((eqangle(X0,X1,X0,X2,X7,X1,X7,X6) & coll(X6,X1,X2) & circle(X7,X0,X1,X2)) => midp(X6,X1,X2))),
  file('./6_GDD_FULL/81-109/97.p',ruleD51)).
fof(f2732,plain,(
  spl32_236 | spl32_86),
  inference(avatar_split_clause,[],[f2719,f1112,f2726])).
fof(f2719,plain,(
  ( ! [X14,X12,X10,X15,X13,X11] : (~midp(X10,X11,X12) | ~midp(X13,X14,X15) | ~cong(X10,X11,X11,X12) | cong(X14,X13,X14,X15)) )),
  inference(resolution,[],[f845,f434])).
fof(f434,plain,(
  ( ! [X6,X4,X2,X0,X7,X5,X3,X1] : (~eqratio(X0,X1,X2,X3,X4,X5,X6,X7) | ~cong(X4,X5,X6,X7) | cong(X0,X1,X2,X3)) )),
  inference(cnf_transformation,[],[f283])).
fof(f283,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6,X7] : (cong(X0,X1,X2,X3) | ~cong(X4,X5,X6,X7) | ~eqratio(X0,X1,X2,X3,X4,X5,X6,X7))),
  inference(flattening,[],[f282])).
fof(f282,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6,X7] : (cong(X0,X1,X2,X3) | (~cong(X4,X5,X6,X7) | ~eqratio(X0,X1,X2,X3,X4,X5,X6,X7)))),
  inference(ennf_transformation,[],[f148])).
fof(f148,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6,X7] : ((cong(X4,X5,X6,X7) & eqratio(X0,X1,X2,X3,X4,X5,X6,X7)) => cong(X0,X1,X2,X3))),
  inference(rectify,[],[f76])).
fof(f76,axiom,(
  ! [X0,X1,X2,X3,X8,X9,X10,X11] : ((cong(X8,X9,X10,X11) & eqratio(X0,X1,X2,X3,X8,X9,X10,X11)) => cong(X0,X1,X2,X3))),
  file('./6_GDD_FULL/81-109/97.p',ruleD75)).
fof(f845,plain,(
  ( ! [X6,X10,X8,X7,X11,X9] : (eqratio(X6,X7,X6,X8,X9,X10,X10,X11) | ~midp(X9,X10,X11) | ~midp(X7,X6,X8)) )),
  inference(resolution,[],[f384,f429])).
fof(f429,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (eqratio(X4,X0,X0,X1,X5,X2,X2,X3) | ~midp(X5,X2,X3) | ~midp(X4,X0,X1)) )),
  inference(cnf_transformation,[],[f273])).
fof(f273,plain,(
  ! [X0,X1,X2,X3,X4,X5] : (eqratio(X4,X0,X0,X1,X5,X2,X2,X3) | ~midp(X5,X2,X3) | ~midp(X4,X0,X1))),
  inference(flattening,[],[f272])).
fof(f272,plain,(
  ! [X0,X1,X2,X3,X4,X5] : (eqratio(X4,X0,X0,X1,X5,X2,X2,X3) | (~midp(X5,X2,X3) | ~midp(X4,X0,X1)))),
  inference(ennf_transformation,[],[f145])).
fof(f145,plain,(
  ! [X0,X1,X2,X3,X4,X5] : ((midp(X5,X2,X3) & midp(X4,X0,X1)) => eqratio(X4,X0,X0,X1,X5,X2,X2,X3))),
  inference(rectify,[],[f71])).
fof(f71,axiom,(
  ! [X0,X1,X2,X3,X6,X16] : ((midp(X16,X2,X3) & midp(X6,X0,X1)) => eqratio(X6,X0,X0,X1,X16,X2,X2,X3))),
  file('./6_GDD_FULL/81-109/97.p',ruleD70)).
fof(f384,plain,(
  ( ! [X6,X4,X2,X0,X7,X5,X3,X1] : (~eqratio(X0,X1,X2,X3,X4,X5,X6,X7) | eqratio(X1,X0,X2,X3,X4,X5,X6,X7)) )),
  inference(cnf_transformation,[],[f202])).
fof(f202,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6,X7] : (eqratio(X1,X0,X2,X3,X4,X5,X6,X7) | ~eqratio(X0,X1,X2,X3,X4,X5,X6,X7))),
  inference(ennf_transformation,[],[f105])).
fof(f105,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6,X7] : (eqratio(X0,X1,X2,X3,X4,X5,X6,X7) => eqratio(X1,X0,X2,X3,X4,X5,X6,X7))),
  inference(rectify,[],[f26])).
fof(f26,axiom,(
  ! [X0,X1,X2,X3,X8,X9,X10,X11] : (eqratio(X0,X1,X2,X3,X8,X9,X10,X11) => eqratio(X1,X0,X2,X3,X8,X9,X10,X11))),
  file('./6_GDD_FULL/81-109/97.p',ruleD26)).
fof(f2343,plain,(
  spl32_3 | ~spl32_192),
  inference(avatar_contradiction_clause,[],[f2330])).
fof(f2330,plain,(
  $false | (~spl32_3 | ~spl32_192)),
  inference(resolution,[],[f2321,f622])).
fof(f622,plain,(
  ~coll(sK23,sK20,sK22) | ~spl32_3),
  inference(avatar_component_clause,[],[f621])).
fof(f621,plain,(
  spl32_3 <=> ~coll(sK23,sK20,sK22)),
  introduced(avatar_definition,[new_symbols(naming,[spl32_3])])).
fof(f2321,plain,(
  ( ! [X2,X0,X1] : (coll(X0,X1,X2)) ) | ~spl32_192),
  inference(subsumption_resolution,[],[f2318,f2300])).
fof(f2300,plain,(
  ( ! [X6,X5] : (coll(X5,X5,X6)) ) | ~spl32_192),
  inference(resolution,[],[f2279,f359])).
fof(f2318,plain,(
  ( ! [X2,X0,X1] : (coll(X0,X1,X2) | ~coll(X2,X2,X0)) ) | ~spl32_192),
  inference(resolution,[],[f2300,f361])).
fof(f2081,plain,(
  spl32_192 | spl32_194),
  inference(avatar_split_clause,[],[f2052,f2079,f2076])).
fof(f2052,plain,(
  ( ! [X2,X0,X3,X1] : (~para(X0,X1,X0,X1) | para(X2,X3,X2,X3)) )),
  inference(resolution,[],[f794,f397])).
fof(f794,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (eqangle(X0,X1,X2,X3,X0,X1,X4,X5) | ~para(X2,X3,X4,X5)) )),
  inference(resolution,[],[f377,f398])).
fof(f377,plain,(
  ( ! [X6,X4,X2,X0,X7,X5,X3,X1] : (~eqangle(X0,X1,X2,X3,X4,X5,X6,X7) | eqangle(X2,X3,X0,X1,X6,X7,X4,X5)) )),
  inference(cnf_transformation,[],[f193])).
fof(f193,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6,X7] : (eqangle(X2,X3,X0,X1,X6,X7,X4,X5) | ~eqangle(X0,X1,X2,X3,X4,X5,X6,X7))),
  inference(ennf_transformation,[],[f101])).
fof(f101,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6,X7] : (eqangle(X0,X1,X2,X3,X4,X5,X6,X7) => eqangle(X2,X3,X0,X1,X6,X7,X4,X5))),
  inference(rectify,[],[f19])).
fof(f19,axiom,(
  ! [X0,X1,X2,X3,X8,X9,X10,X11] : (eqangle(X0,X1,X2,X3,X8,X9,X10,X11) => eqangle(X2,X3,X0,X1,X10,X11,X8,X9))),
  file('./6_GDD_FULL/81-109/97.p',ruleD19)).
fof(f1114,plain,(
  spl32_84 | spl32_86),
  inference(avatar_split_clause,[],[f1107,f1112,f1109])).
fof(f1107,plain,(
  ( ! [X6,X10,X8,X7,X11,X9] : (~cong(X6,X7,X7,X8) | cong(X9,X10,X10,X11) | ~midp(X6,X7,X8) | ~midp(X9,X10,X11)) )),
  inference(resolution,[],[f434,f429])).
fof(f623,plain,(
  spl32_0 | ~spl32_3),
  inference(avatar_split_clause,[],[f608,f621,f615])).
fof(f608,plain,(
  ~coll(sK23,sK20,sK22) | perp(sK20,sK21,sK21,sK22)),
  inference(resolution,[],[f412,f472])).
% SZS output end Proof for 97
% ------------------------------
% Version: Vampire 4.4.0 (commit 7916d27 on 2019-08-23 08:50:16 +0100)
% Termination reason: Refutation

% Memory used [KB]: 24818
% Time elapsed: 0.487 s
% ------------------------------
% ------------------------------
