% Refutation found. Thanks to Tanya!
% SZS status Theorem for M010-26
% SZS output start Proof for M010-26
fof(f569,plain,(
  $false),
  inference(resolution,[],[f566,f481])).
fof(f481,plain,(
  ~cong(sK23,sK22,sK21,sK20)),
  inference(cnf_transformation,[],[f357])).
fof(f357,plain,(
  ~cong(sK23,sK22,sK21,sK20) & ! [X4,X5,X6,X7,X8,X9,X10,X11,X12] : (coll(X12,X8,X9) & coll(X12,X5,X6) & coll(X11,X4,X7) & coll(X11,X8,X9) & coll(X10,X8,X9) & coll(X10,X4,X5) & coll(X9,X6,X7) & midp(X8,X5,X7) & para(X4,X5,X6,X7) & para(X5,X6,X4,X7))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK20,sK21,sK22,sK23])],[f319,f356])).
fof(f356,plain,(
  ? [X0,X1,X2,X3] : (~cong(X3,X2,X1,X0) & ! [X4,X5,X6,X7,X8,X9,X10,X11,X12] : (coll(X12,X8,X9) & coll(X12,X5,X6) & coll(X11,X4,X7) & coll(X11,X8,X9) & coll(X10,X8,X9) & coll(X10,X4,X5) & coll(X9,X6,X7) & midp(X8,X5,X7) & para(X4,X5,X6,X7) & para(X5,X6,X4,X7))) => (~cong(sK23,sK22,sK21,sK20) & ! [X4,X5,X6,X7,X8,X9,X10,X11,X12] : (coll(X12,X8,X9) & coll(X12,X5,X6) & coll(X11,X4,X7) & coll(X11,X8,X9) & coll(X10,X8,X9) & coll(X10,X4,X5) & coll(X9,X6,X7) & midp(X8,X5,X7) & para(X4,X5,X6,X7) & para(X5,X6,X4,X7)))),
  introduced(choice_axiom,[])).
fof(f319,plain,(
  ? [X0,X1,X2,X3] : (~cong(X3,X2,X1,X0) & ! [X4,X5,X6,X7,X8,X9,X10,X11,X12] : (coll(X12,X8,X9) & coll(X12,X5,X6) & coll(X11,X4,X7) & coll(X11,X8,X9) & coll(X10,X8,X9) & coll(X10,X4,X5) & coll(X9,X6,X7) & midp(X8,X5,X7) & para(X4,X5,X6,X7) & para(X5,X6,X4,X7)))),
  inference(ennf_transformation,[],[f167])).
fof(f167,plain,(
  ~! [X0,X1,X2,X3] : (! [X4,X5,X6,X7,X8,X9,X10,X11,X12] : (coll(X12,X8,X9) & coll(X12,X5,X6) & coll(X11,X4,X7) & coll(X11,X8,X9) & coll(X10,X8,X9) & coll(X10,X4,X5) & coll(X9,X6,X7) & midp(X8,X5,X7) & para(X4,X5,X6,X7) & para(X5,X6,X4,X7)) => cong(X3,X2,X1,X0))),
  inference(rectify,[],[f96])).
fof(f96,negated_conjecture,(
  ~! [X18,X12,X13,X5] : (! [X0,X1,X2,X3,X4,X5,X12,X13,X18] : (coll(X18,X4,X5) & coll(X18,X1,X2) & coll(X13,X0,X3) & coll(X13,X4,X5) & coll(X12,X4,X5) & coll(X12,X0,X1) & coll(X5,X2,X3) & midp(X4,X1,X3) & para(X0,X1,X2,X3) & para(X1,X2,X0,X3)) => cong(X5,X13,X12,X18))),
  inference(negated_conjecture,[],[f95])).
fof(f95,conjecture,(
  ( ! [X12,X5,X13,X18] : (! [X0,X1,X2,X3,X4,X5,X12,X13,X18] : (coll(X18,X4,X5) & coll(X18,X1,X2) & coll(X13,X0,X3) & coll(X13,X4,X5) & coll(X12,X4,X5) & coll(X12,X0,X1) & coll(X5,X2,X3) & midp(X4,X1,X3) & para(X0,X1,X2,X3) & para(X1,X2,X0,X3)) => cong(X5,X13,X12,X18)) )),
  file('./Other/gao/M/M010-26.p',exemploOthergaoMM01026)).
fof(f566,plain,(
  ( ! [X10,X8,X7,X9] : (cong(X7,X8,X9,X10)) )),
  inference(resolution,[],[f555,f490])).
fof(f490,plain,(
  ( ! [X2,X0,X1] : (cong(X0,X1,X2,X0)) )),
  inference(subsumption_resolution,[],[f489,f473])).
fof(f473,plain,(
  ( ! [X8,X7,X5] : (midp(X8,X5,X7)) )),
  inference(cnf_transformation,[],[f357])).
fof(f489,plain,(
  ( ! [X2,X0,X1] : (cong(X0,X1,X2,X0) | ~midp(X0,X1,X2)) )),
  inference(resolution,[],[f380,f426])).
fof(f426,plain,(
  ( ! [X2,X0,X1] : (cong(X0,X1,X0,X2) | ~midp(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f270])).
fof(f270,plain,(
  ! [X0,X1,X2] : (cong(X0,X1,X0,X2) | ~midp(X0,X1,X2))),
  inference(ennf_transformation,[],[f69])).
fof(f69,axiom,(
  ! [X0,X1,X2] : (midp(X0,X1,X2) => cong(X0,X1,X0,X2))),
  file('./Other/gao/M/M010-26.p',ruleD68)).
fof(f380,plain,(
  ( ! [X2,X0,X3,X1] : (~cong(X0,X1,X2,X3) | cong(X0,X1,X3,X2)) )),
  inference(cnf_transformation,[],[f198])).
fof(f198,plain,(
  ! [X0,X1,X2,X3] : (cong(X0,X1,X3,X2) | ~cong(X0,X1,X2,X3))),
  inference(ennf_transformation,[],[f23])).
fof(f23,axiom,(
  ! [X0,X1,X2,X3] : (cong(X0,X1,X2,X3) => cong(X0,X1,X3,X2))),
  file('./Other/gao/M/M010-26.p',cong23)).
fof(f555,plain,(
  ( ! [X14,X12,X10,X13,X11] : (~cong(X10,X11,X13,X14) | cong(X10,X11,X12,X13)) )),
  inference(resolution,[],[f382,f490])).
fof(f382,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (~cong(X2,X3,X4,X5) | cong(X0,X1,X4,X5) | ~cong(X0,X1,X2,X3)) )),
  inference(cnf_transformation,[],[f201])).
fof(f201,plain,(
  ! [X0,X1,X2,X3,X4,X5] : (cong(X0,X1,X4,X5) | ~cong(X2,X3,X4,X5) | ~cong(X0,X1,X2,X3))),
  inference(flattening,[],[f200])).
fof(f200,plain,(
  ! [X0,X1,X2,X3,X4,X5] : (cong(X0,X1,X4,X5) | (~cong(X2,X3,X4,X5) | ~cong(X0,X1,X2,X3)))),
  inference(ennf_transformation,[],[f25])).
fof(f25,axiom,(
  ! [X0,X1,X2,X3,X4,X5] : ((cong(X2,X3,X4,X5) & cong(X0,X1,X2,X3)) => cong(X0,X1,X4,X5))),
  file('./Other/gao/M/M010-26.p',ruleD25)).
% SZS output end Proof for M010-26
% ------------------------------
% Version: Vampire 4.4.0 (commit 7916d27 on 2019-08-23 08:50:16 +0100)
% Termination reason: Refutation

% Memory used [KB]: 5117
% Time elapsed: 0.028 s
% ------------------------------
% ------------------------------
