% Refutation found. Thanks to Tanya!
% SZS status Theorem for M021-59
% SZS output start Proof for M021-59
9. ! [X0,X1,X2,X3,X4,X5] : ((perp(X2,X3,X4,X5) & perp(X0,X1,X2,X3)) => para(X0,X1,X4,X5)) [input]
18. ! [X0,X1,X2,X3,X8,X9,X10,X11] : (eqangle(X0,X1,X2,X3,X8,X9,X10,X11) => eqangle(X1,X0,X2,X3,X8,X9,X10,X11)) [input]
19. ! [X0,X1,X2,X3,X8,X9,X10,X11] : (eqangle(X0,X1,X2,X3,X8,X9,X10,X11) => eqangle(X2,X3,X0,X1,X10,X11,X8,X9)) [input]
23. ! [X0,X1,X2,X3] : (cong(X0,X1,X2,X3) => cong(X0,X1,X3,X2)) [input]
25. ! [X0,X1,X2,X3,X4,X5] : ((cong(X2,X3,X4,X5) & cong(X0,X1,X2,X3)) => cong(X0,X1,X4,X5)) [input]
40. ! [X0,X1,X2,X3,X8,X9] : (para(X0,X1,X2,X3) => eqangle(X0,X1,X8,X9,X2,X3,X8,X9)) [input]
51. ! [X0,X1,X2,X7,X6] : ((eqangle(X0,X1,X0,X2,X7,X1,X7,X6) & coll(X6,X1,X2) & circle(X7,X0,X1,X2)) => midp(X6,X1,X2)) [input]
68. ! [X0,X1,X2] : (midp(X0,X1,X2) => cong(X0,X1,X0,X2)) [input]
94. ! [X0,X1,X2,X3,X4,X5,X12,X18,X19,X20] : (coll(X12,X2,X3) & perp(X12,X4,X2,X3) & coll(X5,X2,X3) & perp(X5,X1,X2,X3) & circle(X0,X1,X4,X20) & coll(X4,X1,X0) & circle(X0,X1,X3,X19) & circle(X0,X1,X2,X18)) => cong(X5,X2,X3,X12) [input]
95. ~! [X12,X3,X2,X5] : (! [X0,X1,X2,X3,X4,X5,X12,X18,X19,X20] : (coll(X12,X2,X3) & perp(X12,X4,X2,X3) & coll(X5,X2,X3) & perp(X5,X1,X2,X3) & circle(X0,X1,X4,X20) & coll(X4,X1,X0) & circle(X0,X1,X3,X19) & circle(X0,X1,X2,X18)) => cong(X5,X2,X3,X12)) [negated conjecture 94]
99. ! [X0,X1,X2,X3,X4,X5,X6,X7] : (eqangle(X0,X1,X2,X3,X4,X5,X6,X7) => eqangle(X1,X0,X2,X3,X4,X5,X6,X7)) [rectify 18]
100. ! [X0,X1,X2,X3,X4,X5,X6,X7] : (eqangle(X0,X1,X2,X3,X4,X5,X6,X7) => eqangle(X2,X3,X0,X1,X6,X7,X4,X5)) [rectify 19]
118. ! [X0,X1,X2,X3,X4,X5] : (para(X0,X1,X2,X3) => eqangle(X0,X1,X4,X5,X2,X3,X4,X5)) [rectify 40]
129. ! [X0,X1,X2,X3,X4] : ((eqangle(X0,X1,X0,X2,X3,X1,X3,X4) & coll(X4,X1,X2) & circle(X3,X0,X1,X2)) => midp(X4,X1,X2)) [rectify 51]
165. ~! [X0,X1,X2,X3] : (! [X4,X5,X6,X7,X8,X9,X10,X11,X12,X13] : (coll(X10,X6,X7) & perp(X10,X8,X6,X7) & coll(X9,X6,X7) & perp(X9,X5,X6,X7) & circle(X4,X5,X8,X13) & coll(X8,X5,X4) & circle(X4,X5,X7,X12) & circle(X4,X5,X6,X11)) => cong(X3,X2,X1,X0)) [rectify 95]
176. ! [X0,X1,X2,X3,X4,X5] : (para(X0,X1,X4,X5) | (~perp(X2,X3,X4,X5) | ~perp(X0,X1,X2,X3))) [ennf transformation 9]
177. ! [X0,X1,X2,X3,X4,X5] : (para(X0,X1,X4,X5) | ~perp(X2,X3,X4,X5) | ~perp(X0,X1,X2,X3)) [flattening 176]
190. ! [X0,X1,X2,X3,X4,X5,X6,X7] : (eqangle(X1,X0,X2,X3,X4,X5,X6,X7) | ~eqangle(X0,X1,X2,X3,X4,X5,X6,X7)) [ennf transformation 99]
191. ! [X0,X1,X2,X3,X4,X5,X6,X7] : (eqangle(X2,X3,X0,X1,X6,X7,X4,X5) | ~eqangle(X0,X1,X2,X3,X4,X5,X6,X7)) [ennf transformation 100]
196. ! [X0,X1,X2,X3] : (cong(X0,X1,X3,X2) | ~cong(X0,X1,X2,X3)) [ennf transformation 23]
198. ! [X0,X1,X2,X3,X4,X5] : (cong(X0,X1,X4,X5) | (~cong(X2,X3,X4,X5) | ~cong(X0,X1,X2,X3))) [ennf transformation 25]
199. ! [X0,X1,X2,X3,X4,X5] : (cong(X0,X1,X4,X5) | ~cong(X2,X3,X4,X5) | ~cong(X0,X1,X2,X3)) [flattening 198]
217. ! [X0,X1,X2,X3,X4,X5] : (eqangle(X0,X1,X4,X5,X2,X3,X4,X5) | ~para(X0,X1,X2,X3)) [ennf transformation 118]
236. ! [X0,X1,X2,X3,X4] : (midp(X4,X1,X2) | (~eqangle(X0,X1,X0,X2,X3,X1,X3,X4) | ~coll(X4,X1,X2) | ~circle(X3,X0,X1,X2))) [ennf transformation 129]
237. ! [X0,X1,X2,X3,X4] : (midp(X4,X1,X2) | ~eqangle(X0,X1,X0,X2,X3,X1,X3,X4) | ~coll(X4,X1,X2) | ~circle(X3,X0,X1,X2)) [flattening 236]
266. ! [X0,X1,X2] : (cong(X0,X1,X0,X2) | ~midp(X0,X1,X2)) [ennf transformation 68]
315. ? [X0,X1,X2,X3] : (~cong(X3,X2,X1,X0) & ! [X4,X5,X6,X7,X8,X9,X10,X11,X12,X13] : (coll(X10,X6,X7) & perp(X10,X8,X6,X7) & coll(X9,X6,X7) & perp(X9,X5,X6,X7) & circle(X4,X5,X8,X13) & coll(X8,X5,X4) & circle(X4,X5,X7,X12) & circle(X4,X5,X6,X11))) [ennf transformation 165]
352. ? [X0,X1,X2,X3] : (~cong(X3,X2,X1,X0) & ! [X4,X5,X6,X7,X8,X9,X10,X11,X12,X13] : (coll(X10,X6,X7) & perp(X10,X8,X6,X7) & coll(X9,X6,X7) & perp(X9,X5,X6,X7) & circle(X4,X5,X8,X13) & coll(X8,X5,X4) & circle(X4,X5,X7,X12) & circle(X4,X5,X6,X11))) => (~cong(sK23,sK22,sK21,sK20) & ! [X4,X5,X6,X7,X8,X9,X10,X11,X12,X13] : (coll(X10,X6,X7) & perp(X10,X8,X6,X7) & coll(X9,X6,X7) & perp(X9,X5,X6,X7) & circle(X4,X5,X8,X13) & coll(X8,X5,X4) & circle(X4,X5,X7,X12) & circle(X4,X5,X6,X11))) [choice axiom]
353. ~cong(sK23,sK22,sK21,sK20) & ! [X4,X5,X6,X7,X8,X9,X10,X11,X12,X13] : (coll(X10,X6,X7) & perp(X10,X8,X6,X7) & coll(X9,X6,X7) & perp(X9,X5,X6,X7) & circle(X4,X5,X8,X13) & coll(X8,X5,X4) & circle(X4,X5,X7,X12) & circle(X4,X5,X6,X11)) [skolemisation 315,352]
362. ~perp(X2,X3,X4,X5) | para(X0,X1,X4,X5) | ~perp(X0,X1,X2,X3) [cnf transformation 177]
371. ~eqangle(X0,X1,X2,X3,X4,X5,X6,X7) | eqangle(X1,X0,X2,X3,X4,X5,X6,X7) [cnf transformation 190]
372. ~eqangle(X0,X1,X2,X3,X4,X5,X6,X7) | eqangle(X2,X3,X0,X1,X6,X7,X4,X5) [cnf transformation 191]
376. ~cong(X0,X1,X2,X3) | cong(X0,X1,X3,X2) [cnf transformation 196]
378. ~cong(X2,X3,X4,X5) | cong(X0,X1,X4,X5) | ~cong(X0,X1,X2,X3) [cnf transformation 199]
393. eqangle(X0,X1,X4,X5,X2,X3,X4,X5) | ~para(X0,X1,X2,X3) [cnf transformation 217]
405. ~eqangle(X0,X1,X0,X2,X3,X1,X3,X4) | midp(X4,X1,X2) | ~coll(X4,X1,X2) | ~circle(X3,X0,X1,X2) [cnf transformation 237]
422. cong(X0,X1,X0,X2) | ~midp(X0,X1,X2) [cnf transformation 266]
470. circle(X4,X5,X8,X13) [cnf transformation 353]
473. perp(X10,X8,X6,X7) [cnf transformation 353]
474. coll(X10,X6,X7) [cnf transformation 353]
475. ~cong(sK23,sK22,sK21,sK20) [cnf transformation 353]
480. cong(X0,X1,X2,X0) | ~midp(X0,X1,X2) [resolution 376,422]
532. para(X0,X1,X2,X3) | ~perp(X0,X1,X4,X5) [resolution 362,473]
534. para(X0,X1,X2,X3) [subsumption resolution 532,473]
562. ~cong(X5,X6,X7,X9) | cong(X5,X6,X7,X8) | ~midp(X7,X9,X8) [resolution 378,422]
584. cong(X12,X13,X14,X15) | ~midp(X14,X12,X15) | ~midp(X12,X13,X14) [resolution 562,480]
605. eqangle(X0,X1,X2,X3,X4,X5,X2,X3) | ~para(X1,X0,X4,X5) [resolution 371,393]
609. eqangle(X0,X1,X2,X3,X4,X5,X2,X3) [subsumption resolution 605,534]
612. eqangle(X0,X1,X2,X3,X0,X1,X4,X5) [resolution 372,609]
619. eqangle(X8,X9,X10,X11,X9,X8,X12,X13) [resolution 612,371]
628. eqangle(X8,X9,X10,X11,X12,X13,X11,X10) [resolution 619,372]
706. ~midp(sK21,sK23,sK20) | ~midp(sK23,sK22,sK21) [resolution 584,475]
719. 1 <=> ~midp(sK23,sK22,sK21) [avatar definition]
720. ~midp(sK23,sK22,sK21) <- (1) [avatar component clause 719]
725. 3 <=> ~midp(sK21,sK23,sK20) [avatar definition]
726. ~midp(sK21,sK23,sK20) <- (3) [avatar component clause 725]
727. ~1 | ~3 [avatar split clause 706,725,719]
1240. midp(X6,X7,X8) | ~coll(X6,X7,X8) | ~circle(X8,X6,X7,X8) [resolution 405,628]
1271. midp(X6,X7,X8) | ~circle(X8,X6,X7,X8) [subsumption resolution 1240,474]
1272. midp(X6,X7,X8) [subsumption resolution 1271,470]
1332. $false <- (1) [resolution 1272,720]
1333. 1 [avatar contradiction clause 1332]
1334. $false <- (3) [subsumption resolution 726,1272]
1335. 3 [avatar contradiction clause 1334]
1336. $false [avatar sat refutation 727,1333,1335]
% SZS output end Proof for M021-59
% ------------------------------
% Version: Vampire 4.4.0 (commit 7916d27 on 2019-08-23 08:50:16 +0100)
% Termination reason: Refutation

% Memory used [KB]: 5373
% Time elapsed: 0.036 s
% ------------------------------
% ------------------------------
