% Refutation found. Thanks to Tanya!
% SZS status Theorem for 89
% SZS output start Proof for 89
fof(f35514,plain,(
  $false),
  inference(avatar_sat_refutation,[],[f1863,f2359,f2539,f3077,f3527,f3531,f3535,f3537,f4304,f5116,f5255,f5440,f6274,f7205,f10963,f11203,f13063,f13374,f13381,f13564,f13586,f13696,f13701,f13709,f13724,f14213,f14444,f14732,f14743,f14790,f15081,f17288,f17992,f18016,f18919,f21112,f21114,f26641,f35487])).
fof(f35487,plain,(
  ~spl33_176 | ~spl33_1272),
  inference(avatar_contradiction_clause,[],[f35486])).
fof(f35486,plain,(
  $false | (~spl33_176 | ~spl33_1272)),
  inference(subsumption_resolution,[],[f35485,f33928])).
fof(f33928,plain,(
  ( ! [X26,X24,X27,X25] : (para(X24,X25,X26,X27)) ) | (~spl33_176 | ~spl33_1272)),
  inference(subsumption_resolution,[],[f33861,f33736])).
fof(f33736,plain,(
  ( ! [X2,X3] : (midp(X2,sK24,X3)) ) | (~spl33_176 | ~spl33_1272)),
  inference(resolution,[],[f33708,f370])).
fof(f370,plain,(
  ( ! [X2,X0,X1] : (~midp(X2,X1,X0) | midp(X2,X0,X1)) )),
  inference(cnf_transformation,[],[f183])).
fof(f183,plain,(
  ! [X0,X1,X2] : (midp(X2,X0,X1) | ~midp(X2,X1,X0))),
  inference(ennf_transformation,[],[f97])).
fof(f97,plain,(
  ! [X0,X1,X2] : (midp(X2,X1,X0) => midp(X2,X0,X1))),
  inference(rectify,[],[f11])).
fof(f11,axiom,(
  ! [X0,X1,X6] : (midp(X6,X1,X0) => midp(X6,X0,X1))),
  file('./Other/Auxiliary/aux2/89.p',ruleD11)).
fof(f33708,plain,(
  ( ! [X19,X20] : (midp(X19,X20,sK24)) ) | (~spl33_176 | ~spl33_1272)),
  inference(subsumption_resolution,[],[f33707,f17146])).
fof(f17146,plain,(
  ( ! [X30,X29] : (midp(X29,X30,X30)) ) | ~spl33_176),
  inference(subsumption_resolution,[],[f17114,f5337])).
fof(f5337,plain,(
  ( ! [X4,X3] : (coll(X3,X4,X4)) ) | ~spl33_176),
  inference(resolution,[],[f2355,f426])).
fof(f426,plain,(
  ( ! [X2,X0,X1] : (~para(X0,X1,X0,X2) | coll(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f268])).
fof(f268,plain,(
  ! [X0,X1,X2] : (coll(X0,X1,X2) | ~para(X0,X1,X0,X2))),
  inference(ennf_transformation,[],[f67])).
fof(f67,axiom,(
  ! [X0,X1,X2] : (para(X0,X1,X0,X2) => coll(X0,X1,X2))),
  file('./Other/Auxiliary/aux2/89.p',ruleD66)).
fof(f2355,plain,(
  ( ! [X2,X3] : (para(X2,X3,X2,X3)) ) | ~spl33_176),
  inference(avatar_component_clause,[],[f2354])).
fof(f2354,plain,(
  spl33_176 <=> ! [X3,X2] : para(X2,X3,X2,X3)),
  introduced(avatar_definition,[new_symbols(naming,[spl33_176])])).
fof(f17114,plain,(
  ( ! [X30,X29] : (~coll(X29,X30,X30) | midp(X29,X30,X30)) ) | ~spl33_176),
  inference(resolution,[],[f17103,f427])).
fof(f427,plain,(
  ( ! [X2,X0,X1] : (~cong(X0,X1,X0,X2) | ~coll(X0,X1,X2) | midp(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f270])).
fof(f270,plain,(
  ! [X0,X1,X2] : (midp(X0,X1,X2) | ~coll(X0,X1,X2) | ~cong(X0,X1,X0,X2))),
  inference(flattening,[],[f269])).
fof(f269,plain,(
  ! [X0,X1,X2] : (midp(X0,X1,X2) | (~coll(X0,X1,X2) | ~cong(X0,X1,X0,X2)))),
  inference(ennf_transformation,[],[f68])).
fof(f68,axiom,(
  ! [X0,X1,X2] : ((coll(X0,X1,X2) & cong(X0,X1,X0,X2)) => midp(X0,X1,X2))),
  file('./Other/Auxiliary/aux2/89.p',ruleD67)).
fof(f17103,plain,(
  ( ! [X87,X86] : (cong(X86,X87,X86,X87)) ) | ~spl33_176),
  inference(subsumption_resolution,[],[f17102,f10791])).
fof(f10791,plain,(
  ( ! [X12,X10,X11,X9] : (cyclic(X9,X10,X11,X12)) ) | ~spl33_176),
  inference(subsumption_resolution,[],[f10769,f10639])).
fof(f10639,plain,(
  ( ! [X2,X0,X1] : (cyclic(X0,X0,X1,X2)) ) | ~spl33_176),
  inference(resolution,[],[f4467,f2355])).
fof(f4467,plain,(
  ( ! [X68,X66,X67] : (~para(X66,X67,X66,X67) | cyclic(X67,X67,X66,X68)) )),
  inference(subsumption_resolution,[],[f2522,f2523])).
fof(f2523,plain,(
  ( ! [X70,X71,X69] : (~para(X69,X70,X69,X70) | coll(X69,X71,X70) | cyclic(X70,X70,X69,X71)) )),
  inference(resolution,[],[f770,f401])).
fof(f401,plain,(
  ( ! [X2,X0,X3,X1] : (~eqangle(X2,X0,X2,X1,X3,X0,X3,X1) | coll(X2,X3,X0) | cyclic(X0,X1,X2,X3)) )),
  inference(cnf_transformation,[],[f223])).
fof(f223,plain,(
  ! [X0,X1,X2,X3] : (cyclic(X0,X1,X2,X3) | coll(X2,X3,X0) | ~eqangle(X2,X0,X2,X1,X3,X0,X3,X1))),
  inference(flattening,[],[f222])).
fof(f222,plain,(
  ! [X0,X1,X2,X3] : (cyclic(X0,X1,X2,X3) | (coll(X2,X3,X0) | ~eqangle(X2,X0,X2,X1,X3,X0,X3,X1)))),
  inference(ennf_transformation,[],[f121])).
fof(f121,plain,(
  ! [X0,X1,X2,X3] : ((~coll(X2,X3,X0) & eqangle(X2,X0,X2,X1,X3,X0,X3,X1)) => cyclic(X0,X1,X2,X3))),
  inference(rectify,[],[f42])).
fof(f42,axiom,(
  ! [X0,X1,X8,X9] : ((~coll(X8,X9,X0) & eqangle(X8,X0,X8,X1,X9,X0,X9,X1)) => cyclic(X0,X1,X8,X9))),
  file('./Other/Auxiliary/aux2/89.p',ruleD42a)).
fof(f770,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (eqangle(X0,X1,X2,X3,X4,X5,X4,X5) | ~para(X0,X1,X2,X3)) )),
  inference(resolution,[],[f380,f399])).
fof(f399,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (eqangle(X0,X1,X4,X5,X2,X3,X4,X5) | ~para(X0,X1,X2,X3)) )),
  inference(cnf_transformation,[],[f220])).
fof(f220,plain,(
  ! [X0,X1,X2,X3,X4,X5] : (eqangle(X0,X1,X4,X5,X2,X3,X4,X5) | ~para(X0,X1,X2,X3))),
  inference(ennf_transformation,[],[f119])).
fof(f119,plain,(
  ! [X0,X1,X2,X3,X4,X5] : (para(X0,X1,X2,X3) => eqangle(X0,X1,X4,X5,X2,X3,X4,X5))),
  inference(rectify,[],[f40])).
fof(f40,axiom,(
  ! [X0,X1,X2,X3,X8,X9] : (para(X0,X1,X2,X3) => eqangle(X0,X1,X8,X9,X2,X3,X8,X9))),
  file('./Other/Auxiliary/aux2/89.p',ruleD40)).
fof(f380,plain,(
  ( ! [X6,X4,X2,X0,X7,X5,X3,X1] : (~eqangle(X0,X1,X2,X3,X4,X5,X6,X7) | eqangle(X0,X1,X4,X5,X2,X3,X6,X7)) )),
  inference(cnf_transformation,[],[f196])).
fof(f196,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6,X7] : (eqangle(X0,X1,X4,X5,X2,X3,X6,X7) | ~eqangle(X0,X1,X2,X3,X4,X5,X6,X7))),
  inference(ennf_transformation,[],[f103])).
fof(f103,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6,X7] : (eqangle(X0,X1,X2,X3,X4,X5,X6,X7) => eqangle(X0,X1,X4,X5,X2,X3,X6,X7))),
  inference(rectify,[],[f21])).
fof(f21,axiom,(
  ! [X0,X1,X2,X3,X8,X9,X10,X11] : (eqangle(X0,X1,X2,X3,X8,X9,X10,X11) => eqangle(X0,X1,X8,X9,X2,X3,X10,X11))),
  file('./Other/Auxiliary/aux2/89.p',ruleD21)).
fof(f2522,plain,(
  ( ! [X68,X66,X67] : (~para(X66,X67,X66,X67) | ~coll(X66,X68,X67) | cyclic(X67,X67,X66,X68)) )),
  inference(resolution,[],[f770,f402])).
fof(f402,plain,(
  ( ! [X2,X0,X3,X1] : (~eqangle(X2,X0,X2,X1,X3,X0,X3,X1) | ~coll(X2,X3,X1) | cyclic(X0,X1,X2,X3)) )),
  inference(cnf_transformation,[],[f225])).
fof(f225,plain,(
  ! [X0,X1,X2,X3] : (cyclic(X0,X1,X2,X3) | ~coll(X2,X3,X1) | ~eqangle(X2,X0,X2,X1,X3,X0,X3,X1))),
  inference(flattening,[],[f224])).
fof(f224,plain,(
  ! [X0,X1,X2,X3] : (cyclic(X0,X1,X2,X3) | (~coll(X2,X3,X1) | ~eqangle(X2,X0,X2,X1,X3,X0,X3,X1)))),
  inference(ennf_transformation,[],[f122])).
fof(f122,plain,(
  ! [X0,X1,X2,X3] : ((coll(X2,X3,X1) & eqangle(X2,X0,X2,X1,X3,X0,X3,X1)) => cyclic(X0,X1,X2,X3))),
  inference(rectify,[],[f43])).
fof(f43,axiom,(
  ! [X0,X1,X8,X9] : ((coll(X8,X9,X1) & eqangle(X8,X0,X8,X1,X9,X0,X9,X1)) => cyclic(X0,X1,X8,X9))),
  file('./Other/Auxiliary/aux2/89.p',ruleD42b)).
fof(f10769,plain,(
  ( ! [X12,X10,X11,X9] : (cyclic(X9,X10,X11,X12) | ~cyclic(X9,X9,X10,X11)) ) | ~spl33_176),
  inference(resolution,[],[f10639,f376])).
fof(f376,plain,(
  ( ! [X4,X2,X0,X3,X1] : (~cyclic(X0,X1,X2,X4) | cyclic(X1,X2,X3,X4) | ~cyclic(X0,X1,X2,X3)) )),
  inference(cnf_transformation,[],[f192])).
fof(f192,plain,(
  ! [X0,X1,X2,X3,X4] : (cyclic(X1,X2,X3,X4) | ~cyclic(X0,X1,X2,X4) | ~cyclic(X0,X1,X2,X3))),
  inference(flattening,[],[f191])).
fof(f191,plain,(
  ! [X0,X1,X2,X3,X4] : (cyclic(X1,X2,X3,X4) | (~cyclic(X0,X1,X2,X4) | ~cyclic(X0,X1,X2,X3)))),
  inference(ennf_transformation,[],[f17])).
fof(f17,axiom,(
  ! [X0,X1,X2,X3,X4] : ((cyclic(X0,X1,X2,X4) & cyclic(X0,X1,X2,X3)) => cyclic(X1,X2,X3,X4))),
  file('./Other/Auxiliary/aux2/89.p',ruleD17)).
fof(f17102,plain,(
  ( ! [X88,X87,X86] : (cong(X86,X87,X86,X87) | ~cyclic(X86,X87,X88,X87)) ) | ~spl33_176),
  inference(subsumption_resolution,[],[f17101,f10791])).
fof(f17101,plain,(
  ( ! [X88,X87,X86] : (cong(X86,X87,X86,X87) | ~cyclic(X86,X87,X88,X86) | ~cyclic(X86,X87,X88,X87)) ) | ~spl33_176),
  inference(subsumption_resolution,[],[f17087,f10767])).
fof(f10767,plain,(
  ( ! [X4,X5,X3] : (cyclic(X3,X4,X3,X5)) ) | ~spl33_176),
  inference(resolution,[],[f10639,f374])).
fof(f374,plain,(
  ( ! [X2,X0,X3,X1] : (~cyclic(X0,X1,X2,X3) | cyclic(X0,X2,X1,X3)) )),
  inference(cnf_transformation,[],[f189])).
fof(f189,plain,(
  ! [X0,X1,X2,X3] : (cyclic(X0,X2,X1,X3) | ~cyclic(X0,X1,X2,X3))),
  inference(ennf_transformation,[],[f15])).
fof(f15,axiom,(
  ! [X0,X1,X2,X3] : (cyclic(X0,X1,X2,X3) => cyclic(X0,X2,X1,X3))),
  file('./Other/Auxiliary/aux2/89.p',ruleD15)).
fof(f17087,plain,(
  ( ! [X88,X87,X86] : (~cyclic(X86,X87,X86,X88) | cong(X86,X87,X86,X87) | ~cyclic(X86,X87,X88,X86) | ~cyclic(X86,X87,X88,X87)) )),
  inference(duplicate_literal_removal,[],[f17077])).
fof(f17077,plain,(
  ( ! [X88,X87,X86] : (~cyclic(X86,X87,X86,X88) | cong(X86,X87,X86,X87) | ~cyclic(X86,X87,X88,X86) | ~cyclic(X86,X87,X88,X87) | ~cyclic(X86,X87,X88,X86)) )),
  inference(resolution,[],[f2283,f403])).
fof(f403,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (~eqangle(X2,X0,X2,X1,X5,X3,X5,X4) | cong(X0,X1,X3,X4) | ~cyclic(X0,X1,X2,X5) | ~cyclic(X0,X1,X2,X4) | ~cyclic(X0,X1,X2,X3)) )),
  inference(cnf_transformation,[],[f227])).
fof(f227,plain,(
  ! [X0,X1,X2,X3,X4,X5] : (cong(X0,X1,X3,X4) | ~eqangle(X2,X0,X2,X1,X5,X3,X5,X4) | ~cyclic(X0,X1,X2,X5) | ~cyclic(X0,X1,X2,X4) | ~cyclic(X0,X1,X2,X3))),
  inference(flattening,[],[f226])).
fof(f226,plain,(
  ! [X0,X1,X2,X3,X4,X5] : (cong(X0,X1,X3,X4) | (~eqangle(X2,X0,X2,X1,X5,X3,X5,X4) | ~cyclic(X0,X1,X2,X5) | ~cyclic(X0,X1,X2,X4) | ~cyclic(X0,X1,X2,X3)))),
  inference(ennf_transformation,[],[f123])).
fof(f123,plain,(
  ! [X0,X1,X2,X3,X4,X5] : ((eqangle(X2,X0,X2,X1,X5,X3,X5,X4) & cyclic(X0,X1,X2,X5) & cyclic(X0,X1,X2,X4) & cyclic(X0,X1,X2,X3)) => cong(X0,X1,X3,X4))),
  inference(rectify,[],[f44])).
fof(f44,axiom,(
  ! [X0,X1,X2,X8,X9,X14] : ((eqangle(X2,X0,X2,X1,X14,X8,X14,X9) & cyclic(X0,X1,X2,X14) & cyclic(X0,X1,X2,X9) & cyclic(X0,X1,X2,X8)) => cong(X0,X1,X8,X9))),
  file('./Other/Auxiliary/aux2/89.p',ruleD43)).
fof(f2283,plain,(
  ( ! [X14,X12,X13,X11] : (eqangle(X14,X11,X14,X12,X11,X13,X13,X12) | ~cyclic(X11,X12,X13,X14)) )),
  inference(resolution,[],[f741,f379])).
fof(f379,plain,(
  ( ! [X6,X4,X2,X0,X7,X5,X3,X1] : (~eqangle(X0,X1,X2,X3,X4,X5,X6,X7) | eqangle(X4,X5,X6,X7,X0,X1,X2,X3)) )),
  inference(cnf_transformation,[],[f195])).
fof(f195,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6,X7] : (eqangle(X4,X5,X6,X7,X0,X1,X2,X3) | ~eqangle(X0,X1,X2,X3,X4,X5,X6,X7))),
  inference(ennf_transformation,[],[f102])).
fof(f102,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6,X7] : (eqangle(X0,X1,X2,X3,X4,X5,X6,X7) => eqangle(X4,X5,X6,X7,X0,X1,X2,X3))),
  inference(rectify,[],[f20])).
fof(f20,axiom,(
  ! [X0,X1,X2,X3,X8,X9,X10,X11] : (eqangle(X0,X1,X2,X3,X8,X9,X10,X11) => eqangle(X8,X9,X10,X11,X0,X1,X2,X3))),
  file('./Other/Auxiliary/aux2/89.p',ruleD20)).
fof(f741,plain,(
  ( ! [X6,X8,X7,X9] : (eqangle(X6,X7,X7,X8,X9,X6,X9,X8) | ~cyclic(X6,X8,X7,X9)) )),
  inference(resolution,[],[f377,f400])).
fof(f400,plain,(
  ( ! [X2,X0,X3,X1] : (eqangle(X2,X0,X2,X1,X3,X0,X3,X1) | ~cyclic(X0,X1,X2,X3)) )),
  inference(cnf_transformation,[],[f221])).
fof(f221,plain,(
  ! [X0,X1,X2,X3] : (eqangle(X2,X0,X2,X1,X3,X0,X3,X1) | ~cyclic(X0,X1,X2,X3))),
  inference(ennf_transformation,[],[f120])).
fof(f120,plain,(
  ! [X0,X1,X2,X3] : (cyclic(X0,X1,X2,X3) => eqangle(X2,X0,X2,X1,X3,X0,X3,X1))),
  inference(rectify,[],[f41])).
fof(f41,axiom,(
  ! [X0,X1,X8,X9] : (cyclic(X0,X1,X8,X9) => eqangle(X8,X0,X8,X1,X9,X0,X9,X1))),
  file('./Other/Auxiliary/aux2/89.p',ruleD41)).
fof(f377,plain,(
  ( ! [X6,X4,X2,X0,X7,X5,X3,X1] : (~eqangle(X0,X1,X2,X3,X4,X5,X6,X7) | eqangle(X1,X0,X2,X3,X4,X5,X6,X7)) )),
  inference(cnf_transformation,[],[f193])).
fof(f193,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6,X7] : (eqangle(X1,X0,X2,X3,X4,X5,X6,X7) | ~eqangle(X0,X1,X2,X3,X4,X5,X6,X7))),
  inference(ennf_transformation,[],[f100])).
fof(f100,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6,X7] : (eqangle(X0,X1,X2,X3,X4,X5,X6,X7) => eqangle(X1,X0,X2,X3,X4,X5,X6,X7))),
  inference(rectify,[],[f18])).
fof(f18,axiom,(
  ! [X0,X1,X2,X3,X8,X9,X10,X11] : (eqangle(X0,X1,X2,X3,X8,X9,X10,X11) => eqangle(X1,X0,X2,X3,X8,X9,X10,X11))),
  file('./Other/Auxiliary/aux2/89.p',ruleD18)).
fof(f33707,plain,(
  ( ! [X19,X20] : (midp(X19,X20,sK24) | ~midp(X19,sK23,sK23)) ) | (~spl33_176 | ~spl33_1272)),
  inference(subsumption_resolution,[],[f33680,f33677])).
fof(f33677,plain,(
  ( ! [X13] : (para(sK23,X13,sK23,sK24)) ) | (~spl33_176 | ~spl33_1272)),
  inference(resolution,[],[f21488,f364])).
fof(f364,plain,(
  ( ! [X2,X0,X3,X1] : (~para(X0,X1,X2,X3) | para(X2,X3,X0,X1)) )),
  inference(cnf_transformation,[],[f174])).
fof(f174,plain,(
  ! [X0,X1,X2,X3] : (para(X2,X3,X0,X1) | ~para(X0,X1,X2,X3))),
  inference(ennf_transformation,[],[f5])).
fof(f5,axiom,(
  ! [X0,X1,X2,X3] : (para(X0,X1,X2,X3) => para(X2,X3,X0,X1))),
  file('./Other/Auxiliary/aux2/89.p',ruleD5)).
fof(f21488,plain,(
  ( ! [X11] : (para(sK23,sK24,sK23,X11)) ) | (~spl33_176 | ~spl33_1272)),
  inference(subsumption_resolution,[],[f21481,f17864])).
fof(f17864,plain,(
  ( ! [X0] : (midp(X0,sK26,sK24)) ) | ~spl33_176),
  inference(resolution,[],[f17282,f370])).
fof(f17282,plain,(
  ( ! [X96] : (midp(X96,sK24,sK26)) ) | ~spl33_176),
  inference(subsumption_resolution,[],[f17256,f17146])).
fof(f17256,plain,(
  ( ! [X97,X96] : (midp(X96,sK24,sK26) | ~midp(X96,X97,X97)) ) | ~spl33_176),
  inference(resolution,[],[f17146,f1608])).
fof(f1608,plain,(
  ( ! [X2,X3,X1] : (~midp(sK30,X1,X2) | midp(X3,sK24,sK26) | ~midp(X3,X1,X2)) )),
  inference(subsumption_resolution,[],[f1603,f589])).
fof(f589,plain,(
  ( ! [X19,X18] : (para(X18,sK24,X19,sK26) | ~midp(sK30,X18,X19)) )),
  inference(resolution,[],[f423,f490])).
fof(f490,plain,(
  midp(sK30,sK24,sK26)),
  inference(resolution,[],[f370,f479])).
fof(f479,plain,(
  midp(sK30,sK26,sK24)),
  inference(cnf_transformation,[],[f359])).
fof(f359,plain,(
  ~eqangle(sK31,sK29,sK29,sK32,sK32,sK29,sK29,sK30) & perp(sK29,sK28,sK20,sK32) & midp(sK20,sK29,sK28) & perp(sK27,sK28,sK21,sK32) & midp(sK21,sK27,sK28) & perp(sK27,sK29,sK22,sK32) & midp(sK22,sK27,sK29) & midp(sK31,sK25,sK26) & midp(sK30,sK26,sK24) & midp(sK29,sK24,sK25) & midp(sK28,sK25,sK23) & midp(sK27,sK24,sK23) & eqangle(sK26,sK25,sK25,sK23,sK26,sK25,sK25,sK24) & eqangle(sK26,sK24,sK24,sK25,sK26,sK24,sK24,sK23) & eqangle(sK26,sK23,sK23,sK24,sK26,sK23,sK23,sK25)),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK20,sK21,sK22,sK23,sK24,sK25,sK26,sK27,sK28,sK29,sK30,sK31,sK32])],[f321,f358])).
fof(f358,plain,(
  ? [X0,X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12] : (~eqangle(X11,X9,X9,X12,X12,X9,X9,X10) & perp(X9,X8,X0,X12) & midp(X0,X9,X8) & perp(X7,X8,X1,X12) & midp(X1,X7,X8) & perp(X7,X9,X2,X12) & midp(X2,X7,X9) & midp(X11,X5,X6) & midp(X10,X6,X4) & midp(X9,X4,X5) & midp(X8,X5,X3) & midp(X7,X4,X3) & eqangle(X6,X5,X5,X3,X6,X5,X5,X4) & eqangle(X6,X4,X4,X5,X6,X4,X4,X3) & eqangle(X6,X3,X3,X4,X6,X3,X3,X5)) => (~eqangle(sK31,sK29,sK29,sK32,sK32,sK29,sK29,sK30) & perp(sK29,sK28,sK20,sK32) & midp(sK20,sK29,sK28) & perp(sK27,sK28,sK21,sK32) & midp(sK21,sK27,sK28) & perp(sK27,sK29,sK22,sK32) & midp(sK22,sK27,sK29) & midp(sK31,sK25,sK26) & midp(sK30,sK26,sK24) & midp(sK29,sK24,sK25) & midp(sK28,sK25,sK23) & midp(sK27,sK24,sK23) & eqangle(sK26,sK25,sK25,sK23,sK26,sK25,sK25,sK24) & eqangle(sK26,sK24,sK24,sK25,sK26,sK24,sK24,sK23) & eqangle(sK26,sK23,sK23,sK24,sK26,sK23,sK23,sK25))),
  introduced(choice_axiom,[])).
fof(f321,plain,(
  ? [X0,X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12] : (~eqangle(X11,X9,X9,X12,X12,X9,X9,X10) & perp(X9,X8,X0,X12) & midp(X0,X9,X8) & perp(X7,X8,X1,X12) & midp(X1,X7,X8) & perp(X7,X9,X2,X12) & midp(X2,X7,X9) & midp(X11,X5,X6) & midp(X10,X6,X4) & midp(X9,X4,X5) & midp(X8,X5,X3) & midp(X7,X4,X3) & eqangle(X6,X5,X5,X3,X6,X5,X5,X4) & eqangle(X6,X4,X4,X5,X6,X4,X4,X3) & eqangle(X6,X3,X3,X4,X6,X3,X3,X5))),
  inference(flattening,[],[f320])).
fof(f320,plain,(
  ? [X0,X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12] : (~eqangle(X11,X9,X9,X12,X12,X9,X9,X10) & (perp(X9,X8,X0,X12) & midp(X0,X9,X8) & perp(X7,X8,X1,X12) & midp(X1,X7,X8) & perp(X7,X9,X2,X12) & midp(X2,X7,X9) & midp(X11,X5,X6) & midp(X10,X6,X4) & midp(X9,X4,X5) & midp(X8,X5,X3) & midp(X7,X4,X3) & eqangle(X6,X5,X5,X3,X6,X5,X5,X4) & eqangle(X6,X4,X4,X5,X6,X4,X4,X3) & eqangle(X6,X3,X3,X4,X6,X3,X3,X5)))),
  inference(ennf_transformation,[],[f168])).
fof(f168,plain,(
  ~! [X0,X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12] : ((perp(X9,X8,X0,X12) & midp(X0,X9,X8) & perp(X7,X8,X1,X12) & midp(X1,X7,X8) & perp(X7,X9,X2,X12) & midp(X2,X7,X9) & midp(X11,X5,X6) & midp(X10,X6,X4) & midp(X9,X4,X5) & midp(X8,X5,X3) & midp(X7,X4,X3) & eqangle(X6,X5,X5,X3,X6,X5,X5,X4) & eqangle(X6,X4,X4,X5,X6,X4,X4,X3) & eqangle(X6,X3,X3,X4,X6,X3,X3,X5)) => eqangle(X11,X9,X9,X12,X12,X9,X9,X10))),
  inference(flattening,[],[f167])).
fof(f167,plain,(
  ~! [X0,X1,X2] : ! [X3,X4,X5,X6,X7,X8,X9,X10,X11,X12] : ((perp(X9,X8,X0,X12) & midp(X0,X9,X8) & perp(X7,X8,X1,X12) & midp(X1,X7,X8) & perp(X7,X9,X2,X12) & midp(X2,X7,X9) & midp(X11,X5,X6) & midp(X10,X6,X4) & midp(X9,X4,X5) & midp(X8,X5,X3) & midp(X7,X4,X3) & eqangle(X6,X5,X5,X3,X6,X5,X5,X4) & eqangle(X6,X4,X4,X5,X6,X4,X4,X3) & eqangle(X6,X3,X3,X4,X6,X3,X3,X5)) => eqangle(X11,X9,X9,X12,X12,X9,X9,X10))),
  inference(rectify,[],[f96])).
fof(f96,negated_conjecture,(
  ~! [X22,X21,X20] : ! [X0,X1,X2,X3,X4,X5,X12,X13,X18,X19] : ((perp(X12,X5,X22,X19) & midp(X22,X12,X5) & perp(X4,X5,X21,X19) & midp(X21,X4,X5) & perp(X4,X12,X20,X19) & midp(X20,X4,X12) & midp(X18,X2,X3) & midp(X13,X3,X1) & midp(X12,X1,X2) & midp(X5,X2,X0) & midp(X4,X1,X0) & eqangle(X3,X2,X2,X0,X3,X2,X2,X1) & eqangle(X3,X1,X1,X2,X3,X1,X1,X0) & eqangle(X3,X0,X0,X1,X3,X0,X0,X2)) => eqangle(X18,X12,X12,X19,X19,X12,X12,X13))),
  inference(negated_conjecture,[],[f95])).
fof(f95,conjecture,(
  ( ! [X21,X22,X20] : (! [X0,X1,X2,X3,X4,X5,X12,X13,X18,X19] : ((perp(X12,X5,X22,X19) & midp(X22,X12,X5) & perp(X4,X5,X21,X19) & midp(X21,X4,X5) & perp(X4,X12,X20,X19) & midp(X20,X4,X12) & midp(X18,X2,X3) & midp(X13,X3,X1) & midp(X12,X1,X2) & midp(X5,X2,X0) & midp(X4,X1,X0) & eqangle(X3,X2,X2,X0,X3,X2,X2,X1) & eqangle(X3,X1,X1,X2,X3,X1,X1,X0) & eqangle(X3,X0,X0,X1,X3,X0,X0,X2)) => eqangle(X18,X12,X12,X19,X19,X12,X12,X13))) )),
  file('./Other/Auxiliary/aux2/89.p',exemploOtherAuxiliaryaux289)).
fof(f423,plain,(
  ( ! [X4,X2,X0,X3,X1] : (~midp(X4,X2,X3) | para(X0,X2,X1,X3) | ~midp(X4,X0,X1)) )),
  inference(cnf_transformation,[],[f263])).
fof(f263,plain,(
  ! [X0,X1,X2,X3,X4] : (para(X0,X2,X1,X3) | ~midp(X4,X2,X3) | ~midp(X4,X0,X1))),
  inference(flattening,[],[f262])).
fof(f262,plain,(
  ! [X0,X1,X2,X3,X4] : (para(X0,X2,X1,X3) | (~midp(X4,X2,X3) | ~midp(X4,X0,X1)))),
  inference(ennf_transformation,[],[f142])).
fof(f142,plain,(
  ! [X0,X1,X2,X3,X4] : ((midp(X4,X2,X3) & midp(X4,X0,X1)) => para(X0,X2,X1,X3))),
  inference(rectify,[],[f64])).
fof(f64,axiom,(
  ! [X0,X1,X2,X3,X6] : ((midp(X6,X2,X3) & midp(X6,X0,X1)) => para(X0,X2,X1,X3))),
  file('./Other/Auxiliary/aux2/89.p',ruleD63)).
fof(f1603,plain,(
  ( ! [X2,X3,X1] : (~midp(sK30,X1,X2) | midp(X3,sK24,sK26) | ~para(X1,sK24,X2,sK26) | ~midp(X3,X1,X2)) )),
  inference(resolution,[],[f582,f424])).
fof(f424,plain,(
  ( ! [X4,X2,X0,X3,X1] : (~para(X0,X3,X1,X2) | midp(X4,X2,X3) | ~para(X0,X2,X1,X3) | ~midp(X4,X0,X1)) )),
  inference(cnf_transformation,[],[f265])).
fof(f265,plain,(
  ! [X0,X1,X2,X3,X4] : (midp(X4,X2,X3) | ~para(X0,X3,X1,X2) | ~para(X0,X2,X1,X3) | ~midp(X4,X0,X1))),
  inference(flattening,[],[f264])).
fof(f264,plain,(
  ! [X0,X1,X2,X3,X4] : (midp(X4,X2,X3) | (~para(X0,X3,X1,X2) | ~para(X0,X2,X1,X3) | ~midp(X4,X0,X1)))),
  inference(ennf_transformation,[],[f143])).
fof(f143,plain,(
  ! [X0,X1,X2,X3,X4] : ((para(X0,X3,X1,X2) & para(X0,X2,X1,X3) & midp(X4,X0,X1)) => midp(X4,X2,X3))),
  inference(rectify,[],[f65])).
fof(f65,axiom,(
  ! [X0,X1,X2,X3,X6] : ((para(X0,X3,X1,X2) & para(X0,X2,X1,X3) & midp(X6,X0,X1)) => midp(X6,X2,X3))),
  file('./Other/Auxiliary/aux2/89.p',ruleD64)).
fof(f582,plain,(
  ( ! [X4,X5] : (para(X4,sK26,X5,sK24) | ~midp(sK30,X4,X5)) )),
  inference(resolution,[],[f423,f479])).
fof(f21481,plain,(
  ( ! [X11] : (~midp(X11,sK26,sK24) | para(sK23,sK24,sK23,X11)) ) | ~spl33_1272),
  inference(resolution,[],[f14579,f3333])).
fof(f3333,plain,(
  ( ! [X2,X0,X3,X1] : (~circle(X0,X0,X1,X2) | ~midp(X3,X1,X2) | para(X0,X2,X0,X3)) )),
  inference(resolution,[],[f822,f398])).
fof(f398,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (~eqangle(X0,X1,X4,X5,X2,X3,X4,X5) | para(X0,X1,X2,X3)) )),
  inference(cnf_transformation,[],[f219])).
fof(f219,plain,(
  ! [X0,X1,X2,X3,X4,X5] : (para(X0,X1,X2,X3) | ~eqangle(X0,X1,X4,X5,X2,X3,X4,X5))),
  inference(ennf_transformation,[],[f118])).
fof(f118,plain,(
  ! [X0,X1,X2,X3,X4,X5] : (eqangle(X0,X1,X4,X5,X2,X3,X4,X5) => para(X0,X1,X2,X3))),
  inference(rectify,[],[f39])).
fof(f39,axiom,(
  ! [X0,X1,X2,X3,X8,X9] : (eqangle(X0,X1,X8,X9,X2,X3,X8,X9) => para(X0,X1,X2,X3))),
  file('./Other/Auxiliary/aux2/89.p',ruleD39)).
fof(f822,plain,(
  ( ! [X14,X17,X15,X13,X16] : (eqangle(X17,X15,X17,X14,X16,X13,X16,X14) | ~circle(X16,X17,X14,X15) | ~midp(X13,X14,X15)) )),
  inference(resolution,[],[f410,f378])).
fof(f378,plain,(
  ( ! [X6,X4,X2,X0,X7,X5,X3,X1] : (~eqangle(X0,X1,X2,X3,X4,X5,X6,X7) | eqangle(X2,X3,X0,X1,X6,X7,X4,X5)) )),
  inference(cnf_transformation,[],[f194])).
fof(f194,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6,X7] : (eqangle(X2,X3,X0,X1,X6,X7,X4,X5) | ~eqangle(X0,X1,X2,X3,X4,X5,X6,X7))),
  inference(ennf_transformation,[],[f101])).
fof(f101,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6,X7] : (eqangle(X0,X1,X2,X3,X4,X5,X6,X7) => eqangle(X2,X3,X0,X1,X6,X7,X4,X5))),
  inference(rectify,[],[f19])).
fof(f19,axiom,(
  ! [X0,X1,X2,X3,X8,X9,X10,X11] : (eqangle(X0,X1,X2,X3,X8,X9,X10,X11) => eqangle(X2,X3,X0,X1,X10,X11,X8,X9))),
  file('./Other/Auxiliary/aux2/89.p',ruleD19)).
fof(f410,plain,(
  ( ! [X4,X2,X0,X3,X1] : (eqangle(X0,X1,X0,X2,X3,X1,X3,X4) | ~midp(X4,X1,X2) | ~circle(X3,X0,X1,X2)) )),
  inference(cnf_transformation,[],[f240])).
fof(f240,plain,(
  ! [X0,X1,X2,X3,X4] : (eqangle(X0,X1,X0,X2,X3,X1,X3,X4) | ~midp(X4,X1,X2) | ~circle(X3,X0,X1,X2))),
  inference(flattening,[],[f239])).
fof(f239,plain,(
  ! [X0,X1,X2,X3,X4] : (eqangle(X0,X1,X0,X2,X3,X1,X3,X4) | (~midp(X4,X1,X2) | ~circle(X3,X0,X1,X2)))),
  inference(ennf_transformation,[],[f130])).
fof(f130,plain,(
  ! [X0,X1,X2,X3,X4] : ((midp(X4,X1,X2) & circle(X3,X0,X1,X2)) => eqangle(X0,X1,X0,X2,X3,X1,X3,X4))),
  inference(rectify,[],[f51])).
fof(f51,axiom,(
  ! [X0,X1,X2,X7,X6] : ((midp(X6,X1,X2) & circle(X7,X0,X1,X2)) => eqangle(X0,X1,X0,X2,X7,X1,X7,X6))),
  file('./Other/Auxiliary/aux2/89.p',ruleD50)).
fof(f14579,plain,(
  circle(sK23,sK23,sK26,sK24) | ~spl33_1272),
  inference(avatar_component_clause,[],[f14578])).
fof(f14578,plain,(
  spl33_1272 <=> circle(sK23,sK23,sK26,sK24)),
  introduced(avatar_definition,[new_symbols(naming,[spl33_1272])])).
fof(f33680,plain,(
  ( ! [X19,X20] : (midp(X19,X20,sK24) | ~para(sK23,X20,sK23,sK24) | ~midp(X19,sK23,sK23)) ) | (~spl33_176 | ~spl33_1272)),
  inference(resolution,[],[f21488,f424])).
fof(f33861,plain,(
  ( ! [X26,X24,X27,X25] : (para(X24,X25,X26,X27) | ~midp(X24,sK24,X26)) ) | (~spl33_176 | ~spl33_1272)),
  inference(resolution,[],[f33736,f404])).
fof(f404,plain,(
  ( ! [X4,X2,X0,X3,X1] : (~midp(X4,X0,X2) | para(X3,X4,X1,X2) | ~midp(X3,X0,X1)) )),
  inference(cnf_transformation,[],[f229])).
fof(f229,plain,(
  ! [X0,X1,X2,X3,X4] : (para(X3,X4,X1,X2) | ~midp(X4,X0,X2) | ~midp(X3,X0,X1))),
  inference(flattening,[],[f228])).
fof(f228,plain,(
  ! [X0,X1,X2,X3,X4] : (para(X3,X4,X1,X2) | (~midp(X4,X0,X2) | ~midp(X3,X0,X1)))),
  inference(ennf_transformation,[],[f124])).
fof(f124,plain,(
  ! [X0,X1,X2,X3,X4] : ((midp(X4,X0,X2) & midp(X3,X0,X1)) => para(X3,X4,X1,X2))),
  inference(rectify,[],[f45])).
fof(f45,axiom,(
  ! [X0,X1,X2,X4,X5] : ((midp(X5,X0,X2) & midp(X4,X0,X1)) => para(X4,X5,X1,X2))),
  file('./Other/Auxiliary/aux2/89.p',ruleD44)).
fof(f35485,plain,(
  ~para(sK29,sK32,sK29,sK30) | (~spl33_176 | ~spl33_1272)),
  inference(subsumption_resolution,[],[f35435,f33928])).
fof(f35435,plain,(
  ~para(sK31,sK29,sK32,sK29) | ~para(sK29,sK32,sK29,sK30)),
  inference(resolution,[],[f5304,f487])).
fof(f487,plain,(
  ~eqangle(sK31,sK29,sK29,sK32,sK32,sK29,sK29,sK30)),
  inference(cnf_transformation,[],[f359])).
fof(f5304,plain,(
  ( ! [X80,X87,X85,X83,X81,X86,X84,X82] : (eqangle(X80,X81,X82,X83,X84,X85,X86,X87) | ~para(X80,X81,X84,X85) | ~para(X82,X83,X86,X87)) )),
  inference(resolution,[],[f1357,f750])).
fof(f750,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (eqangle(X0,X1,X2,X3,X0,X1,X4,X5) | ~para(X2,X3,X4,X5)) )),
  inference(resolution,[],[f378,f399])).
fof(f1357,plain,(
  ( ! [X14,X12,X21,X19,X17,X15,X13,X20,X18,X16] : (~eqangle(X12,X13,X14,X15,X20,X21,X18,X19) | eqangle(X12,X13,X14,X15,X16,X17,X18,X19) | ~para(X20,X21,X16,X17)) )),
  inference(resolution,[],[f381,f399])).
fof(f381,plain,(
  ( ! [X6,X4,X2,X0,X10,X8,X7,X5,X3,X1,X11,X9] : (~eqangle(X4,X5,X6,X7,X8,X9,X10,X11) | eqangle(X0,X1,X2,X3,X8,X9,X10,X11) | ~eqangle(X0,X1,X2,X3,X4,X5,X6,X7)) )),
  inference(cnf_transformation,[],[f198])).
fof(f198,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11] : (eqangle(X0,X1,X2,X3,X8,X9,X10,X11) | ~eqangle(X4,X5,X6,X7,X8,X9,X10,X11) | ~eqangle(X0,X1,X2,X3,X4,X5,X6,X7))),
  inference(flattening,[],[f197])).
fof(f197,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11] : (eqangle(X0,X1,X2,X3,X8,X9,X10,X11) | (~eqangle(X4,X5,X6,X7,X8,X9,X10,X11) | ~eqangle(X0,X1,X2,X3,X4,X5,X6,X7)))),
  inference(ennf_transformation,[],[f104])).
fof(f104,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11] : ((eqangle(X4,X5,X6,X7,X8,X9,X10,X11) & eqangle(X0,X1,X2,X3,X4,X5,X6,X7)) => eqangle(X0,X1,X2,X3,X8,X9,X10,X11))),
  inference(rectify,[],[f22])).
fof(f22,axiom,(
  ! [X0,X1,X2,X3,X8,X9,X10,X11,X4,X5,X12,X13] : ((eqangle(X8,X9,X10,X11,X4,X5,X12,X13) & eqangle(X0,X1,X2,X3,X8,X9,X10,X11)) => eqangle(X0,X1,X2,X3,X4,X5,X12,X13))),
  file('./Other/Auxiliary/aux2/89.p',ruleD22)).
fof(f26641,plain,(
  spl33_753 | ~spl33_1122),
  inference(avatar_contradiction_clause,[],[f26640])).
fof(f26640,plain,(
  $false | (~spl33_753 | ~spl33_1122)),
  inference(subsumption_resolution,[],[f26639,f13239])).
fof(f13239,plain,(
  midp(sK25,sK25,sK23) | ~spl33_1122),
  inference(avatar_component_clause,[],[f13238])).
fof(f13238,plain,(
  spl33_1122 <=> midp(sK25,sK25,sK23)),
  introduced(avatar_definition,[new_symbols(naming,[spl33_1122])])).
fof(f26639,plain,(
  ~midp(sK25,sK25,sK23) | ~spl33_753),
  inference(duplicate_literal_removal,[],[f26637])).
fof(f26637,plain,(
  ~midp(sK25,sK25,sK23) | ~midp(sK25,sK25,sK23) | ~spl33_753),
  inference(resolution,[],[f7204,f2198])).
fof(f2198,plain,(
  ( ! [X6,X4,X7,X5] : (circle(X4,X5,X6,X7) | ~midp(X4,X5,X7) | ~midp(X4,X5,X6)) )),
  inference(resolution,[],[f702,f428])).
fof(f428,plain,(
  ( ! [X2,X0,X1] : (cong(X0,X1,X0,X2) | ~midp(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f271])).
fof(f271,plain,(
  ! [X0,X1,X2] : (cong(X0,X1,X0,X2) | ~midp(X0,X1,X2))),
  inference(ennf_transformation,[],[f69])).
fof(f69,axiom,(
  ! [X0,X1,X2] : (midp(X0,X1,X2) => cong(X0,X1,X0,X2))),
  file('./Other/Auxiliary/aux2/89.p',ruleD68)).
fof(f702,plain,(
  ( ! [X2,X0,X3,X1] : (~cong(X0,X1,X0,X2) | circle(X0,X1,X2,X3) | ~midp(X0,X1,X3)) )),
  inference(resolution,[],[f371,f428])).
fof(f371,plain,(
  ( ! [X2,X0,X3,X1] : (~cong(X3,X0,X3,X2) | circle(X3,X0,X1,X2) | ~cong(X3,X0,X3,X1)) )),
  inference(cnf_transformation,[],[f185])).
fof(f185,plain,(
  ! [X0,X1,X2,X3] : (circle(X3,X0,X1,X2) | ~cong(X3,X0,X3,X2) | ~cong(X3,X0,X3,X1))),
  inference(flattening,[],[f184])).
fof(f184,plain,(
  ! [X0,X1,X2,X3] : (circle(X3,X0,X1,X2) | (~cong(X3,X0,X3,X2) | ~cong(X3,X0,X3,X1)))),
  inference(ennf_transformation,[],[f98])).
fof(f98,plain,(
  ! [X0,X1,X2,X3] : ((cong(X3,X0,X3,X2) & cong(X3,X0,X3,X1)) => circle(X3,X0,X1,X2))),
  inference(rectify,[],[f12])).
fof(f12,axiom,(
  ! [X0,X1,X2,X7] : ((cong(X7,X0,X7,X2) & cong(X7,X0,X7,X1)) => circle(X7,X0,X1,X2))),
  file('./Other/Auxiliary/aux2/89.p',ruleD12)).
fof(f7204,plain,(
  ~circle(sK25,sK25,sK23,sK23) | ~spl33_753),
  inference(avatar_component_clause,[],[f7203])).
fof(f7203,plain,(
  spl33_753 <=> ~circle(sK25,sK25,sK23,sK23)),
  introduced(avatar_definition,[new_symbols(naming,[spl33_753])])).
fof(f21114,plain,(
  ~spl33_859 | ~spl33_1166 | spl33_1273),
  inference(avatar_split_clause,[],[f21113,f14581,f13624,f8716])).
fof(f8716,plain,(
  spl33_859 <=> ~midp(sK23,sK26,sK23)),
  introduced(avatar_definition,[new_symbols(naming,[spl33_859])])).
fof(f13624,plain,(
  spl33_1166 <=> midp(sK23,sK23,sK24)),
  introduced(avatar_definition,[new_symbols(naming,[spl33_1166])])).
fof(f14581,plain,(
  spl33_1273 <=> ~circle(sK23,sK23,sK26,sK24)),
  introduced(avatar_definition,[new_symbols(naming,[spl33_1273])])).
fof(f21113,plain,(
  ~midp(sK23,sK26,sK23) | (~spl33_1166 | ~spl33_1273)),
  inference(subsumption_resolution,[],[f16279,f13625])).
fof(f13625,plain,(
  midp(sK23,sK23,sK24) | ~spl33_1166),
  inference(avatar_component_clause,[],[f13624])).
fof(f16279,plain,(
  ~midp(sK23,sK23,sK24) | ~midp(sK23,sK26,sK23) | ~spl33_1273),
  inference(resolution,[],[f14582,f2197])).
fof(f2197,plain,(
  ( ! [X2,X0,X3,X1] : (circle(X0,X1,X2,X3) | ~midp(X0,X1,X3) | ~midp(X0,X2,X1)) )),
  inference(resolution,[],[f702,f543])).
fof(f543,plain,(
  ( ! [X2,X0,X1] : (cong(X0,X1,X0,X2) | ~midp(X0,X2,X1)) )),
  inference(resolution,[],[f383,f428])).
fof(f383,plain,(
  ( ! [X2,X0,X3,X1] : (~cong(X0,X1,X2,X3) | cong(X2,X3,X0,X1)) )),
  inference(cnf_transformation,[],[f200])).
fof(f200,plain,(
  ! [X0,X1,X2,X3] : (cong(X2,X3,X0,X1) | ~cong(X0,X1,X2,X3))),
  inference(ennf_transformation,[],[f24])).
fof(f24,axiom,(
  ! [X0,X1,X2,X3] : (cong(X0,X1,X2,X3) => cong(X2,X3,X0,X1))),
  file('./Other/Auxiliary/aux2/89.p',ruleD24)).
fof(f14582,plain,(
  ~circle(sK23,sK23,sK26,sK24) | ~spl33_1273),
  inference(avatar_component_clause,[],[f14581])).
fof(f21112,plain,(
  spl33_858 | ~spl33_176 | ~spl33_1286),
  inference(avatar_split_clause,[],[f21111,f14661,f2354,f8719])).
fof(f8719,plain,(
  spl33_858 <=> midp(sK23,sK26,sK23)),
  introduced(avatar_definition,[new_symbols(naming,[spl33_858])])).
fof(f14661,plain,(
  spl33_1286 <=> circle(sK25,sK25,sK26,sK23)),
  introduced(avatar_definition,[new_symbols(naming,[spl33_1286])])).
fof(f21111,plain,(
  midp(sK23,sK26,sK23) | (~spl33_176 | ~spl33_1286)),
  inference(subsumption_resolution,[],[f21107,f2355])).
fof(f21107,plain,(
  midp(sK23,sK26,sK23) | ~para(sK25,sK26,sK25,sK26) | ~spl33_1286),
  inference(subsumption_resolution,[],[f21104,f2122])).
fof(f2122,plain,(
  coll(sK23,sK26,sK23)),
  inference(resolution,[],[f2108,f361])).
fof(f361,plain,(
  ( ! [X2,X0,X1] : (~coll(X0,X1,X2) | coll(X1,X0,X2)) )),
  inference(cnf_transformation,[],[f170])).
fof(f170,plain,(
  ! [X0,X1,X2] : (coll(X1,X0,X2) | ~coll(X0,X1,X2))),
  inference(ennf_transformation,[],[f2])).
fof(f2,axiom,(
  ! [X0,X1,X2] : (coll(X0,X1,X2) => coll(X1,X0,X2))),
  file('./Other/Auxiliary/aux2/89.p',ruleD2)).
fof(f2108,plain,(
  coll(sK26,sK23,sK23)),
  inference(resolution,[],[f1881,f426])).
fof(f1881,plain,(
  para(sK26,sK23,sK26,sK23)),
  inference(subsumption_resolution,[],[f1878,f1697])).
fof(f1697,plain,(
  para(sK23,sK24,sK23,sK25)),
  inference(resolution,[],[f747,f398])).
fof(f747,plain,(
  eqangle(sK23,sK24,sK26,sK23,sK23,sK25,sK26,sK23)),
  inference(resolution,[],[f378,f473])).
fof(f473,plain,(
  eqangle(sK26,sK23,sK23,sK24,sK26,sK23,sK23,sK25)),
  inference(cnf_transformation,[],[f359])).
fof(f1878,plain,(
  ~para(sK23,sK24,sK23,sK25) | para(sK26,sK23,sK26,sK23)),
  inference(resolution,[],[f767,f433])).
fof(f433,plain,(
  ( ! [X6,X4,X2,X0,X7,X5,X3,X1] : (~eqangle(X0,X1,X2,X3,X4,X5,X6,X7) | ~para(X4,X5,X6,X7) | para(X0,X1,X2,X3)) )),
  inference(cnf_transformation,[],[f280])).
fof(f280,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6,X7] : (para(X0,X1,X2,X3) | ~para(X4,X5,X6,X7) | ~eqangle(X0,X1,X2,X3,X4,X5,X6,X7))),
  inference(flattening,[],[f279])).
fof(f279,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6,X7] : (para(X0,X1,X2,X3) | (~para(X4,X5,X6,X7) | ~eqangle(X0,X1,X2,X3,X4,X5,X6,X7)))),
  inference(ennf_transformation,[],[f146])).
fof(f146,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6,X7] : ((para(X4,X5,X6,X7) & eqangle(X0,X1,X2,X3,X4,X5,X6,X7)) => para(X0,X1,X2,X3))),
  inference(rectify,[],[f74])).
fof(f74,axiom,(
  ! [X0,X1,X2,X3,X8,X9,X10,X11] : ((para(X8,X9,X10,X11) & eqangle(X0,X1,X2,X3,X8,X9,X10,X11)) => para(X0,X1,X2,X3))),
  file('./Other/Auxiliary/aux2/89.p',ruleD73)).
fof(f767,plain,(
  eqangle(sK26,sK23,sK26,sK23,sK23,sK24,sK23,sK25)),
  inference(resolution,[],[f380,f473])).
fof(f21104,plain,(
  ~coll(sK23,sK26,sK23) | midp(sK23,sK26,sK23) | ~para(sK25,sK26,sK25,sK26) | ~spl33_1286),
  inference(resolution,[],[f14662,f1202])).
fof(f1202,plain,(
  ( ! [X2,X0,X1] : (~circle(X2,X2,X1,X0) | ~coll(X0,X1,X0) | midp(X0,X1,X0) | ~para(X2,X1,X2,X1)) )),
  inference(resolution,[],[f411,f399])).
fof(f411,plain,(
  ( ! [X4,X2,X0,X3,X1] : (~eqangle(X0,X1,X0,X2,X3,X1,X3,X4) | midp(X4,X1,X2) | ~coll(X4,X1,X2) | ~circle(X3,X0,X1,X2)) )),
  inference(cnf_transformation,[],[f242])).
fof(f242,plain,(
  ! [X0,X1,X2,X3,X4] : (midp(X4,X1,X2) | ~eqangle(X0,X1,X0,X2,X3,X1,X3,X4) | ~coll(X4,X1,X2) | ~circle(X3,X0,X1,X2))),
  inference(flattening,[],[f241])).
fof(f241,plain,(
  ! [X0,X1,X2,X3,X4] : (midp(X4,X1,X2) | (~eqangle(X0,X1,X0,X2,X3,X1,X3,X4) | ~coll(X4,X1,X2) | ~circle(X3,X0,X1,X2)))),
  inference(ennf_transformation,[],[f131])).
fof(f131,plain,(
  ! [X0,X1,X2,X3,X4] : ((eqangle(X0,X1,X0,X2,X3,X1,X3,X4) & coll(X4,X1,X2) & circle(X3,X0,X1,X2)) => midp(X4,X1,X2))),
  inference(rectify,[],[f52])).
fof(f52,axiom,(
  ! [X0,X1,X2,X7,X6] : ((eqangle(X0,X1,X0,X2,X7,X1,X7,X6) & coll(X6,X1,X2) & circle(X7,X0,X1,X2)) => midp(X6,X1,X2))),
  file('./Other/Auxiliary/aux2/89.p',ruleD51)).
fof(f14662,plain,(
  circle(sK25,sK25,sK26,sK23) | ~spl33_1286),
  inference(avatar_component_clause,[],[f14661])).
fof(f18919,plain,(
  ~spl33_176 | spl33_1301),
  inference(avatar_contradiction_clause,[],[f18907])).
fof(f18907,plain,(
  $false | (~spl33_176 | ~spl33_1301)),
  inference(resolution,[],[f17864,f14738])).
fof(f14738,plain,(
  ~midp(sK23,sK26,sK24) | ~spl33_1301),
  inference(avatar_component_clause,[],[f14737])).
fof(f14737,plain,(
  spl33_1301 <=> ~midp(sK23,sK26,sK24)),
  introduced(avatar_definition,[new_symbols(naming,[spl33_1301])])).
fof(f18016,plain,(
  ~spl33_176 | ~spl33_1122 | spl33_1287),
  inference(avatar_contradiction_clause,[],[f18015])).
fof(f18015,plain,(
  $false | (~spl33_176 | ~spl33_1122 | ~spl33_1287)),
  inference(subsumption_resolution,[],[f18014,f17283])).
fof(f17283,plain,(
  ( ! [X99] : (midp(X99,sK26,sK25)) ) | ~spl33_176),
  inference(subsumption_resolution,[],[f17258,f17146])).
fof(f17258,plain,(
  ( ! [X99,X100] : (midp(X99,sK26,sK25) | ~midp(X99,X100,X100)) ) | ~spl33_176),
  inference(resolution,[],[f17146,f1629])).
fof(f1629,plain,(
  ( ! [X2,X3,X1] : (~midp(sK31,X1,X2) | midp(X3,sK26,sK25) | ~midp(X3,X1,X2)) )),
  inference(subsumption_resolution,[],[f1624,f587])).
fof(f587,plain,(
  ( ! [X14,X15] : (para(X14,sK26,X15,sK25) | ~midp(sK31,X14,X15)) )),
  inference(resolution,[],[f423,f492])).
fof(f492,plain,(
  midp(sK31,sK26,sK25)),
  inference(resolution,[],[f370,f480])).
fof(f480,plain,(
  midp(sK31,sK25,sK26)),
  inference(cnf_transformation,[],[f359])).
fof(f1624,plain,(
  ( ! [X2,X3,X1] : (~midp(sK31,X1,X2) | midp(X3,sK26,sK25) | ~para(X1,sK26,X2,sK25) | ~midp(X3,X1,X2)) )),
  inference(resolution,[],[f588,f424])).
fof(f588,plain,(
  ( ! [X17,X16] : (para(X16,sK25,X17,sK26) | ~midp(sK31,X16,X17)) )),
  inference(resolution,[],[f423,f480])).
fof(f18014,plain,(
  ~midp(sK25,sK26,sK25) | (~spl33_1122 | ~spl33_1287)),
  inference(subsumption_resolution,[],[f16435,f13239])).
fof(f16435,plain,(
  ~midp(sK25,sK25,sK23) | ~midp(sK25,sK26,sK25) | ~spl33_1287),
  inference(resolution,[],[f14665,f2197])).
fof(f14665,plain,(
  ~circle(sK25,sK25,sK26,sK23) | ~spl33_1287),
  inference(avatar_component_clause,[],[f14664])).
fof(f14664,plain,(
  spl33_1287 <=> ~circle(sK25,sK25,sK26,sK23)),
  introduced(avatar_definition,[new_symbols(naming,[spl33_1287])])).
fof(f17992,plain,(
  ~spl33_176 | spl33_1133),
  inference(avatar_contradiction_clause,[],[f17980])).
fof(f17980,plain,(
  $false | (~spl33_176 | ~spl33_1133)),
  inference(resolution,[],[f17283,f13380])).
fof(f13380,plain,(
  ~midp(sK25,sK26,sK25) | ~spl33_1133),
  inference(avatar_component_clause,[],[f13379])).
fof(f13379,plain,(
  spl33_1133 <=> ~midp(sK25,sK26,sK25)),
  introduced(avatar_definition,[new_symbols(naming,[spl33_1133])])).
fof(f17288,plain,(
  spl33_1122 | ~spl33_44 | ~spl33_176),
  inference(avatar_split_clause,[],[f17240,f2354,f1028,f13238])).
fof(f1028,plain,(
  spl33_44 <=> para(sK26,sK25,sK25,sK23)),
  introduced(avatar_definition,[new_symbols(naming,[spl33_44])])).
fof(f17240,plain,(
  midp(sK25,sK25,sK23) | (~spl33_44 | ~spl33_176)),
  inference(resolution,[],[f17146,f13685])).
fof(f13685,plain,(
  ( ! [X2] : (~midp(sK26,X2,sK25) | midp(sK25,X2,sK23)) ) | (~spl33_44 | ~spl33_176)),
  inference(subsumption_resolution,[],[f13679,f5598])).
fof(f5598,plain,(
  ( ! [X2,X0,X1] : (coll(X0,X1,X2)) ) | ~spl33_176),
  inference(subsumption_resolution,[],[f5595,f5566])).
fof(f5566,plain,(
  ( ! [X6,X5] : (coll(X5,X5,X6)) ) | ~spl33_176),
  inference(resolution,[],[f5383,f360])).
fof(f360,plain,(
  ( ! [X2,X0,X1] : (~coll(X0,X1,X2) | coll(X0,X2,X1)) )),
  inference(cnf_transformation,[],[f169])).
fof(f169,plain,(
  ! [X0,X1,X2] : (coll(X0,X2,X1) | ~coll(X0,X1,X2))),
  inference(ennf_transformation,[],[f1])).
fof(f1,axiom,(
  ! [X0,X1,X2] : (coll(X0,X1,X2) => coll(X0,X2,X1))),
  file('./Other/Auxiliary/aux2/89.p',ruleD1)).
fof(f5383,plain,(
  ( ! [X4,X3] : (coll(X3,X4,X3)) ) | ~spl33_176),
  inference(resolution,[],[f5337,f361])).
fof(f5595,plain,(
  ( ! [X2,X0,X1] : (coll(X0,X1,X2) | ~coll(X2,X2,X0)) ) | ~spl33_176),
  inference(resolution,[],[f5566,f362])).
fof(f362,plain,(
  ( ! [X2,X0,X3,X1] : (~coll(X0,X1,X3) | coll(X2,X3,X0) | ~coll(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f172])).
fof(f172,plain,(
  ! [X0,X1,X2,X3] : (coll(X2,X3,X0) | ~coll(X0,X1,X3) | ~coll(X0,X1,X2))),
  inference(flattening,[],[f171])).
fof(f171,plain,(
  ! [X0,X1,X2,X3] : (coll(X2,X3,X0) | (~coll(X0,X1,X3) | ~coll(X0,X1,X2)))),
  inference(ennf_transformation,[],[f3])).
fof(f3,axiom,(
  ! [X0,X1,X2,X3] : ((coll(X0,X1,X3) & coll(X0,X1,X2)) => coll(X2,X3,X0))),
  file('./Other/Auxiliary/aux2/89.p',ruleD3)).
fof(f13679,plain,(
  ( ! [X2] : (~coll(sK25,X2,sK23) | midp(sK25,X2,sK23) | ~midp(sK26,X2,sK25)) ) | ~spl33_44),
  inference(resolution,[],[f1029,f405])).
fof(f405,plain,(
  ( ! [X4,X2,X0,X3,X1] : (~para(X3,X4,X1,X2) | ~coll(X4,X0,X2) | midp(X4,X0,X2) | ~midp(X3,X0,X1)) )),
  inference(cnf_transformation,[],[f231])).
fof(f231,plain,(
  ! [X0,X1,X2,X3,X4] : (midp(X4,X0,X2) | ~coll(X4,X0,X2) | ~para(X3,X4,X1,X2) | ~midp(X3,X0,X1))),
  inference(flattening,[],[f230])).
fof(f230,plain,(
  ! [X0,X1,X2,X3,X4] : (midp(X4,X0,X2) | (~coll(X4,X0,X2) | ~para(X3,X4,X1,X2) | ~midp(X3,X0,X1)))),
  inference(ennf_transformation,[],[f125])).
fof(f125,plain,(
  ! [X0,X1,X2,X3,X4] : ((coll(X4,X0,X2) & para(X3,X4,X1,X2) & midp(X3,X0,X1)) => midp(X4,X0,X2))),
  inference(rectify,[],[f46])).
fof(f46,axiom,(
  ! [X0,X1,X2,X4,X5] : ((coll(X5,X0,X2) & para(X4,X5,X1,X2) & midp(X4,X0,X1)) => midp(X5,X0,X2))),
  file('./Other/Auxiliary/aux2/89.p',ruleD45)).
fof(f1029,plain,(
  para(sK26,sK25,sK25,sK23) | ~spl33_44),
  inference(avatar_component_clause,[],[f1028])).
fof(f15081,plain,(
  spl33_1166 | ~spl33_36 | ~spl33_176 | ~spl33_634),
  inference(avatar_split_clause,[],[f15079,f6396,f2354,f1002,f13624])).
fof(f1002,plain,(
  spl33_36 <=> para(sK26,sK23,sK23,sK24)),
  introduced(avatar_definition,[new_symbols(naming,[spl33_36])])).
fof(f6396,plain,(
  spl33_634 <=> midp(sK26,sK23,sK23)),
  introduced(avatar_definition,[new_symbols(naming,[spl33_634])])).
fof(f15079,plain,(
  midp(sK23,sK23,sK24) | (~spl33_36 | ~spl33_176 | ~spl33_634)),
  inference(resolution,[],[f14829,f6397])).
fof(f6397,plain,(
  midp(sK26,sK23,sK23) | ~spl33_634),
  inference(avatar_component_clause,[],[f6396])).
fof(f14829,plain,(
  ( ! [X4] : (~midp(sK26,X4,sK23) | midp(sK23,X4,sK24)) ) | (~spl33_36 | ~spl33_176)),
  inference(subsumption_resolution,[],[f14822,f5598])).
fof(f14822,plain,(
  ( ! [X4] : (~coll(sK23,X4,sK24) | midp(sK23,X4,sK24) | ~midp(sK26,X4,sK23)) ) | ~spl33_36),
  inference(resolution,[],[f1003,f405])).
fof(f1003,plain,(
  para(sK26,sK23,sK23,sK24) | ~spl33_36),
  inference(avatar_component_clause,[],[f1002])).
fof(f14790,plain,(
  spl33_36 | ~spl33_264 | ~spl33_448),
  inference(avatar_split_clause,[],[f14789,f4832,f3181,f1002])).
fof(f3181,plain,(
  spl33_264 <=> cong(sK26,sK23,sK26,sK24)),
  introduced(avatar_definition,[new_symbols(naming,[spl33_264])])).
fof(f4832,plain,(
  spl33_448 <=> para(sK23,sK24,sK26,sK24)),
  introduced(avatar_definition,[new_symbols(naming,[spl33_448])])).
fof(f14789,plain,(
  para(sK26,sK23,sK23,sK24) | (~spl33_264 | ~spl33_448)),
  inference(subsumption_resolution,[],[f14771,f4833])).
fof(f4833,plain,(
  para(sK23,sK24,sK26,sK24) | ~spl33_448),
  inference(avatar_component_clause,[],[f4832])).
fof(f14771,plain,(
  para(sK26,sK23,sK23,sK24) | ~para(sK23,sK24,sK26,sK24) | ~spl33_264),
  inference(resolution,[],[f3182,f996])).
fof(f996,plain,(
  ( ! [X21,X19,X20] : (~cong(X21,X19,X21,X20) | para(X21,X19,X19,X20) | ~para(X19,X20,X21,X20)) )),
  inference(resolution,[],[f433,f406])).
fof(f406,plain,(
  ( ! [X2,X0,X1] : (eqangle(X2,X0,X0,X1,X0,X1,X2,X1) | ~cong(X2,X0,X2,X1)) )),
  inference(cnf_transformation,[],[f232])).
fof(f232,plain,(
  ! [X0,X1,X2] : (eqangle(X2,X0,X0,X1,X0,X1,X2,X1) | ~cong(X2,X0,X2,X1))),
  inference(ennf_transformation,[],[f126])).
fof(f126,plain,(
  ! [X0,X1,X2] : (cong(X2,X0,X2,X1) => eqangle(X2,X0,X0,X1,X0,X1,X2,X1))),
  inference(rectify,[],[f47])).
fof(f47,axiom,(
  ! [X0,X1,X7] : (cong(X7,X0,X7,X1) => eqangle(X7,X0,X0,X1,X0,X1,X7,X1))),
  file('./Other/Auxiliary/aux2/89.p',ruleD46)).
fof(f3182,plain,(
  cong(sK26,sK23,sK26,sK24) | ~spl33_264),
  inference(avatar_component_clause,[],[f3181])).
fof(f14743,plain,(
  ~spl33_1301 | spl33_535),
  inference(avatar_split_clause,[],[f11031,f5438,f14737])).
fof(f5438,plain,(
  spl33_535 <=> ~circle(sK23,sK26,sK24,sK24)),
  introduced(avatar_definition,[new_symbols(naming,[spl33_535])])).
fof(f11031,plain,(
  ~midp(sK23,sK26,sK24) | ~spl33_535),
  inference(duplicate_literal_removal,[],[f11011])).
fof(f11011,plain,(
  ~midp(sK23,sK26,sK24) | ~midp(sK23,sK26,sK24) | ~spl33_535),
  inference(resolution,[],[f2198,f5439])).
fof(f5439,plain,(
  ~circle(sK23,sK26,sK24,sK24) | ~spl33_535),
  inference(avatar_component_clause,[],[f5438])).
fof(f14732,plain,(
  spl33_264 | ~spl33_546),
  inference(avatar_split_clause,[],[f14706,f5518,f3181])).
fof(f5518,plain,(
  spl33_546 <=> cong(sK26,sK24,sK26,sK23)),
  introduced(avatar_definition,[new_symbols(naming,[spl33_546])])).
fof(f14706,plain,(
  cong(sK26,sK23,sK26,sK24) | ~spl33_546),
  inference(resolution,[],[f5519,f383])).
fof(f5519,plain,(
  cong(sK26,sK24,sK26,sK23) | ~spl33_546),
  inference(avatar_component_clause,[],[f5518])).
fof(f14444,plain,(
  spl33_448 | ~spl33_110),
  inference(avatar_split_clause,[],[f14441,f1749,f4832])).
fof(f1749,plain,(
  spl33_110 <=> para(sK24,sK23,sK26,sK24)),
  introduced(avatar_definition,[new_symbols(naming,[spl33_110])])).
fof(f14441,plain,(
  para(sK23,sK24,sK26,sK24) | ~spl33_110),
  inference(resolution,[],[f1750,f2254])).
fof(f2254,plain,(
  ( ! [X2,X0,X3,X1] : (~para(X0,X1,X2,X3) | para(X1,X0,X2,X3)) )),
  inference(resolution,[],[f740,f398])).
fof(f740,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (eqangle(X0,X1,X2,X3,X4,X5,X2,X3) | ~para(X1,X0,X4,X5)) )),
  inference(resolution,[],[f377,f399])).
fof(f1750,plain,(
  para(sK24,sK23,sK26,sK24) | ~spl33_110),
  inference(avatar_component_clause,[],[f1749])).
fof(f14213,plain,(
  spl33_110 | ~spl33_578),
  inference(avatar_split_clause,[],[f14197,f5980,f1749])).
fof(f5980,plain,(
  spl33_578 <=> para(sK24,sK23,sK24,sK26)),
  introduced(avatar_definition,[new_symbols(naming,[spl33_578])])).
fof(f14197,plain,(
  para(sK24,sK23,sK26,sK24) | ~spl33_578),
  inference(resolution,[],[f5981,f363])).
fof(f363,plain,(
  ( ! [X2,X0,X3,X1] : (~para(X0,X1,X2,X3) | para(X0,X1,X3,X2)) )),
  inference(cnf_transformation,[],[f173])).
fof(f173,plain,(
  ! [X0,X1,X2,X3] : (para(X0,X1,X3,X2) | ~para(X0,X1,X2,X3))),
  inference(ennf_transformation,[],[f4])).
fof(f4,axiom,(
  ! [X0,X1,X2,X3] : (para(X0,X1,X2,X3) => para(X0,X1,X3,X2))),
  file('./Other/Auxiliary/aux2/89.p',ruleD4)).
fof(f5981,plain,(
  para(sK24,sK23,sK24,sK26) | ~spl33_578),
  inference(avatar_component_clause,[],[f5980])).
fof(f13724,plain,(
  spl33_578 | ~spl33_132 | ~spl33_580),
  inference(avatar_split_clause,[],[f13723,f5989,f1896,f5980])).
fof(f1896,plain,(
  spl33_132 <=> cyclic(sK23,sK23,sK26,sK24)),
  introduced(avatar_definition,[new_symbols(naming,[spl33_132])])).
fof(f5989,plain,(
  spl33_580 <=> para(sK23,sK23,sK23,sK26)),
  introduced(avatar_definition,[new_symbols(naming,[spl33_580])])).
fof(f13723,plain,(
  para(sK24,sK23,sK24,sK26) | (~spl33_132 | ~spl33_580)),
  inference(subsumption_resolution,[],[f8619,f5990])).
fof(f5990,plain,(
  para(sK23,sK23,sK23,sK26) | ~spl33_580),
  inference(avatar_component_clause,[],[f5989])).
fof(f8619,plain,(
  para(sK24,sK23,sK24,sK26) | ~para(sK23,sK23,sK23,sK26) | ~spl33_132),
  inference(resolution,[],[f5947,f994])).
fof(f994,plain,(
  ( ! [X12,X10,X13,X11] : (~cyclic(X11,X12,X13,X10) | para(X13,X11,X13,X12) | ~para(X10,X11,X10,X12)) )),
  inference(resolution,[],[f433,f400])).
fof(f5947,plain,(
  cyclic(sK23,sK26,sK24,sK23) | ~spl33_132),
  inference(resolution,[],[f4791,f373])).
fof(f373,plain,(
  ( ! [X2,X0,X3,X1] : (~cyclic(X0,X1,X2,X3) | cyclic(X0,X1,X3,X2)) )),
  inference(cnf_transformation,[],[f188])).
fof(f188,plain,(
  ! [X0,X1,X2,X3] : (cyclic(X0,X1,X3,X2) | ~cyclic(X0,X1,X2,X3))),
  inference(ennf_transformation,[],[f14])).
fof(f14,axiom,(
  ! [X0,X1,X2,X3] : (cyclic(X0,X1,X2,X3) => cyclic(X0,X1,X3,X2))),
  file('./Other/Auxiliary/aux2/89.p',ruleD14)).
fof(f4791,plain,(
  cyclic(sK23,sK26,sK23,sK24) | ~spl33_132),
  inference(resolution,[],[f1897,f374])).
fof(f1897,plain,(
  cyclic(sK23,sK23,sK26,sK24) | ~spl33_132),
  inference(avatar_component_clause,[],[f1896])).
fof(f13709,plain,(
  spl33_612 | ~spl33_118),
  inference(avatar_split_clause,[],[f13708,f1786,f6269])).
fof(f6269,plain,(
  spl33_612 <=> para(sK25,sK23,sK25,sK26)),
  introduced(avatar_definition,[new_symbols(naming,[spl33_612])])).
fof(f1786,plain,(
  spl33_118 <=> para(sK25,sK24,sK26,sK25)),
  introduced(avatar_definition,[new_symbols(naming,[spl33_118])])).
fof(f13708,plain,(
  para(sK25,sK23,sK25,sK26) | ~spl33_118),
  inference(subsumption_resolution,[],[f9302,f1787])).
fof(f1787,plain,(
  para(sK25,sK24,sK26,sK25) | ~spl33_118),
  inference(avatar_component_clause,[],[f1786])).
fof(f9302,plain,(
  ~para(sK25,sK24,sK26,sK25) | para(sK25,sK23,sK25,sK26)),
  inference(resolution,[],[f1688,f433])).
fof(f1688,plain,(
  eqangle(sK25,sK23,sK25,sK26,sK25,sK24,sK26,sK25)),
  inference(resolution,[],[f739,f378])).
fof(f739,plain,(
  eqangle(sK25,sK26,sK25,sK23,sK26,sK25,sK25,sK24)),
  inference(resolution,[],[f377,f475])).
fof(f475,plain,(
  eqangle(sK26,sK25,sK25,sK23,sK26,sK25,sK25,sK24)),
  inference(cnf_transformation,[],[f359])).
fof(f13701,plain,(
  spl33_118 | ~spl33_117),
  inference(avatar_split_clause,[],[f9447,f1780,f1786])).
fof(f1780,plain,(
  spl33_117 <=> ~para(sK25,sK23,sK26,sK25)),
  introduced(avatar_definition,[new_symbols(naming,[spl33_117])])).
fof(f9447,plain,(
  ~para(sK25,sK23,sK26,sK25) | para(sK25,sK24,sK26,sK25)),
  inference(resolution,[],[f1774,f433])).
fof(f1774,plain,(
  eqangle(sK25,sK24,sK26,sK25,sK25,sK23,sK26,sK25)),
  inference(resolution,[],[f749,f379])).
fof(f749,plain,(
  eqangle(sK25,sK23,sK26,sK25,sK25,sK24,sK26,sK25)),
  inference(resolution,[],[f378,f475])).
fof(f13696,plain,(
  spl33_116 | ~spl33_44),
  inference(avatar_split_clause,[],[f13677,f1028,f1783])).
fof(f1783,plain,(
  spl33_116 <=> para(sK25,sK23,sK26,sK25)),
  introduced(avatar_definition,[new_symbols(naming,[spl33_116])])).
fof(f13677,plain,(
  para(sK25,sK23,sK26,sK25) | ~spl33_44),
  inference(resolution,[],[f1029,f364])).
fof(f13586,plain,(
  spl33_602 | ~spl33_46),
  inference(avatar_split_clause,[],[f13585,f1031,f6221])).
fof(f6221,plain,(
  spl33_602 <=> para(sK25,sK26,sK25,sK23)),
  introduced(avatar_definition,[new_symbols(naming,[spl33_602])])).
fof(f1031,plain,(
  spl33_46 <=> para(sK26,sK25,sK25,sK24)),
  introduced(avatar_definition,[new_symbols(naming,[spl33_46])])).
fof(f13585,plain,(
  para(sK25,sK26,sK25,sK23) | ~spl33_46),
  inference(subsumption_resolution,[],[f1692,f1032])).
fof(f1032,plain,(
  para(sK26,sK25,sK25,sK24) | ~spl33_46),
  inference(avatar_component_clause,[],[f1031])).
fof(f1692,plain,(
  ~para(sK26,sK25,sK25,sK24) | para(sK25,sK26,sK25,sK23)),
  inference(resolution,[],[f739,f433])).
fof(f13564,plain,(
  spl33_44 | ~spl33_56 | ~spl33_176 | ~spl33_182),
  inference(avatar_split_clause,[],[f13556,f2537,f2354,f1076,f1028])).
fof(f1076,plain,(
  spl33_56 <=> perp(sK26,sK25,sK25,sK23)),
  introduced(avatar_definition,[new_symbols(naming,[spl33_56])])).
fof(f2537,plain,(
  spl33_182 <=> ! [X44,X46,X45,X47] : (~para(X44,X45,X46,X47) | perp(X44,X45,X46,X47))),
  introduced(avatar_definition,[new_symbols(naming,[spl33_182])])).
fof(f13556,plain,(
  para(sK26,sK25,sK25,sK23) | (~spl33_56 | ~spl33_176 | ~spl33_182)),
  inference(resolution,[],[f13407,f11269])).
fof(f11269,plain,(
  ( ! [X0,X1] : (perp(X0,X1,X0,X1)) ) | (~spl33_176 | ~spl33_182)),
  inference(resolution,[],[f2538,f2355])).
fof(f2538,plain,(
  ( ! [X47,X45,X46,X44] : (~para(X44,X45,X46,X47) | perp(X44,X45,X46,X47)) ) | ~spl33_182),
  inference(avatar_component_clause,[],[f2537])).
fof(f13407,plain,(
  ( ! [X10,X9] : (~perp(X9,X10,sK26,sK25) | para(X9,X10,sK25,sK23)) ) | ~spl33_56),
  inference(resolution,[],[f1077,f368])).
fof(f368,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (~perp(X2,X3,X4,X5) | para(X0,X1,X4,X5) | ~perp(X0,X1,X2,X3)) )),
  inference(cnf_transformation,[],[f180])).
fof(f180,plain,(
  ! [X0,X1,X2,X3,X4,X5] : (para(X0,X1,X4,X5) | ~perp(X2,X3,X4,X5) | ~perp(X0,X1,X2,X3))),
  inference(flattening,[],[f179])).
fof(f179,plain,(
  ! [X0,X1,X2,X3,X4,X5] : (para(X0,X1,X4,X5) | (~perp(X2,X3,X4,X5) | ~perp(X0,X1,X2,X3)))),
  inference(ennf_transformation,[],[f9])).
fof(f9,axiom,(
  ! [X0,X1,X2,X3,X4,X5] : ((perp(X2,X3,X4,X5) & perp(X0,X1,X2,X3)) => para(X0,X1,X4,X5))),
  file('./Other/Auxiliary/aux2/89.p',ruleD9)).
fof(f1077,plain,(
  perp(sK26,sK25,sK25,sK23) | ~spl33_56),
  inference(avatar_component_clause,[],[f1076])).
fof(f13381,plain,(
  ~spl33_1133 | ~spl33_1123 | ~spl33_176 | ~spl33_256 | spl33_547),
  inference(avatar_split_clause,[],[f13230,f5521,f3071,f2354,f13241,f13379])).
fof(f13241,plain,(
  spl33_1123 <=> ~midp(sK25,sK25,sK23)),
  introduced(avatar_definition,[new_symbols(naming,[spl33_1123])])).
fof(f3071,plain,(
  spl33_256 <=> ! [X13,X15,X14] : (~midp(X13,X14,X15) | cong(X14,X13,X14,X15))),
  introduced(avatar_definition,[new_symbols(naming,[spl33_256])])).
fof(f5521,plain,(
  spl33_547 <=> ~cong(sK26,sK24,sK26,sK23)),
  introduced(avatar_definition,[new_symbols(naming,[spl33_547])])).
fof(f13230,plain,(
  ~midp(sK25,sK25,sK23) | ~midp(sK25,sK26,sK25) | (~spl33_176 | ~spl33_256 | ~spl33_547)),
  inference(resolution,[],[f13085,f2197])).
fof(f13085,plain,(
  ~circle(sK25,sK25,sK26,sK23) | (~spl33_176 | ~spl33_256 | ~spl33_547)),
  inference(subsumption_resolution,[],[f13084,f5598])).
fof(f13084,plain,(
  ~coll(sK24,sK26,sK23) | ~circle(sK25,sK25,sK26,sK23) | (~spl33_256 | ~spl33_547)),
  inference(subsumption_resolution,[],[f13083,f10125])).
fof(f10125,plain,(
  ~midp(sK24,sK26,sK23) | (~spl33_256 | ~spl33_547)),
  inference(resolution,[],[f3072,f5522])).
fof(f5522,plain,(
  ~cong(sK26,sK24,sK26,sK23) | ~spl33_547),
  inference(avatar_component_clause,[],[f5521])).
fof(f3072,plain,(
  ( ! [X14,X15,X13] : (cong(X14,X13,X14,X15) | ~midp(X13,X14,X15)) ) | ~spl33_256),
  inference(avatar_component_clause,[],[f3071])).
fof(f13083,plain,(
  midp(sK24,sK26,sK23) | ~coll(sK24,sK26,sK23) | ~circle(sK25,sK25,sK26,sK23)),
  inference(subsumption_resolution,[],[f13077,f2213])).
fof(f2213,plain,(
  para(sK26,sK25,sK25,sK26)),
  inference(resolution,[],[f1976,f363])).
fof(f1976,plain,(
  para(sK26,sK25,sK26,sK25)),
  inference(subsumption_resolution,[],[f1973,f1771])).
fof(f1771,plain,(
  para(sK25,sK23,sK25,sK24)),
  inference(resolution,[],[f749,f398])).
fof(f1973,plain,(
  ~para(sK25,sK23,sK25,sK24) | para(sK26,sK25,sK26,sK25)),
  inference(resolution,[],[f769,f433])).
fof(f769,plain,(
  eqangle(sK26,sK25,sK26,sK25,sK25,sK23,sK25,sK24)),
  inference(resolution,[],[f380,f475])).
fof(f13077,plain,(
  ~para(sK26,sK25,sK25,sK26) | midp(sK24,sK26,sK23) | ~coll(sK24,sK26,sK23) | ~circle(sK25,sK25,sK26,sK23)),
  inference(resolution,[],[f5325,f411])).
fof(f5325,plain,(
  ( ! [X204,X205] : (eqangle(sK25,sK26,sK25,sK23,X204,X205,sK25,sK24) | ~para(sK26,sK25,X204,X205)) )),
  inference(resolution,[],[f1357,f739])).
fof(f13374,plain,(
  spl33_496 | ~spl33_156 | ~spl33_176),
  inference(avatar_split_clause,[],[f13373,f2354,f1991,f5105])).
fof(f5105,plain,(
  spl33_496 <=> cong(sK25,sK26,sK25,sK26)),
  introduced(avatar_definition,[new_symbols(naming,[spl33_496])])).
fof(f1991,plain,(
  spl33_156 <=> cyclic(sK25,sK25,sK26,sK23)),
  introduced(avatar_definition,[new_symbols(naming,[spl33_156])])).
fof(f13373,plain,(
  cong(sK25,sK26,sK25,sK26) | (~spl33_156 | ~spl33_176)),
  inference(subsumption_resolution,[],[f13372,f7109])).
fof(f7109,plain,(
  cyclic(sK25,sK26,sK23,sK25) | ~spl33_156),
  inference(resolution,[],[f5092,f373])).
fof(f5092,plain,(
  cyclic(sK25,sK26,sK25,sK23) | ~spl33_156),
  inference(resolution,[],[f1992,f374])).
fof(f1992,plain,(
  cyclic(sK25,sK25,sK26,sK23) | ~spl33_156),
  inference(avatar_component_clause,[],[f1991])).
fof(f13372,plain,(
  cong(sK25,sK26,sK25,sK26) | ~cyclic(sK25,sK26,sK23,sK25) | (~spl33_156 | ~spl33_176)),
  inference(subsumption_resolution,[],[f13342,f10791])).
fof(f13342,plain,(
  cong(sK25,sK26,sK25,sK26) | ~cyclic(sK25,sK26,sK23,sK26) | ~cyclic(sK25,sK26,sK23,sK25) | ~spl33_156),
  inference(resolution,[],[f7109,f1455])).
fof(f1455,plain,(
  ( ! [X10,X8,X7,X9] : (~cyclic(X7,X8,X9,X10) | cong(X7,X8,X7,X8) | ~cyclic(X7,X8,X9,X8) | ~cyclic(X7,X8,X9,X7)) )),
  inference(duplicate_literal_removal,[],[f1449])).
fof(f1449,plain,(
  ( ! [X10,X8,X7,X9] : (cong(X7,X8,X7,X8) | ~cyclic(X7,X8,X9,X10) | ~cyclic(X7,X8,X9,X8) | ~cyclic(X7,X8,X9,X7) | ~cyclic(X7,X8,X9,X10)) )),
  inference(resolution,[],[f403,f400])).
fof(f13063,plain,(
  spl33_498 | ~spl33_156 | ~spl33_176),
  inference(avatar_split_clause,[],[f13062,f2354,f1991,f5111])).
fof(f5111,plain,(
  spl33_498 <=> cong(sK25,sK23,sK25,sK23)),
  introduced(avatar_definition,[new_symbols(naming,[spl33_498])])).
fof(f13062,plain,(
  cong(sK25,sK23,sK25,sK23) | (~spl33_156 | ~spl33_176)),
  inference(subsumption_resolution,[],[f13061,f10767])).
fof(f13061,plain,(
  cong(sK25,sK23,sK25,sK23) | ~cyclic(sK25,sK23,sK25,sK25) | (~spl33_156 | ~spl33_176)),
  inference(subsumption_resolution,[],[f13022,f10767])).
fof(f13022,plain,(
  cong(sK25,sK23,sK25,sK23) | ~cyclic(sK25,sK23,sK25,sK23) | ~cyclic(sK25,sK23,sK25,sK25) | ~spl33_156),
  inference(resolution,[],[f7052,f1455])).
fof(f7052,plain,(
  cyclic(sK25,sK23,sK25,sK26) | ~spl33_156),
  inference(resolution,[],[f5091,f374])).
fof(f5091,plain,(
  cyclic(sK25,sK25,sK23,sK26) | ~spl33_156),
  inference(resolution,[],[f1992,f373])).
fof(f11203,plain,(
  ~spl33_176 | ~spl33_180 | ~spl33_434),
  inference(avatar_contradiction_clause,[],[f11202])).
fof(f11202,plain,(
  $false | (~spl33_176 | ~spl33_180 | ~spl33_434)),
  inference(subsumption_resolution,[],[f11201,f2535])).
fof(f2535,plain,(
  ( ! [X48,X49] : (~perp(X48,X49,X48,X49)) ) | ~spl33_180),
  inference(avatar_component_clause,[],[f2534])).
fof(f2534,plain,(
  spl33_180 <=> ! [X49,X48] : ~perp(X48,X49,X48,X49)),
  introduced(avatar_definition,[new_symbols(naming,[spl33_180])])).
fof(f11201,plain,(
  perp(sK24,sK24,sK24,sK24) | (~spl33_176 | ~spl33_434)),
  inference(subsumption_resolution,[],[f11184,f5383])).
fof(f11184,plain,(
  ~coll(sK24,sK24,sK24) | perp(sK24,sK24,sK24,sK24) | ~spl33_434),
  inference(resolution,[],[f4715,f413])).
fof(f413,plain,(
  ( ! [X2,X0,X3,X1] : (~circle(X3,X0,X1,X2) | ~coll(X3,X0,X2) | perp(X0,X1,X1,X2)) )),
  inference(cnf_transformation,[],[f246])).
fof(f246,plain,(
  ! [X0,X1,X2,X3] : (perp(X0,X1,X1,X2) | ~coll(X3,X0,X2) | ~circle(X3,X0,X1,X2))),
  inference(flattening,[],[f245])).
fof(f245,plain,(
  ! [X0,X1,X2,X3] : (perp(X0,X1,X1,X2) | (~coll(X3,X0,X2) | ~circle(X3,X0,X1,X2)))),
  inference(ennf_transformation,[],[f133])).
fof(f133,plain,(
  ! [X0,X1,X2,X3] : ((coll(X3,X0,X2) & circle(X3,X0,X1,X2)) => perp(X0,X1,X1,X2))),
  inference(rectify,[],[f54])).
fof(f54,axiom,(
  ! [X0,X1,X2,X7] : ((coll(X7,X0,X2) & circle(X7,X0,X1,X2)) => perp(X0,X1,X1,X2))),
  file('./Other/Auxiliary/aux2/89.p',ruleD53)).
fof(f4715,plain,(
  circle(sK24,sK24,sK24,sK24) | ~spl33_434),
  inference(avatar_component_clause,[],[f4714])).
fof(f4714,plain,(
  spl33_434 <=> circle(sK24,sK24,sK24,sK24)),
  introduced(avatar_definition,[new_symbols(naming,[spl33_434])])).
fof(f10963,plain,(
  spl33_435 | ~spl33_530),
  inference(avatar_contradiction_clause,[],[f10962])).
fof(f10962,plain,(
  $false | (~spl33_435 | ~spl33_530)),
  inference(subsumption_resolution,[],[f10955,f5423])).
fof(f5423,plain,(
  midp(sK24,sK24,sK24) | ~spl33_530),
  inference(avatar_component_clause,[],[f5422])).
fof(f5422,plain,(
  spl33_530 <=> midp(sK24,sK24,sK24)),
  introduced(avatar_definition,[new_symbols(naming,[spl33_530])])).
fof(f10955,plain,(
  ~midp(sK24,sK24,sK24) | ~spl33_435),
  inference(duplicate_literal_removal,[],[f10944])).
fof(f10944,plain,(
  ~midp(sK24,sK24,sK24) | ~midp(sK24,sK24,sK24) | ~spl33_435),
  inference(resolution,[],[f2197,f4718])).
fof(f4718,plain,(
  ~circle(sK24,sK24,sK24,sK24) | ~spl33_435),
  inference(avatar_component_clause,[],[f4717])).
fof(f4717,plain,(
  spl33_435 <=> ~circle(sK24,sK24,sK24,sK24)),
  introduced(avatar_definition,[new_symbols(naming,[spl33_435])])).
fof(f7205,plain,(
  spl33_634 | ~spl33_603 | ~spl33_753 | ~spl33_156),
  inference(avatar_split_clause,[],[f7198,f1991,f7203,f6224,f6396])).
fof(f6224,plain,(
  spl33_603 <=> ~para(sK25,sK26,sK25,sK23)),
  introduced(avatar_definition,[new_symbols(naming,[spl33_603])])).
fof(f7198,plain,(
  ~circle(sK25,sK25,sK23,sK23) | ~para(sK25,sK26,sK25,sK23) | midp(sK26,sK23,sK23) | ~spl33_156),
  inference(subsumption_resolution,[],[f7122,f2108])).
fof(f7122,plain,(
  ~coll(sK26,sK23,sK23) | ~circle(sK25,sK25,sK23,sK23) | ~para(sK25,sK26,sK25,sK23) | midp(sK26,sK23,sK23) | ~spl33_156),
  inference(resolution,[],[f5092,f1203])).
fof(f1203,plain,(
  ( ! [X4,X5,X3] : (~cyclic(X5,X3,X5,X4) | ~coll(X3,X4,X4) | ~circle(X5,X5,X4,X4) | ~para(X5,X3,X5,X4) | midp(X3,X4,X4)) )),
  inference(resolution,[],[f411,f414])).
fof(f414,plain,(
  ( ! [X2,X0,X3,X1] : (eqangle(X0,X3,X2,X3,X2,X3,X2,X1) | ~para(X0,X1,X2,X3) | ~cyclic(X0,X1,X2,X3)) )),
  inference(cnf_transformation,[],[f248])).
fof(f248,plain,(
  ! [X0,X1,X2,X3] : (eqangle(X0,X3,X2,X3,X2,X3,X2,X1) | ~para(X0,X1,X2,X3) | ~cyclic(X0,X1,X2,X3))),
  inference(flattening,[],[f247])).
fof(f247,plain,(
  ! [X0,X1,X2,X3] : (eqangle(X0,X3,X2,X3,X2,X3,X2,X1) | (~para(X0,X1,X2,X3) | ~cyclic(X0,X1,X2,X3)))),
  inference(ennf_transformation,[],[f55])).
fof(f55,axiom,(
  ! [X0,X1,X2,X3] : ((para(X0,X1,X2,X3) & cyclic(X0,X1,X2,X3)) => eqangle(X0,X3,X2,X3,X2,X3,X2,X1))),
  file('./Other/Auxiliary/aux2/89.p',ruleD54)).
fof(f6274,plain,(
  ~spl33_613 | spl33_580 | ~spl33_134),
  inference(avatar_split_clause,[],[f6261,f1902,f5989,f6272])).
fof(f6272,plain,(
  spl33_613 <=> ~para(sK25,sK23,sK25,sK26)),
  introduced(avatar_definition,[new_symbols(naming,[spl33_613])])).
fof(f1902,plain,(
  spl33_134 <=> cyclic(sK23,sK23,sK26,sK25)),
  introduced(avatar_definition,[new_symbols(naming,[spl33_134])])).
fof(f6261,plain,(
  para(sK23,sK23,sK23,sK26) | ~para(sK25,sK23,sK25,sK26) | ~spl33_134),
  inference(resolution,[],[f4880,f994])).
fof(f4880,plain,(
  cyclic(sK23,sK26,sK23,sK25) | ~spl33_134),
  inference(resolution,[],[f1903,f374])).
fof(f1903,plain,(
  cyclic(sK23,sK23,sK26,sK25) | ~spl33_134),
  inference(avatar_component_clause,[],[f1902])).
fof(f5440,plain,(
  spl33_530 | ~spl33_535 | ~spl33_146 | ~spl33_176),
  inference(avatar_split_clause,[],[f5433,f2354,f1950,f5438,f5422])).
fof(f1950,plain,(
  spl33_146 <=> cyclic(sK24,sK24,sK26,sK23)),
  introduced(avatar_definition,[new_symbols(naming,[spl33_146])])).
fof(f5433,plain,(
  ~circle(sK23,sK26,sK24,sK24) | midp(sK24,sK24,sK24) | (~spl33_146 | ~spl33_176)),
  inference(subsumption_resolution,[],[f5037,f5383])).
fof(f5037,plain,(
  ~coll(sK24,sK24,sK24) | ~circle(sK23,sK26,sK24,sK24) | midp(sK24,sK24,sK24) | ~spl33_146),
  inference(resolution,[],[f1951,f1204])).
fof(f1204,plain,(
  ( ! [X6,X8,X7,X9] : (~cyclic(X7,X6,X9,X8) | ~coll(X6,X7,X6) | ~circle(X8,X9,X7,X6) | midp(X6,X7,X6)) )),
  inference(resolution,[],[f411,f400])).
fof(f1951,plain,(
  cyclic(sK24,sK24,sK26,sK23) | ~spl33_146),
  inference(avatar_component_clause,[],[f1950])).
fof(f5255,plain,(
  ~spl33_178),
  inference(avatar_contradiction_clause,[],[f5244])).
fof(f5244,plain,(
  $false | ~spl33_178),
  inference(resolution,[],[f2358,f2093])).
fof(f2093,plain,(
  para(sK20,sK32,sK20,sK32)),
  inference(resolution,[],[f690,f531])).
fof(f531,plain,(
  perp(sK20,sK32,sK29,sK28)),
  inference(resolution,[],[f367,f486])).
fof(f486,plain,(
  perp(sK29,sK28,sK20,sK32)),
  inference(cnf_transformation,[],[f359])).
fof(f367,plain,(
  ( ! [X2,X0,X3,X1] : (~perp(X0,X1,X2,X3) | perp(X2,X3,X0,X1)) )),
  inference(cnf_transformation,[],[f178])).
fof(f178,plain,(
  ! [X0,X1,X2,X3] : (perp(X2,X3,X0,X1) | ~perp(X0,X1,X2,X3))),
  inference(ennf_transformation,[],[f8])).
fof(f8,axiom,(
  ! [X0,X1,X2,X3] : (perp(X0,X1,X2,X3) => perp(X2,X3,X0,X1))),
  file('./Other/Auxiliary/aux2/89.p',ruleD8)).
fof(f690,plain,(
  ( ! [X4,X5] : (~perp(X4,X5,sK29,sK28) | para(X4,X5,sK20,sK32)) )),
  inference(resolution,[],[f368,f486])).
fof(f2358,plain,(
  ( ! [X0,X1] : (~para(X0,X1,X0,X1)) ) | ~spl33_178),
  inference(avatar_component_clause,[],[f2357])).
fof(f2357,plain,(
  spl33_178 <=> ! [X1,X0] : ~para(X0,X1,X0,X1)),
  introduced(avatar_definition,[new_symbols(naming,[spl33_178])])).
fof(f5116,plain,(
  ~spl33_497 | ~spl33_499 | spl33_57 | ~spl33_156),
  inference(avatar_split_clause,[],[f5103,f1991,f1073,f5114,f5108])).
fof(f5108,plain,(
  spl33_497 <=> ~cong(sK25,sK26,sK25,sK26)),
  introduced(avatar_definition,[new_symbols(naming,[spl33_497])])).
fof(f5114,plain,(
  spl33_499 <=> ~cong(sK25,sK23,sK25,sK23)),
  introduced(avatar_definition,[new_symbols(naming,[spl33_499])])).
fof(f1073,plain,(
  spl33_57 <=> ~perp(sK26,sK25,sK25,sK23)),
  introduced(avatar_definition,[new_symbols(naming,[spl33_57])])).
fof(f5103,plain,(
  ~cong(sK25,sK23,sK25,sK23) | ~cong(sK25,sK26,sK25,sK26) | (~spl33_57 | ~spl33_156)),
  inference(subsumption_resolution,[],[f5095,f1074])).
fof(f1074,plain,(
  ~perp(sK26,sK25,sK25,sK23) | ~spl33_57),
  inference(avatar_component_clause,[],[f1073])).
fof(f5095,plain,(
  perp(sK26,sK25,sK25,sK23) | ~cong(sK25,sK23,sK25,sK23) | ~cong(sK25,sK26,sK25,sK26) | ~spl33_156),
  inference(resolution,[],[f1992,f417])).
fof(f417,plain,(
  ( ! [X2,X0,X3,X1] : (~cyclic(X0,X1,X2,X3) | perp(X2,X0,X0,X3) | ~cong(X0,X3,X1,X3) | ~cong(X0,X2,X1,X2)) )),
  inference(cnf_transformation,[],[f254])).
fof(f254,plain,(
  ! [X0,X1,X2,X3] : (perp(X2,X0,X0,X3) | ~cyclic(X0,X1,X2,X3) | ~cong(X0,X3,X1,X3) | ~cong(X0,X2,X1,X2))),
  inference(flattening,[],[f253])).
fof(f253,plain,(
  ! [X0,X1,X2,X3] : (perp(X2,X0,X0,X3) | (~cyclic(X0,X1,X2,X3) | ~cong(X0,X3,X1,X3) | ~cong(X0,X2,X1,X2)))),
  inference(ennf_transformation,[],[f136])).
fof(f136,plain,(
  ! [X0,X1,X2,X3] : ((cyclic(X0,X1,X2,X3) & cong(X0,X3,X1,X3) & cong(X0,X2,X1,X2)) => perp(X2,X0,X0,X3))),
  inference(rectify,[],[f58])).
fof(f58,axiom,(
  ! [X0,X1,X8,X9] : ((cyclic(X0,X1,X8,X9) & cong(X0,X9,X1,X9) & cong(X0,X8,X1,X8)) => perp(X8,X0,X0,X9))),
  file('./Other/Auxiliary/aux2/89.p',ruleD57)).
fof(f4304,plain,(
  ~spl33_62 | ~spl33_176),
  inference(avatar_contradiction_clause,[],[f4303])).
fof(f4303,plain,(
  $false | (~spl33_62 | ~spl33_176)),
  inference(subsumption_resolution,[],[f4302,f3117])).
fof(f3117,plain,(
  ( ! [X0,X1] : (~midp(X0,X1,X0)) ) | ~spl33_62),
  inference(subsumption_resolution,[],[f3106,f370])).
fof(f3106,plain,(
  ( ! [X0,X1] : (~midp(X0,X0,X1) | ~midp(X0,X1,X0)) ) | ~spl33_62),
  inference(resolution,[],[f1092,f543])).
fof(f1092,plain,(
  ( ! [X6,X8,X7] : (~cong(X6,X7,X7,X8) | ~midp(X6,X7,X8)) ) | ~spl33_62),
  inference(avatar_component_clause,[],[f1091])).
fof(f1091,plain,(
  spl33_62 <=> ! [X7,X8,X6] : (~cong(X6,X7,X7,X8) | ~midp(X6,X7,X8))),
  introduced(avatar_definition,[new_symbols(naming,[spl33_62])])).
fof(f4302,plain,(
  ( ! [X14] : (midp(X14,X14,X14)) ) | ~spl33_176),
  inference(subsumption_resolution,[],[f4284,f3045])).
fof(f3045,plain,(
  ( ! [X4,X3] : (coll(X3,X4,X3)) ) | ~spl33_176),
  inference(resolution,[],[f3032,f361])).
fof(f3032,plain,(
  ( ! [X0,X1] : (coll(X0,X1,X1)) ) | ~spl33_176),
  inference(resolution,[],[f2355,f426])).
fof(f4284,plain,(
  ( ! [X14] : (~coll(X14,X14,X14) | midp(X14,X14,X14)) ) | ~spl33_176),
  inference(resolution,[],[f4274,f427])).
fof(f4274,plain,(
  ( ! [X0] : (cong(X0,X0,X0,X0)) ) | ~spl33_176),
  inference(subsumption_resolution,[],[f4273,f2543])).
fof(f2543,plain,(
  ( ! [X70,X71,X69] : (cyclic(X70,X70,X69,X71)) ) | ~spl33_176),
  inference(subsumption_resolution,[],[f2542,f2541])).
fof(f2541,plain,(
  ( ! [X68,X66,X67] : (~coll(X66,X68,X67) | cyclic(X67,X67,X66,X68)) ) | ~spl33_176),
  inference(subsumption_resolution,[],[f2522,f2355])).
fof(f2542,plain,(
  ( ! [X70,X71,X69] : (coll(X69,X71,X70) | cyclic(X70,X70,X69,X71)) ) | ~spl33_176),
  inference(subsumption_resolution,[],[f2523,f2355])).
fof(f4273,plain,(
  ( ! [X0,X1] : (cong(X0,X0,X0,X0) | ~cyclic(X0,X0,X1,X0)) ) | ~spl33_176),
  inference(duplicate_literal_removal,[],[f4267])).
fof(f4267,plain,(
  ( ! [X0,X1] : (cong(X0,X0,X0,X0) | ~cyclic(X0,X0,X1,X0) | ~cyclic(X0,X0,X1,X0)) ) | ~spl33_176),
  inference(resolution,[],[f1455,f2543])).
fof(f3537,plain,(
  spl33_135 | ~spl33_176),
  inference(avatar_contradiction_clause,[],[f3536])).
fof(f3536,plain,(
  $false | (~spl33_135 | ~spl33_176)),
  inference(subsumption_resolution,[],[f1906,f2543])).
fof(f1906,plain,(
  ~cyclic(sK23,sK23,sK26,sK25) | ~spl33_135),
  inference(avatar_component_clause,[],[f1905])).
fof(f1905,plain,(
  spl33_135 <=> ~cyclic(sK23,sK23,sK26,sK25)),
  introduced(avatar_definition,[new_symbols(naming,[spl33_135])])).
fof(f3535,plain,(
  spl33_147 | ~spl33_176),
  inference(avatar_contradiction_clause,[],[f3534])).
fof(f3534,plain,(
  $false | (~spl33_147 | ~spl33_176)),
  inference(subsumption_resolution,[],[f1954,f2543])).
fof(f1954,plain,(
  ~cyclic(sK24,sK24,sK26,sK23) | ~spl33_147),
  inference(avatar_component_clause,[],[f1953])).
fof(f1953,plain,(
  spl33_147 <=> ~cyclic(sK24,sK24,sK26,sK23)),
  introduced(avatar_definition,[new_symbols(naming,[spl33_147])])).
fof(f3531,plain,(
  spl33_133 | ~spl33_176),
  inference(avatar_contradiction_clause,[],[f3530])).
fof(f3530,plain,(
  $false | (~spl33_133 | ~spl33_176)),
  inference(subsumption_resolution,[],[f1900,f2543])).
fof(f1900,plain,(
  ~cyclic(sK23,sK23,sK26,sK24) | ~spl33_133),
  inference(avatar_component_clause,[],[f1899])).
fof(f1899,plain,(
  spl33_133 <=> ~cyclic(sK23,sK23,sK26,sK24)),
  introduced(avatar_definition,[new_symbols(naming,[spl33_133])])).
fof(f3527,plain,(
  spl33_157 | ~spl33_176),
  inference(avatar_contradiction_clause,[],[f3526])).
fof(f3526,plain,(
  $false | (~spl33_157 | ~spl33_176)),
  inference(subsumption_resolution,[],[f1995,f2543])).
fof(f1995,plain,(
  ~cyclic(sK25,sK25,sK26,sK23) | ~spl33_157),
  inference(avatar_component_clause,[],[f1994])).
fof(f1994,plain,(
  spl33_157 <=> ~cyclic(sK25,sK25,sK26,sK23)),
  introduced(avatar_definition,[new_symbols(naming,[spl33_157])])).
fof(f3077,plain,(
  spl33_256 | spl33_62),
  inference(avatar_split_clause,[],[f3064,f1091,f3071])).
fof(f3064,plain,(
  ( ! [X14,X12,X10,X15,X13,X11] : (~midp(X10,X11,X12) | ~midp(X13,X14,X15) | ~cong(X10,X11,X11,X12) | cong(X14,X13,X14,X15)) )),
  inference(resolution,[],[f778,f435])).
fof(f435,plain,(
  ( ! [X6,X4,X2,X0,X7,X5,X3,X1] : (~eqratio(X0,X1,X2,X3,X4,X5,X6,X7) | ~cong(X4,X5,X6,X7) | cong(X0,X1,X2,X3)) )),
  inference(cnf_transformation,[],[f284])).
fof(f284,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6,X7] : (cong(X0,X1,X2,X3) | ~cong(X4,X5,X6,X7) | ~eqratio(X0,X1,X2,X3,X4,X5,X6,X7))),
  inference(flattening,[],[f283])).
fof(f283,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6,X7] : (cong(X0,X1,X2,X3) | (~cong(X4,X5,X6,X7) | ~eqratio(X0,X1,X2,X3,X4,X5,X6,X7)))),
  inference(ennf_transformation,[],[f148])).
fof(f148,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6,X7] : ((cong(X4,X5,X6,X7) & eqratio(X0,X1,X2,X3,X4,X5,X6,X7)) => cong(X0,X1,X2,X3))),
  inference(rectify,[],[f76])).
fof(f76,axiom,(
  ! [X0,X1,X2,X3,X8,X9,X10,X11] : ((cong(X8,X9,X10,X11) & eqratio(X0,X1,X2,X3,X8,X9,X10,X11)) => cong(X0,X1,X2,X3))),
  file('./Other/Auxiliary/aux2/89.p',ruleD75)).
fof(f778,plain,(
  ( ! [X6,X10,X8,X7,X11,X9] : (eqratio(X6,X7,X6,X8,X9,X10,X10,X11) | ~midp(X9,X10,X11) | ~midp(X7,X6,X8)) )),
  inference(resolution,[],[f385,f430])).
fof(f430,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (eqratio(X4,X0,X0,X1,X5,X2,X2,X3) | ~midp(X5,X2,X3) | ~midp(X4,X0,X1)) )),
  inference(cnf_transformation,[],[f274])).
fof(f274,plain,(
  ! [X0,X1,X2,X3,X4,X5] : (eqratio(X4,X0,X0,X1,X5,X2,X2,X3) | ~midp(X5,X2,X3) | ~midp(X4,X0,X1))),
  inference(flattening,[],[f273])).
fof(f273,plain,(
  ! [X0,X1,X2,X3,X4,X5] : (eqratio(X4,X0,X0,X1,X5,X2,X2,X3) | (~midp(X5,X2,X3) | ~midp(X4,X0,X1)))),
  inference(ennf_transformation,[],[f145])).
fof(f145,plain,(
  ! [X0,X1,X2,X3,X4,X5] : ((midp(X5,X2,X3) & midp(X4,X0,X1)) => eqratio(X4,X0,X0,X1,X5,X2,X2,X3))),
  inference(rectify,[],[f71])).
fof(f71,axiom,(
  ! [X0,X1,X2,X3,X6,X16] : ((midp(X16,X2,X3) & midp(X6,X0,X1)) => eqratio(X6,X0,X0,X1,X16,X2,X2,X3))),
  file('./Other/Auxiliary/aux2/89.p',ruleD70)).
fof(f385,plain,(
  ( ! [X6,X4,X2,X0,X7,X5,X3,X1] : (~eqratio(X0,X1,X2,X3,X4,X5,X6,X7) | eqratio(X1,X0,X2,X3,X4,X5,X6,X7)) )),
  inference(cnf_transformation,[],[f203])).
fof(f203,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6,X7] : (eqratio(X1,X0,X2,X3,X4,X5,X6,X7) | ~eqratio(X0,X1,X2,X3,X4,X5,X6,X7))),
  inference(ennf_transformation,[],[f105])).
fof(f105,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6,X7] : (eqratio(X0,X1,X2,X3,X4,X5,X6,X7) => eqratio(X1,X0,X2,X3,X4,X5,X6,X7))),
  inference(rectify,[],[f26])).
fof(f26,axiom,(
  ! [X0,X1,X2,X3,X8,X9,X10,X11] : (eqratio(X0,X1,X2,X3,X8,X9,X10,X11) => eqratio(X1,X0,X2,X3,X8,X9,X10,X11))),
  file('./Other/Auxiliary/aux2/89.p',ruleD26)).
fof(f2539,plain,(
  spl33_180 | spl33_182),
  inference(avatar_split_clause,[],[f2516,f2537,f2534])).
fof(f2516,plain,(
  ( ! [X47,X45,X48,X46,X44,X49] : (~para(X44,X45,X46,X47) | ~perp(X48,X49,X48,X49) | perp(X44,X45,X46,X47)) )),
  inference(resolution,[],[f770,f434])).
fof(f434,plain,(
  ( ! [X6,X4,X2,X0,X7,X5,X3,X1] : (~eqangle(X0,X1,X2,X3,X4,X5,X6,X7) | ~perp(X4,X5,X6,X7) | perp(X0,X1,X2,X3)) )),
  inference(cnf_transformation,[],[f282])).
fof(f282,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6,X7] : (perp(X0,X1,X2,X3) | ~perp(X4,X5,X6,X7) | ~eqangle(X0,X1,X2,X3,X4,X5,X6,X7))),
  inference(flattening,[],[f281])).
fof(f281,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6,X7] : (perp(X0,X1,X2,X3) | (~perp(X4,X5,X6,X7) | ~eqangle(X0,X1,X2,X3,X4,X5,X6,X7)))),
  inference(ennf_transformation,[],[f147])).
fof(f147,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6,X7] : ((perp(X4,X5,X6,X7) & eqangle(X0,X1,X2,X3,X4,X5,X6,X7)) => perp(X0,X1,X2,X3))),
  inference(rectify,[],[f75])).
fof(f75,axiom,(
  ! [X0,X1,X2,X3,X8,X9,X10,X11] : ((perp(X8,X9,X10,X11) & eqangle(X0,X1,X2,X3,X8,X9,X10,X11)) => perp(X0,X1,X2,X3))),
  file('./Other/Auxiliary/aux2/89.p',ruleD74)).
fof(f2359,plain,(
  spl33_176 | spl33_178),
  inference(avatar_split_clause,[],[f2330,f2357,f2354])).
fof(f2330,plain,(
  ( ! [X2,X0,X3,X1] : (~para(X0,X1,X0,X1) | para(X2,X3,X2,X3)) )),
  inference(resolution,[],[f750,f398])).
fof(f1863,plain,(
  ~spl33_45 | spl33_47),
  inference(avatar_split_clause,[],[f1862,f1034,f1025])).
fof(f1025,plain,(
  spl33_45 <=> ~para(sK26,sK25,sK25,sK23)),
  introduced(avatar_definition,[new_symbols(naming,[spl33_45])])).
fof(f1034,plain,(
  spl33_47 <=> ~para(sK26,sK25,sK25,sK24)),
  introduced(avatar_definition,[new_symbols(naming,[spl33_47])])).
fof(f1862,plain,(
  ~para(sK26,sK25,sK25,sK23) | ~spl33_47),
  inference(subsumption_resolution,[],[f1860,f1035])).
fof(f1035,plain,(
  ~para(sK26,sK25,sK25,sK24) | ~spl33_47),
  inference(avatar_component_clause,[],[f1034])).
fof(f1860,plain,(
  ~para(sK26,sK25,sK25,sK23) | para(sK26,sK25,sK25,sK24)),
  inference(resolution,[],[f759,f433])).
fof(f759,plain,(
  eqangle(sK26,sK25,sK25,sK24,sK26,sK25,sK25,sK23)),
  inference(resolution,[],[f379,f475])).
% SZS output end Proof for 89
% ------------------------------
% Version: Vampire 4.4.0 (commit 7916d27 on 2019-08-23 08:50:16 +0100)
% Termination reason: Refutation

% Memory used [KB]: 24050
% Time elapsed: 0.748 s
% ------------------------------
% ------------------------------
