% Refutation found. Thanks to Tanya!
% SZS status Theorem for trapezoid
% SZS output start Proof for trapezoid
fof(f8771,plain,(
  $false),
  inference(avatar_sat_refutation,[],[f782,f1927,f3555,f7336,f7702,f8712,f8755])).
fof(f8755,plain,(
  ~spl26_6 | ~spl26_156),
  inference(avatar_contradiction_clause,[],[f8754])).
fof(f8754,plain,(
  $false | (~spl26_6 | ~spl26_156)),
  inference(resolution,[],[f8714,f475])).
fof(f475,plain,(
  ~midp(sK25,sK21,sK22)),
  inference(cnf_transformation,[],[f357])).
fof(f357,plain,(
  ~midp(sK25,sK21,sK22) & coll(sK25,sK21,sK22) & para(sK20,sK21,sK25,sK24) & midp(sK24,sK23,sK20) & para(sK20,sK21,sK22,sK23)),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK20,sK21,sK22,sK23,sK24,sK25])],[f319,f356])).
fof(f356,plain,(
  ? [X0,X1,X2,X3,X4,X5] : (~midp(X5,X1,X2) & coll(X5,X1,X2) & para(X0,X1,X5,X4) & midp(X4,X3,X0) & para(X0,X1,X2,X3)) => (~midp(sK25,sK21,sK22) & coll(sK25,sK21,sK22) & para(sK20,sK21,sK25,sK24) & midp(sK24,sK23,sK20) & para(sK20,sK21,sK22,sK23))),
  introduced(choice_axiom,[])).
fof(f319,plain,(
  ? [X0,X1,X2,X3,X4,X5] : (~midp(X5,X1,X2) & coll(X5,X1,X2) & para(X0,X1,X5,X4) & midp(X4,X3,X0) & para(X0,X1,X2,X3))),
  inference(flattening,[],[f318])).
fof(f318,plain,(
  ? [X0,X1,X2,X3,X4,X5] : (~midp(X5,X1,X2) & (coll(X5,X1,X2) & para(X0,X1,X5,X4) & midp(X4,X3,X0) & para(X0,X1,X2,X3)))),
  inference(ennf_transformation,[],[f96])).
fof(f96,negated_conjecture,(
  ~! [X0,X1,X2,X3,X4,X5] : ((coll(X5,X1,X2) & para(X0,X1,X5,X4) & midp(X4,X3,X0) & para(X0,X1,X2,X3)) => midp(X5,X1,X2))),
  inference(negated_conjecture,[],[f95])).
fof(f95,conjecture,(
  ! [X0,X1,X2,X3,X4,X5] : ((coll(X5,X1,X2) & para(X0,X1,X5,X4) & midp(X4,X3,X0) & para(X0,X1,X2,X3)) => midp(X5,X1,X2))),
  file('./Other/Auxiliary/aux2/trapezoid.p',exemploOtherAuxiliaryaux2trapezoid)).
fof(f8714,plain,(
  ( ! [X0] : (midp(X0,sK21,sK22)) ) | (~spl26_6 | ~spl26_156)),
  inference(resolution,[],[f8713,f368])).
fof(f368,plain,(
  ( ! [X2,X0,X1] : (~midp(X2,X1,X0) | midp(X2,X0,X1)) )),
  inference(cnf_transformation,[],[f181])).
fof(f181,plain,(
  ! [X0,X1,X2] : (midp(X2,X0,X1) | ~midp(X2,X1,X0))),
  inference(ennf_transformation,[],[f97])).
fof(f97,plain,(
  ! [X0,X1,X2] : (midp(X2,X1,X0) => midp(X2,X0,X1))),
  inference(rectify,[],[f11])).
fof(f11,axiom,(
  ! [X0,X1,X6] : (midp(X6,X1,X0) => midp(X6,X0,X1))),
  file('./Other/Auxiliary/aux2/trapezoid.p',ruleD11)).
fof(f8713,plain,(
  ( ! [X1] : (midp(X1,sK22,sK21)) ) | (~spl26_6 | ~spl26_156)),
  inference(subsumption_resolution,[],[f781,f8463])).
fof(f8463,plain,(
  ( ! [X46] : (midp(X46,sK20,sK23)) ) | ~spl26_156),
  inference(subsumption_resolution,[],[f8459,f8450])).
fof(f8450,plain,(
  ( ! [X26,X27] : (midp(X26,X27,X27)) ) | ~spl26_156),
  inference(resolution,[],[f8375,f3688])).
fof(f3688,plain,(
  ( ! [X12,X13,X11] : (~midp(X11,X13,X13) | midp(X11,X12,X12)) ) | ~spl26_156),
  inference(subsumption_resolution,[],[f3680,f1923])).
fof(f1923,plain,(
  ( ! [X2,X3] : (para(X2,X3,X2,X3)) ) | ~spl26_156),
  inference(avatar_component_clause,[],[f1922])).
fof(f1922,plain,(
  spl26_156 <=> ! [X3,X2] : para(X2,X3,X2,X3)),
  introduced(avatar_definition,[new_symbols(naming,[spl26_156])])).
fof(f3680,plain,(
  ( ! [X12,X13,X11] : (midp(X11,X12,X12) | ~para(X13,X12,X13,X12) | ~midp(X11,X13,X13)) ) | ~spl26_156),
  inference(resolution,[],[f1923,f422])).
fof(f422,plain,(
  ( ! [X4,X2,X0,X3,X1] : (~para(X0,X3,X1,X2) | midp(X4,X2,X3) | ~para(X0,X2,X1,X3) | ~midp(X4,X0,X1)) )),
  inference(cnf_transformation,[],[f263])).
fof(f263,plain,(
  ! [X0,X1,X2,X3,X4] : (midp(X4,X2,X3) | ~para(X0,X3,X1,X2) | ~para(X0,X2,X1,X3) | ~midp(X4,X0,X1))),
  inference(flattening,[],[f262])).
fof(f262,plain,(
  ! [X0,X1,X2,X3,X4] : (midp(X4,X2,X3) | (~para(X0,X3,X1,X2) | ~para(X0,X2,X1,X3) | ~midp(X4,X0,X1)))),
  inference(ennf_transformation,[],[f143])).
fof(f143,plain,(
  ! [X0,X1,X2,X3,X4] : ((para(X0,X3,X1,X2) & para(X0,X2,X1,X3) & midp(X4,X0,X1)) => midp(X4,X2,X3))),
  inference(rectify,[],[f65])).
fof(f65,axiom,(
  ! [X0,X1,X2,X3,X6] : ((para(X0,X3,X1,X2) & para(X0,X2,X1,X3) & midp(X6,X0,X1)) => midp(X6,X2,X3))),
  file('./Other/Auxiliary/aux2/trapezoid.p',ruleD64)).
fof(f8375,plain,(
  ( ! [X19] : (midp(X19,X19,X19)) ) | ~spl26_156),
  inference(subsumption_resolution,[],[f8347,f3677])).
fof(f3677,plain,(
  ( ! [X4,X3] : (coll(X3,X4,X4)) ) | ~spl26_156),
  inference(resolution,[],[f1923,f424])).
fof(f424,plain,(
  ( ! [X2,X0,X1] : (~para(X0,X1,X0,X2) | coll(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f266])).
fof(f266,plain,(
  ! [X0,X1,X2] : (coll(X0,X1,X2) | ~para(X0,X1,X0,X2))),
  inference(ennf_transformation,[],[f67])).
fof(f67,axiom,(
  ! [X0,X1,X2] : (para(X0,X1,X0,X2) => coll(X0,X1,X2))),
  file('./Other/Auxiliary/aux2/trapezoid.p',ruleD66)).
fof(f8347,plain,(
  ( ! [X19] : (~coll(X19,X19,X19) | midp(X19,X19,X19)) ) | ~spl26_156),
  inference(resolution,[],[f7275,f425])).
fof(f425,plain,(
  ( ! [X2,X0,X1] : (~cong(X0,X1,X0,X2) | ~coll(X0,X1,X2) | midp(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f268])).
fof(f268,plain,(
  ! [X0,X1,X2] : (midp(X0,X1,X2) | ~coll(X0,X1,X2) | ~cong(X0,X1,X0,X2))),
  inference(flattening,[],[f267])).
fof(f267,plain,(
  ! [X0,X1,X2] : (midp(X0,X1,X2) | (~coll(X0,X1,X2) | ~cong(X0,X1,X0,X2)))),
  inference(ennf_transformation,[],[f68])).
fof(f68,axiom,(
  ! [X0,X1,X2] : ((coll(X0,X1,X2) & cong(X0,X1,X0,X2)) => midp(X0,X1,X2))),
  file('./Other/Auxiliary/aux2/trapezoid.p',ruleD67)).
fof(f7275,plain,(
  ( ! [X35] : (cong(X35,X35,X35,X35)) ) | ~spl26_156),
  inference(subsumption_resolution,[],[f7272,f7238])).
fof(f7238,plain,(
  ( ! [X2,X0,X1] : (cyclic(X0,X0,X1,X2)) ) | ~spl26_156),
  inference(resolution,[],[f3545,f1923])).
fof(f3545,plain,(
  ( ! [X68,X66,X67] : (~para(X66,X67,X66,X67) | cyclic(X67,X67,X66,X68)) )),
  inference(subsumption_resolution,[],[f2089,f2090])).
fof(f2090,plain,(
  ( ! [X70,X71,X69] : (~para(X69,X70,X69,X70) | coll(X69,X71,X70) | cyclic(X70,X70,X69,X71)) )),
  inference(resolution,[],[f686,f399])).
fof(f399,plain,(
  ( ! [X2,X0,X3,X1] : (~eqangle(X2,X0,X2,X1,X3,X0,X3,X1) | coll(X2,X3,X0) | cyclic(X0,X1,X2,X3)) )),
  inference(cnf_transformation,[],[f221])).
fof(f221,plain,(
  ! [X0,X1,X2,X3] : (cyclic(X0,X1,X2,X3) | coll(X2,X3,X0) | ~eqangle(X2,X0,X2,X1,X3,X0,X3,X1))),
  inference(flattening,[],[f220])).
fof(f220,plain,(
  ! [X0,X1,X2,X3] : (cyclic(X0,X1,X2,X3) | (coll(X2,X3,X0) | ~eqangle(X2,X0,X2,X1,X3,X0,X3,X1)))),
  inference(ennf_transformation,[],[f121])).
fof(f121,plain,(
  ! [X0,X1,X2,X3] : ((~coll(X2,X3,X0) & eqangle(X2,X0,X2,X1,X3,X0,X3,X1)) => cyclic(X0,X1,X2,X3))),
  inference(rectify,[],[f42])).
fof(f42,axiom,(
  ! [X0,X1,X8,X9] : ((~coll(X8,X9,X0) & eqangle(X8,X0,X8,X1,X9,X0,X9,X1)) => cyclic(X0,X1,X8,X9))),
  file('./Other/Auxiliary/aux2/trapezoid.p',ruleD42a)).
fof(f686,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (eqangle(X0,X1,X2,X3,X4,X5,X4,X5) | ~para(X0,X1,X2,X3)) )),
  inference(resolution,[],[f378,f397])).
fof(f397,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (eqangle(X0,X1,X4,X5,X2,X3,X4,X5) | ~para(X0,X1,X2,X3)) )),
  inference(cnf_transformation,[],[f218])).
fof(f218,plain,(
  ! [X0,X1,X2,X3,X4,X5] : (eqangle(X0,X1,X4,X5,X2,X3,X4,X5) | ~para(X0,X1,X2,X3))),
  inference(ennf_transformation,[],[f119])).
fof(f119,plain,(
  ! [X0,X1,X2,X3,X4,X5] : (para(X0,X1,X2,X3) => eqangle(X0,X1,X4,X5,X2,X3,X4,X5))),
  inference(rectify,[],[f40])).
fof(f40,axiom,(
  ! [X0,X1,X2,X3,X8,X9] : (para(X0,X1,X2,X3) => eqangle(X0,X1,X8,X9,X2,X3,X8,X9))),
  file('./Other/Auxiliary/aux2/trapezoid.p',ruleD40)).
fof(f378,plain,(
  ( ! [X6,X4,X2,X0,X7,X5,X3,X1] : (~eqangle(X0,X1,X2,X3,X4,X5,X6,X7) | eqangle(X0,X1,X4,X5,X2,X3,X6,X7)) )),
  inference(cnf_transformation,[],[f194])).
fof(f194,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6,X7] : (eqangle(X0,X1,X4,X5,X2,X3,X6,X7) | ~eqangle(X0,X1,X2,X3,X4,X5,X6,X7))),
  inference(ennf_transformation,[],[f103])).
fof(f103,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6,X7] : (eqangle(X0,X1,X2,X3,X4,X5,X6,X7) => eqangle(X0,X1,X4,X5,X2,X3,X6,X7))),
  inference(rectify,[],[f21])).
fof(f21,axiom,(
  ! [X0,X1,X2,X3,X8,X9,X10,X11] : (eqangle(X0,X1,X2,X3,X8,X9,X10,X11) => eqangle(X0,X1,X8,X9,X2,X3,X10,X11))),
  file('./Other/Auxiliary/aux2/trapezoid.p',ruleD21)).
fof(f2089,plain,(
  ( ! [X68,X66,X67] : (~para(X66,X67,X66,X67) | ~coll(X66,X68,X67) | cyclic(X67,X67,X66,X68)) )),
  inference(resolution,[],[f686,f400])).
fof(f400,plain,(
  ( ! [X2,X0,X3,X1] : (~eqangle(X2,X0,X2,X1,X3,X0,X3,X1) | ~coll(X2,X3,X1) | cyclic(X0,X1,X2,X3)) )),
  inference(cnf_transformation,[],[f223])).
fof(f223,plain,(
  ! [X0,X1,X2,X3] : (cyclic(X0,X1,X2,X3) | ~coll(X2,X3,X1) | ~eqangle(X2,X0,X2,X1,X3,X0,X3,X1))),
  inference(flattening,[],[f222])).
fof(f222,plain,(
  ! [X0,X1,X2,X3] : (cyclic(X0,X1,X2,X3) | (~coll(X2,X3,X1) | ~eqangle(X2,X0,X2,X1,X3,X0,X3,X1)))),
  inference(ennf_transformation,[],[f122])).
fof(f122,plain,(
  ! [X0,X1,X2,X3] : ((coll(X2,X3,X1) & eqangle(X2,X0,X2,X1,X3,X0,X3,X1)) => cyclic(X0,X1,X2,X3))),
  inference(rectify,[],[f43])).
fof(f43,axiom,(
  ! [X0,X1,X8,X9] : ((coll(X8,X9,X1) & eqangle(X8,X0,X8,X1,X9,X0,X9,X1)) => cyclic(X0,X1,X8,X9))),
  file('./Other/Auxiliary/aux2/trapezoid.p',ruleD42b)).
fof(f7272,plain,(
  ( ! [X35,X36] : (cong(X35,X35,X35,X35) | ~cyclic(X35,X35,X36,X35)) ) | ~spl26_156),
  inference(duplicate_literal_removal,[],[f7259])).
fof(f7259,plain,(
  ( ! [X35,X36] : (cong(X35,X35,X35,X35) | ~cyclic(X35,X35,X36,X35) | ~cyclic(X35,X35,X36,X35)) ) | ~spl26_156),
  inference(resolution,[],[f7238,f1264])).
fof(f1264,plain,(
  ( ! [X10,X8,X7,X9] : (~cyclic(X7,X8,X9,X10) | cong(X7,X8,X7,X8) | ~cyclic(X7,X8,X9,X8) | ~cyclic(X7,X8,X9,X7)) )),
  inference(duplicate_literal_removal,[],[f1258])).
fof(f1258,plain,(
  ( ! [X10,X8,X7,X9] : (cong(X7,X8,X7,X8) | ~cyclic(X7,X8,X9,X10) | ~cyclic(X7,X8,X9,X8) | ~cyclic(X7,X8,X9,X7) | ~cyclic(X7,X8,X9,X10)) )),
  inference(resolution,[],[f401,f398])).
fof(f398,plain,(
  ( ! [X2,X0,X3,X1] : (eqangle(X2,X0,X2,X1,X3,X0,X3,X1) | ~cyclic(X0,X1,X2,X3)) )),
  inference(cnf_transformation,[],[f219])).
fof(f219,plain,(
  ! [X0,X1,X2,X3] : (eqangle(X2,X0,X2,X1,X3,X0,X3,X1) | ~cyclic(X0,X1,X2,X3))),
  inference(ennf_transformation,[],[f120])).
fof(f120,plain,(
  ! [X0,X1,X2,X3] : (cyclic(X0,X1,X2,X3) => eqangle(X2,X0,X2,X1,X3,X0,X3,X1))),
  inference(rectify,[],[f41])).
fof(f41,axiom,(
  ! [X0,X1,X8,X9] : (cyclic(X0,X1,X8,X9) => eqangle(X8,X0,X8,X1,X9,X0,X9,X1))),
  file('./Other/Auxiliary/aux2/trapezoid.p',ruleD41)).
fof(f401,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (~eqangle(X2,X0,X2,X1,X5,X3,X5,X4) | cong(X0,X1,X3,X4) | ~cyclic(X0,X1,X2,X5) | ~cyclic(X0,X1,X2,X4) | ~cyclic(X0,X1,X2,X3)) )),
  inference(cnf_transformation,[],[f225])).
fof(f225,plain,(
  ! [X0,X1,X2,X3,X4,X5] : (cong(X0,X1,X3,X4) | ~eqangle(X2,X0,X2,X1,X5,X3,X5,X4) | ~cyclic(X0,X1,X2,X5) | ~cyclic(X0,X1,X2,X4) | ~cyclic(X0,X1,X2,X3))),
  inference(flattening,[],[f224])).
fof(f224,plain,(
  ! [X0,X1,X2,X3,X4,X5] : (cong(X0,X1,X3,X4) | (~eqangle(X2,X0,X2,X1,X5,X3,X5,X4) | ~cyclic(X0,X1,X2,X5) | ~cyclic(X0,X1,X2,X4) | ~cyclic(X0,X1,X2,X3)))),
  inference(ennf_transformation,[],[f123])).
fof(f123,plain,(
  ! [X0,X1,X2,X3,X4,X5] : ((eqangle(X2,X0,X2,X1,X5,X3,X5,X4) & cyclic(X0,X1,X2,X5) & cyclic(X0,X1,X2,X4) & cyclic(X0,X1,X2,X3)) => cong(X0,X1,X3,X4))),
  inference(rectify,[],[f44])).
fof(f44,axiom,(
  ! [X0,X1,X2,X8,X9,X14] : ((eqangle(X2,X0,X2,X1,X14,X8,X14,X9) & cyclic(X0,X1,X2,X14) & cyclic(X0,X1,X2,X9) & cyclic(X0,X1,X2,X8)) => cong(X0,X1,X8,X9))),
  file('./Other/Auxiliary/aux2/trapezoid.p',ruleD43)).
fof(f8459,plain,(
  ( ! [X46] : (midp(X46,sK20,sK23) | ~midp(X46,sK24,sK24)) ) | ~spl26_156),
  inference(resolution,[],[f8375,f1326])).
fof(f1326,plain,(
  ( ! [X2,X3,X1] : (~midp(sK24,X1,X2) | midp(X3,sK20,sK23) | ~midp(X3,X1,X2)) )),
  inference(subsumption_resolution,[],[f1320,f537])).
fof(f537,plain,(
  ( ! [X2,X3] : (para(X2,sK20,X3,sK23) | ~midp(sK24,X2,X3)) )),
  inference(resolution,[],[f421,f482])).
fof(f482,plain,(
  midp(sK24,sK20,sK23)),
  inference(resolution,[],[f368,f472])).
fof(f472,plain,(
  midp(sK24,sK23,sK20)),
  inference(cnf_transformation,[],[f357])).
fof(f421,plain,(
  ( ! [X4,X2,X0,X3,X1] : (~midp(X4,X2,X3) | para(X0,X2,X1,X3) | ~midp(X4,X0,X1)) )),
  inference(cnf_transformation,[],[f261])).
fof(f261,plain,(
  ! [X0,X1,X2,X3,X4] : (para(X0,X2,X1,X3) | ~midp(X4,X2,X3) | ~midp(X4,X0,X1))),
  inference(flattening,[],[f260])).
fof(f260,plain,(
  ! [X0,X1,X2,X3,X4] : (para(X0,X2,X1,X3) | (~midp(X4,X2,X3) | ~midp(X4,X0,X1)))),
  inference(ennf_transformation,[],[f142])).
fof(f142,plain,(
  ! [X0,X1,X2,X3,X4] : ((midp(X4,X2,X3) & midp(X4,X0,X1)) => para(X0,X2,X1,X3))),
  inference(rectify,[],[f64])).
fof(f64,axiom,(
  ! [X0,X1,X2,X3,X6] : ((midp(X6,X2,X3) & midp(X6,X0,X1)) => para(X0,X2,X1,X3))),
  file('./Other/Auxiliary/aux2/trapezoid.p',ruleD63)).
fof(f1320,plain,(
  ( ! [X2,X3,X1] : (~midp(sK24,X1,X2) | midp(X3,sK20,sK23) | ~para(X1,sK20,X2,sK23) | ~midp(X3,X1,X2)) )),
  inference(resolution,[],[f536,f422])).
fof(f536,plain,(
  ( ! [X0,X1] : (para(X0,sK23,X1,sK20) | ~midp(sK24,X0,X1)) )),
  inference(resolution,[],[f421,f472])).
fof(f781,plain,(
  ( ! [X1] : (~midp(X1,sK20,sK23) | midp(X1,sK22,sK21)) ) | ~spl26_6),
  inference(avatar_component_clause,[],[f780])).
fof(f780,plain,(
  spl26_6 <=> ! [X1] : (midp(X1,sK22,sK21) | ~midp(X1,sK20,sK23))),
  introduced(avatar_definition,[new_symbols(naming,[spl26_6])])).
fof(f8712,plain,(
  ~spl26_156 | spl26_185),
  inference(avatar_contradiction_clause,[],[f8711])).
fof(f8711,plain,(
  $false | (~spl26_156 | ~spl26_185)),
  inference(subsumption_resolution,[],[f2510,f4096])).
fof(f4096,plain,(
  ( ! [X2,X0,X1] : (coll(X0,X1,X2)) ) | ~spl26_156),
  inference(subsumption_resolution,[],[f4086,f4054])).
fof(f4054,plain,(
  ( ! [X6,X5] : (coll(X5,X5,X6)) ) | ~spl26_156),
  inference(resolution,[],[f4016,f358])).
fof(f358,plain,(
  ( ! [X2,X0,X1] : (~coll(X0,X1,X2) | coll(X0,X2,X1)) )),
  inference(cnf_transformation,[],[f167])).
fof(f167,plain,(
  ! [X0,X1,X2] : (coll(X0,X2,X1) | ~coll(X0,X1,X2))),
  inference(ennf_transformation,[],[f1])).
fof(f1,axiom,(
  ! [X0,X1,X2] : (coll(X0,X1,X2) => coll(X0,X2,X1))),
  file('./Other/Auxiliary/aux2/trapezoid.p',ruleD1)).
fof(f4016,plain,(
  ( ! [X4,X3] : (coll(X3,X4,X3)) ) | ~spl26_156),
  inference(resolution,[],[f3677,f359])).
fof(f359,plain,(
  ( ! [X2,X0,X1] : (~coll(X0,X1,X2) | coll(X1,X0,X2)) )),
  inference(cnf_transformation,[],[f168])).
fof(f168,plain,(
  ! [X0,X1,X2] : (coll(X1,X0,X2) | ~coll(X0,X1,X2))),
  inference(ennf_transformation,[],[f2])).
fof(f2,axiom,(
  ! [X0,X1,X2] : (coll(X0,X1,X2) => coll(X1,X0,X2))),
  file('./Other/Auxiliary/aux2/trapezoid.p',ruleD2)).
fof(f4086,plain,(
  ( ! [X2,X0,X1] : (coll(X0,X1,X2) | ~coll(X2,X2,X0)) ) | ~spl26_156),
  inference(resolution,[],[f4054,f360])).
fof(f360,plain,(
  ( ! [X2,X0,X3,X1] : (~coll(X0,X1,X3) | coll(X2,X3,X0) | ~coll(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f170])).
fof(f170,plain,(
  ! [X0,X1,X2,X3] : (coll(X2,X3,X0) | ~coll(X0,X1,X3) | ~coll(X0,X1,X2))),
  inference(flattening,[],[f169])).
fof(f169,plain,(
  ! [X0,X1,X2,X3] : (coll(X2,X3,X0) | (~coll(X0,X1,X3) | ~coll(X0,X1,X2)))),
  inference(ennf_transformation,[],[f3])).
fof(f3,axiom,(
  ! [X0,X1,X2,X3] : ((coll(X0,X1,X3) & coll(X0,X1,X2)) => coll(X2,X3,X0))),
  file('./Other/Auxiliary/aux2/trapezoid.p',ruleD3)).
fof(f2510,plain,(
  ~coll(sK25,sK25,sK24) | ~spl26_185),
  inference(avatar_component_clause,[],[f2509])).
fof(f2509,plain,(
  spl26_185 <=> ~coll(sK25,sK25,sK24)),
  introduced(avatar_definition,[new_symbols(naming,[spl26_185])])).
fof(f7702,plain,(
  spl26_4 | ~spl26_36),
  inference(avatar_split_clause,[],[f7688,f854,f774])).
fof(f774,plain,(
  spl26_4 <=> para(sK20,sK22,sK23,sK21)),
  introduced(avatar_definition,[new_symbols(naming,[spl26_4])])).
fof(f854,plain,(
  spl26_36 <=> para(sK22,sK20,sK21,sK23)),
  introduced(avatar_definition,[new_symbols(naming,[spl26_36])])).
fof(f7688,plain,(
  para(sK20,sK22,sK23,sK21) | ~spl26_36),
  inference(resolution,[],[f7430,f361])).
fof(f361,plain,(
  ( ! [X2,X0,X3,X1] : (~para(X0,X1,X2,X3) | para(X0,X1,X3,X2)) )),
  inference(cnf_transformation,[],[f171])).
fof(f171,plain,(
  ! [X0,X1,X2,X3] : (para(X0,X1,X3,X2) | ~para(X0,X1,X2,X3))),
  inference(ennf_transformation,[],[f4])).
fof(f4,axiom,(
  ! [X0,X1,X2,X3] : (para(X0,X1,X2,X3) => para(X0,X1,X3,X2))),
  file('./Other/Auxiliary/aux2/trapezoid.p',ruleD4)).
fof(f7430,plain,(
  para(sK20,sK22,sK21,sK23) | ~spl26_36),
  inference(resolution,[],[f855,f1816])).
fof(f1816,plain,(
  ( ! [X2,X0,X3,X1] : (~para(X0,X1,X2,X3) | para(X1,X0,X2,X3)) )),
  inference(resolution,[],[f663,f396])).
fof(f396,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (~eqangle(X0,X1,X4,X5,X2,X3,X4,X5) | para(X0,X1,X2,X3)) )),
  inference(cnf_transformation,[],[f217])).
fof(f217,plain,(
  ! [X0,X1,X2,X3,X4,X5] : (para(X0,X1,X2,X3) | ~eqangle(X0,X1,X4,X5,X2,X3,X4,X5))),
  inference(ennf_transformation,[],[f118])).
fof(f118,plain,(
  ! [X0,X1,X2,X3,X4,X5] : (eqangle(X0,X1,X4,X5,X2,X3,X4,X5) => para(X0,X1,X2,X3))),
  inference(rectify,[],[f39])).
fof(f39,axiom,(
  ! [X0,X1,X2,X3,X8,X9] : (eqangle(X0,X1,X8,X9,X2,X3,X8,X9) => para(X0,X1,X2,X3))),
  file('./Other/Auxiliary/aux2/trapezoid.p',ruleD39)).
fof(f663,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (eqangle(X0,X1,X2,X3,X4,X5,X2,X3) | ~para(X1,X0,X4,X5)) )),
  inference(resolution,[],[f375,f397])).
fof(f375,plain,(
  ( ! [X6,X4,X2,X0,X7,X5,X3,X1] : (~eqangle(X0,X1,X2,X3,X4,X5,X6,X7) | eqangle(X1,X0,X2,X3,X4,X5,X6,X7)) )),
  inference(cnf_transformation,[],[f191])).
fof(f191,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6,X7] : (eqangle(X1,X0,X2,X3,X4,X5,X6,X7) | ~eqangle(X0,X1,X2,X3,X4,X5,X6,X7))),
  inference(ennf_transformation,[],[f100])).
fof(f100,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6,X7] : (eqangle(X0,X1,X2,X3,X4,X5,X6,X7) => eqangle(X1,X0,X2,X3,X4,X5,X6,X7))),
  inference(rectify,[],[f18])).
fof(f18,axiom,(
  ! [X0,X1,X2,X3,X8,X9,X10,X11] : (eqangle(X0,X1,X2,X3,X8,X9,X10,X11) => eqangle(X1,X0,X2,X3,X8,X9,X10,X11))),
  file('./Other/Auxiliary/aux2/trapezoid.p',ruleD18)).
fof(f855,plain,(
  para(sK22,sK20,sK21,sK23) | ~spl26_36),
  inference(avatar_component_clause,[],[f854])).
fof(f7336,plain,(
  spl26_37 | ~spl26_156 | ~spl26_184),
  inference(avatar_contradiction_clause,[],[f7335])).
fof(f7335,plain,(
  $false | (~spl26_37 | ~spl26_156 | ~spl26_184)),
  inference(subsumption_resolution,[],[f7330,f6432])).
fof(f6432,plain,(
  ~midp(sK24,sK22,sK21) | ~spl26_37),
  inference(resolution,[],[f537,f858])).
fof(f858,plain,(
  ~para(sK22,sK20,sK21,sK23) | ~spl26_37),
  inference(avatar_component_clause,[],[f857])).
fof(f857,plain,(
  spl26_37 <=> ~para(sK22,sK20,sK21,sK23)),
  introduced(avatar_definition,[new_symbols(naming,[spl26_37])])).
fof(f7330,plain,(
  midp(sK24,sK22,sK21) | (~spl26_156 | ~spl26_184)),
  inference(resolution,[],[f7289,f4183])).
fof(f4183,plain,(
  ( ! [X12,X10,X11] : (~midp(X12,X11,X10) | midp(X10,X11,X12)) ) | ~spl26_156),
  inference(subsumption_resolution,[],[f4176,f4096])).
fof(f4176,plain,(
  ( ! [X12,X10,X11] : (~coll(X10,X11,X12) | midp(X10,X11,X12) | ~midp(X12,X11,X10)) ) | ~spl26_156),
  inference(resolution,[],[f3684,f403])).
fof(f403,plain,(
  ( ! [X4,X2,X0,X3,X1] : (~para(X3,X4,X1,X2) | ~coll(X4,X0,X2) | midp(X4,X0,X2) | ~midp(X3,X0,X1)) )),
  inference(cnf_transformation,[],[f229])).
fof(f229,plain,(
  ! [X0,X1,X2,X3,X4] : (midp(X4,X0,X2) | ~coll(X4,X0,X2) | ~para(X3,X4,X1,X2) | ~midp(X3,X0,X1))),
  inference(flattening,[],[f228])).
fof(f228,plain,(
  ! [X0,X1,X2,X3,X4] : (midp(X4,X0,X2) | (~coll(X4,X0,X2) | ~para(X3,X4,X1,X2) | ~midp(X3,X0,X1)))),
  inference(ennf_transformation,[],[f125])).
fof(f125,plain,(
  ! [X0,X1,X2,X3,X4] : ((coll(X4,X0,X2) & para(X3,X4,X1,X2) & midp(X3,X0,X1)) => midp(X4,X0,X2))),
  inference(rectify,[],[f46])).
fof(f46,axiom,(
  ! [X0,X1,X2,X4,X5] : ((coll(X5,X0,X2) & para(X4,X5,X1,X2) & midp(X4,X0,X1)) => midp(X5,X0,X2))),
  file('./Other/Auxiliary/aux2/trapezoid.p',ruleD45)).
fof(f3684,plain,(
  ( ! [X24,X23] : (para(X23,X24,X24,X23)) ) | ~spl26_156),
  inference(resolution,[],[f1923,f361])).
fof(f7289,plain,(
  midp(sK21,sK22,sK24) | (~spl26_156 | ~spl26_184)),
  inference(resolution,[],[f7283,f368])).
fof(f7283,plain,(
  midp(sK21,sK24,sK22) | (~spl26_156 | ~spl26_184)),
  inference(subsumption_resolution,[],[f7281,f2727])).
fof(f2727,plain,(
  coll(sK21,sK24,sK22) | ~spl26_184),
  inference(resolution,[],[f2700,f359])).
fof(f2700,plain,(
  coll(sK24,sK21,sK22) | ~spl26_184),
  inference(resolution,[],[f2660,f358])).
fof(f2660,plain,(
  coll(sK24,sK22,sK21) | ~spl26_184),
  inference(resolution,[],[f2609,f518])).
fof(f518,plain,(
  ( ! [X2] : (~coll(sK21,sK25,X2) | coll(X2,sK22,sK21)) )),
  inference(resolution,[],[f360,f478])).
fof(f478,plain,(
  coll(sK21,sK25,sK22)),
  inference(resolution,[],[f359,f474])).
fof(f474,plain,(
  coll(sK25,sK21,sK22)),
  inference(cnf_transformation,[],[f357])).
fof(f2609,plain,(
  coll(sK21,sK25,sK24) | ~spl26_184),
  inference(resolution,[],[f2578,f358])).
fof(f2578,plain,(
  coll(sK21,sK24,sK25) | ~spl26_184),
  inference(resolution,[],[f2561,f359])).
fof(f2561,plain,(
  coll(sK24,sK21,sK25) | ~spl26_184),
  inference(resolution,[],[f2513,f667])).
fof(f667,plain,(
  ( ! [X0] : (~coll(sK25,sK25,X0) | coll(X0,sK21,sK25)) )),
  inference(resolution,[],[f659,f360])).
fof(f659,plain,(
  coll(sK25,sK25,sK21)),
  inference(resolution,[],[f519,f481])).
fof(f481,plain,(
  coll(sK21,sK22,sK25)),
  inference(resolution,[],[f478,f358])).
fof(f519,plain,(
  ( ! [X3] : (~coll(sK21,sK22,X3) | coll(X3,sK25,sK21)) )),
  inference(resolution,[],[f360,f481])).
fof(f2513,plain,(
  coll(sK25,sK25,sK24) | ~spl26_184),
  inference(avatar_component_clause,[],[f2512])).
fof(f2512,plain,(
  spl26_184 <=> coll(sK25,sK25,sK24)),
  introduced(avatar_definition,[new_symbols(naming,[spl26_184])])).
fof(f7281,plain,(
  midp(sK21,sK24,sK22) | ~coll(sK21,sK24,sK22) | ~spl26_156),
  inference(resolution,[],[f631,f5448])).
fof(f5448,plain,(
  midp(sK20,sK24,sK23) | ~spl26_156),
  inference(resolution,[],[f5409,f368])).
fof(f5409,plain,(
  midp(sK20,sK23,sK24) | ~spl26_156),
  inference(resolution,[],[f4183,f472])).
fof(f631,plain,(
  ( ! [X1] : (~midp(sK20,X1,sK23) | midp(sK21,X1,sK22) | ~coll(sK21,X1,sK22)) )),
  inference(resolution,[],[f403,f492])).
fof(f492,plain,(
  para(sK20,sK21,sK23,sK22)),
  inference(resolution,[],[f361,f471])).
fof(f471,plain,(
  para(sK20,sK21,sK22,sK23)),
  inference(cnf_transformation,[],[f357])).
fof(f3555,plain,(
  ~spl26_158),
  inference(avatar_contradiction_clause,[],[f3552])).
fof(f3552,plain,(
  $false | ~spl26_158),
  inference(resolution,[],[f1926,f1367])).
fof(f1367,plain,(
  para(sK22,sK23,sK22,sK23)),
  inference(resolution,[],[f582,f496])).
fof(f496,plain,(
  para(sK22,sK23,sK20,sK21)),
  inference(resolution,[],[f362,f471])).
fof(f362,plain,(
  ( ! [X2,X0,X3,X1] : (~para(X0,X1,X2,X3) | para(X2,X3,X0,X1)) )),
  inference(cnf_transformation,[],[f172])).
fof(f172,plain,(
  ! [X0,X1,X2,X3] : (para(X2,X3,X0,X1) | ~para(X0,X1,X2,X3))),
  inference(ennf_transformation,[],[f5])).
fof(f5,axiom,(
  ! [X0,X1,X2,X3] : (para(X0,X1,X2,X3) => para(X2,X3,X0,X1))),
  file('./Other/Auxiliary/aux2/trapezoid.p',ruleD5)).
fof(f582,plain,(
  ( ! [X0,X1] : (~para(X0,X1,sK20,sK21) | para(X0,X1,sK22,sK23)) )),
  inference(resolution,[],[f363,f471])).
fof(f363,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (~para(X2,X3,X4,X5) | para(X0,X1,X4,X5) | ~para(X0,X1,X2,X3)) )),
  inference(cnf_transformation,[],[f174])).
fof(f174,plain,(
  ! [X0,X1,X2,X3,X4,X5] : (para(X0,X1,X4,X5) | ~para(X2,X3,X4,X5) | ~para(X0,X1,X2,X3))),
  inference(flattening,[],[f173])).
fof(f173,plain,(
  ! [X0,X1,X2,X3,X4,X5] : (para(X0,X1,X4,X5) | (~para(X2,X3,X4,X5) | ~para(X0,X1,X2,X3)))),
  inference(ennf_transformation,[],[f6])).
fof(f6,axiom,(
  ! [X0,X1,X2,X3,X4,X5] : ((para(X2,X3,X4,X5) & para(X0,X1,X2,X3)) => para(X0,X1,X4,X5))),
  file('./Other/Auxiliary/aux2/trapezoid.p',ruleD6)).
fof(f1926,plain,(
  ( ! [X0,X1] : (~para(X0,X1,X0,X1)) ) | ~spl26_158),
  inference(avatar_component_clause,[],[f1925])).
fof(f1925,plain,(
  spl26_158 <=> ! [X1,X0] : ~para(X0,X1,X0,X1)),
  introduced(avatar_definition,[new_symbols(naming,[spl26_158])])).
fof(f1927,plain,(
  spl26_156 | spl26_158),
  inference(avatar_split_clause,[],[f1898,f1925,f1922])).
fof(f1898,plain,(
  ( ! [X2,X0,X3,X1] : (~para(X0,X1,X0,X1) | para(X2,X3,X2,X3)) )),
  inference(resolution,[],[f670,f396])).
fof(f670,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (eqangle(X0,X1,X2,X3,X0,X1,X4,X5) | ~para(X2,X3,X4,X5)) )),
  inference(resolution,[],[f376,f397])).
fof(f376,plain,(
  ( ! [X6,X4,X2,X0,X7,X5,X3,X1] : (~eqangle(X0,X1,X2,X3,X4,X5,X6,X7) | eqangle(X2,X3,X0,X1,X6,X7,X4,X5)) )),
  inference(cnf_transformation,[],[f192])).
fof(f192,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6,X7] : (eqangle(X2,X3,X0,X1,X6,X7,X4,X5) | ~eqangle(X0,X1,X2,X3,X4,X5,X6,X7))),
  inference(ennf_transformation,[],[f101])).
fof(f101,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6,X7] : (eqangle(X0,X1,X2,X3,X4,X5,X6,X7) => eqangle(X2,X3,X0,X1,X6,X7,X4,X5))),
  inference(rectify,[],[f19])).
fof(f19,axiom,(
  ! [X0,X1,X2,X3,X8,X9,X10,X11] : (eqangle(X0,X1,X2,X3,X8,X9,X10,X11) => eqangle(X2,X3,X0,X1,X10,X11,X8,X9))),
  file('./Other/Auxiliary/aux2/trapezoid.p',ruleD19)).
fof(f782,plain,(
  ~spl26_5 | spl26_6),
  inference(avatar_split_clause,[],[f748,f780,f777])).
fof(f777,plain,(
  spl26_5 <=> ~para(sK20,sK22,sK23,sK21)),
  introduced(avatar_definition,[new_symbols(naming,[spl26_5])])).
fof(f748,plain,(
  ( ! [X1] : (midp(X1,sK22,sK21) | ~para(sK20,sK22,sK23,sK21) | ~midp(X1,sK20,sK23)) )),
  inference(resolution,[],[f422,f492])).
% SZS output end Proof for trapezoid
% ------------------------------
% Version: Vampire 4.4.0 (commit 7916d27 on 2019-08-23 08:50:16 +0100)
% Termination reason: Refutation

% Memory used [KB]: 9850
% Time elapsed: 0.176 s
% ------------------------------
% ------------------------------
