% Refutation found. Thanks to Tanya!
% SZS status Theorem for E059-56
% SZS output start Proof for E059-56
9. ! [X0,X1,X2,X3,X4,X5] : ((perp(X2,X3,X4,X5) & perp(X0,X1,X2,X3)) => para(X0,X1,X4,X5)) [input]
18. ! [X0,X1,X2,X3,X8,X9,X10,X11] : (eqangle(X0,X1,X2,X3,X8,X9,X10,X11) => eqangle(X1,X0,X2,X3,X8,X9,X10,X11)) [input]
19. ! [X0,X1,X2,X3,X8,X9,X10,X11] : (eqangle(X0,X1,X2,X3,X8,X9,X10,X11) => eqangle(X2,X3,X0,X1,X10,X11,X8,X9)) [input]
23. ! [X0,X1,X2,X3] : (cong(X0,X1,X2,X3) => cong(X0,X1,X3,X2)) [input]
25. ! [X0,X1,X2,X3,X4,X5] : ((cong(X2,X3,X4,X5) & cong(X0,X1,X2,X3)) => cong(X0,X1,X4,X5)) [input]
40. ! [X0,X1,X2,X3,X8,X9] : (para(X0,X1,X2,X3) => eqangle(X0,X1,X8,X9,X2,X3,X8,X9)) [input]
51. ! [X0,X1,X2,X7,X6] : ((eqangle(X0,X1,X0,X2,X7,X1,X7,X6) & coll(X6,X1,X2) & circle(X7,X0,X1,X2)) => midp(X6,X1,X2)) [input]
68. ! [X0,X1,X2] : (midp(X0,X1,X2) => cong(X0,X1,X0,X2)) [input]
94. ! [X0,X1,X2,X3,X4,X5,X12,X13,X18,X19,X20] : (circle(X12,X0,X18,X20) & coll(X18,X0,X2) & circle(X12,X0,X13,X19) & coll(X13,X1,X2) & circle(X12,X1,X5,X0) & coll(X5,X0,X4) & coll(X5,X1,X3) & coll(X4,X1,X2) & perp(X4,X0,X1,X2) & coll(X3,X0,X2) & perp(X3,X1,X0,X2)) => cong(X5,X2,X5,X18) [input]
95. ~! [X18,X5,X2] : (! [X0,X1,X2,X3,X4,X5,X12,X13,X18,X19,X20] : (circle(X12,X0,X18,X20) & coll(X18,X0,X2) & circle(X12,X0,X13,X19) & coll(X13,X1,X2) & circle(X12,X1,X5,X0) & coll(X5,X0,X4) & coll(X5,X1,X3) & coll(X4,X1,X2) & perp(X4,X0,X1,X2) & coll(X3,X0,X2) & perp(X3,X1,X0,X2)) => cong(X5,X2,X5,X18)) [negated conjecture 94]
99. ! [X0,X1,X2,X3,X4,X5,X6,X7] : (eqangle(X0,X1,X2,X3,X4,X5,X6,X7) => eqangle(X1,X0,X2,X3,X4,X5,X6,X7)) [rectify 18]
100. ! [X0,X1,X2,X3,X4,X5,X6,X7] : (eqangle(X0,X1,X2,X3,X4,X5,X6,X7) => eqangle(X2,X3,X0,X1,X6,X7,X4,X5)) [rectify 19]
118. ! [X0,X1,X2,X3,X4,X5] : (para(X0,X1,X2,X3) => eqangle(X0,X1,X4,X5,X2,X3,X4,X5)) [rectify 40]
129. ! [X0,X1,X2,X3,X4] : ((eqangle(X0,X1,X0,X2,X3,X1,X3,X4) & coll(X4,X1,X2) & circle(X3,X0,X1,X2)) => midp(X4,X1,X2)) [rectify 51]
165. ~! [X0,X1,X2] : (! [X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13] : (circle(X9,X3,X11,X13) & coll(X11,X3,X5) & circle(X9,X3,X10,X12) & coll(X10,X4,X5) & circle(X9,X4,X8,X3) & coll(X8,X3,X7) & coll(X8,X4,X6) & coll(X7,X4,X5) & perp(X7,X3,X4,X5) & coll(X6,X3,X5) & perp(X6,X4,X3,X5)) => cong(X1,X2,X1,X0)) [rectify 95]
176. ! [X0,X1,X2,X3,X4,X5] : (para(X0,X1,X4,X5) | (~perp(X2,X3,X4,X5) | ~perp(X0,X1,X2,X3))) [ennf transformation 9]
177. ! [X0,X1,X2,X3,X4,X5] : (para(X0,X1,X4,X5) | ~perp(X2,X3,X4,X5) | ~perp(X0,X1,X2,X3)) [flattening 176]
190. ! [X0,X1,X2,X3,X4,X5,X6,X7] : (eqangle(X1,X0,X2,X3,X4,X5,X6,X7) | ~eqangle(X0,X1,X2,X3,X4,X5,X6,X7)) [ennf transformation 99]
191. ! [X0,X1,X2,X3,X4,X5,X6,X7] : (eqangle(X2,X3,X0,X1,X6,X7,X4,X5) | ~eqangle(X0,X1,X2,X3,X4,X5,X6,X7)) [ennf transformation 100]
196. ! [X0,X1,X2,X3] : (cong(X0,X1,X3,X2) | ~cong(X0,X1,X2,X3)) [ennf transformation 23]
198. ! [X0,X1,X2,X3,X4,X5] : (cong(X0,X1,X4,X5) | (~cong(X2,X3,X4,X5) | ~cong(X0,X1,X2,X3))) [ennf transformation 25]
199. ! [X0,X1,X2,X3,X4,X5] : (cong(X0,X1,X4,X5) | ~cong(X2,X3,X4,X5) | ~cong(X0,X1,X2,X3)) [flattening 198]
217. ! [X0,X1,X2,X3,X4,X5] : (eqangle(X0,X1,X4,X5,X2,X3,X4,X5) | ~para(X0,X1,X2,X3)) [ennf transformation 118]
236. ! [X0,X1,X2,X3,X4] : (midp(X4,X1,X2) | (~eqangle(X0,X1,X0,X2,X3,X1,X3,X4) | ~coll(X4,X1,X2) | ~circle(X3,X0,X1,X2))) [ennf transformation 129]
237. ! [X0,X1,X2,X3,X4] : (midp(X4,X1,X2) | ~eqangle(X0,X1,X0,X2,X3,X1,X3,X4) | ~coll(X4,X1,X2) | ~circle(X3,X0,X1,X2)) [flattening 236]
266. ! [X0,X1,X2] : (cong(X0,X1,X0,X2) | ~midp(X0,X1,X2)) [ennf transformation 68]
315. ? [X0,X1,X2] : (~cong(X1,X2,X1,X0) & ! [X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13] : (circle(X9,X3,X11,X13) & coll(X11,X3,X5) & circle(X9,X3,X10,X12) & coll(X10,X4,X5) & circle(X9,X4,X8,X3) & coll(X8,X3,X7) & coll(X8,X4,X6) & coll(X7,X4,X5) & perp(X7,X3,X4,X5) & coll(X6,X3,X5) & perp(X6,X4,X3,X5))) [ennf transformation 165]
352. ? [X0,X1,X2] : (~cong(X1,X2,X1,X0) & ! [X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13] : (circle(X9,X3,X11,X13) & coll(X11,X3,X5) & circle(X9,X3,X10,X12) & coll(X10,X4,X5) & circle(X9,X4,X8,X3) & coll(X8,X3,X7) & coll(X8,X4,X6) & coll(X7,X4,X5) & perp(X7,X3,X4,X5) & coll(X6,X3,X5) & perp(X6,X4,X3,X5))) => (~cong(sK21,sK22,sK21,sK20) & ! [X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13] : (circle(X9,X3,X11,X13) & coll(X11,X3,X5) & circle(X9,X3,X10,X12) & coll(X10,X4,X5) & circle(X9,X4,X8,X3) & coll(X8,X3,X7) & coll(X8,X4,X6) & coll(X7,X4,X5) & perp(X7,X3,X4,X5) & coll(X6,X3,X5) & perp(X6,X4,X3,X5))) [choice axiom]
353. ~cong(sK21,sK22,sK21,sK20) & ! [X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13] : (circle(X9,X3,X11,X13) & coll(X11,X3,X5) & circle(X9,X3,X10,X12) & coll(X10,X4,X5) & circle(X9,X4,X8,X3) & coll(X8,X3,X7) & coll(X8,X4,X6) & coll(X7,X4,X5) & perp(X7,X3,X4,X5) & coll(X6,X3,X5) & perp(X6,X4,X3,X5)) [skolemisation 315,352]
362. ~perp(X2,X3,X4,X5) | para(X0,X1,X4,X5) | ~perp(X0,X1,X2,X3) [cnf transformation 177]
371. ~eqangle(X0,X1,X2,X3,X4,X5,X6,X7) | eqangle(X1,X0,X2,X3,X4,X5,X6,X7) [cnf transformation 190]
372. ~eqangle(X0,X1,X2,X3,X4,X5,X6,X7) | eqangle(X2,X3,X0,X1,X6,X7,X4,X5) [cnf transformation 191]
376. ~cong(X0,X1,X2,X3) | cong(X0,X1,X3,X2) [cnf transformation 196]
378. ~cong(X2,X3,X4,X5) | cong(X0,X1,X4,X5) | ~cong(X0,X1,X2,X3) [cnf transformation 199]
393. eqangle(X0,X1,X4,X5,X2,X3,X4,X5) | ~para(X0,X1,X2,X3) [cnf transformation 217]
405. ~eqangle(X0,X1,X0,X2,X3,X1,X3,X4) | midp(X4,X1,X2) | ~coll(X4,X1,X2) | ~circle(X3,X0,X1,X2) [cnf transformation 237]
422. cong(X0,X1,X0,X2) | ~midp(X0,X1,X2) [cnf transformation 266]
469. perp(X7,X3,X4,X5) [cnf transformation 353]
476. coll(X11,X3,X5) [cnf transformation 353]
477. circle(X9,X3,X11,X13) [cnf transformation 353]
478. ~cong(sK21,sK22,sK21,sK20) [cnf transformation 353]
484. cong(X0,X1,X2,X0) | ~midp(X0,X1,X2) [resolution 376,422]
537. para(X0,X1,X2,X3) | ~perp(X0,X1,X4,X5) [resolution 362,469]
539. para(X0,X1,X2,X3) [subsumption resolution 537,469]
567. ~cong(X5,X6,X7,X9) | cong(X5,X6,X7,X8) | ~midp(X7,X9,X8) [resolution 378,422]
589. cong(X12,X13,X14,X15) | ~midp(X14,X12,X15) | ~midp(X12,X13,X14) [resolution 567,484]
610. eqangle(X0,X1,X2,X3,X4,X5,X2,X3) | ~para(X1,X0,X4,X5) [resolution 371,393]
614. eqangle(X0,X1,X2,X3,X4,X5,X2,X3) [subsumption resolution 610,539]
617. eqangle(X0,X1,X2,X3,X0,X1,X4,X5) [resolution 372,614]
624. eqangle(X8,X9,X10,X11,X9,X8,X12,X13) [resolution 617,371]
633. eqangle(X8,X9,X10,X11,X12,X13,X11,X10) [resolution 624,372]
711. ~midp(sK21,sK21,sK20) | ~midp(sK21,sK22,sK21) [resolution 589,478]
724. 1 <=> ~midp(sK21,sK22,sK21) [avatar definition]
725. ~midp(sK21,sK22,sK21) <- (1) [avatar component clause 724]
730. 3 <=> ~midp(sK21,sK21,sK20) [avatar definition]
731. ~midp(sK21,sK21,sK20) <- (3) [avatar component clause 730]
732. ~1 | ~3 [avatar split clause 711,730,724]
1243. midp(X0,X1,X0) | ~coll(X0,X1,X0) | ~circle(X2,X2,X1,X0) [resolution 405,614]
1245. midp(X6,X7,X8) | ~coll(X6,X7,X8) | ~circle(X8,X6,X7,X8) [resolution 405,633]
1274. midp(X0,X1,X0) | ~circle(X2,X2,X1,X0) [subsumption resolution 1243,476]
1275. midp(X0,X1,X0) [subsumption resolution 1274,477]
1276. midp(X6,X7,X8) | ~circle(X8,X6,X7,X8) [subsumption resolution 1245,476]
1277. midp(X6,X7,X8) [subsumption resolution 1276,477]
1299. $false <- (1) [resolution 1275,725]
1300. 1 [avatar contradiction clause 1299]
1301. $false <- (3) [subsumption resolution 731,1277]
1302. 3 [avatar contradiction clause 1301]
1303. $false [avatar sat refutation 732,1300,1302]
% SZS output end Proof for E059-56
% ------------------------------
% Version: Vampire 4.4.0 (commit 7916d27 on 2019-08-23 08:50:16 +0100)
% Termination reason: Refutation

% Memory used [KB]: 5373
% Time elapsed: 0.037 s
% ------------------------------
% ------------------------------
