% Refutation found. Thanks to Tanya!
% SZS status Theorem for E059-58
% SZS output start Proof for E059-58
fof(f14163,plain,(
  $false),
  inference(avatar_sat_refutation,[],[f2065,f2241,f3575,f8080,f13909,f13911,f13976,f14139])).
fof(f14139,plain,(
  ~spl33_136 | ~spl33_180 | ~spl33_624),
  inference(avatar_contradiction_clause,[],[f14136])).
fof(f14136,plain,(
  $false | (~spl33_136 | ~spl33_180 | ~spl33_624)),
  inference(resolution,[],[f14115,f8525])).
fof(f8525,plain,(
  ~midp(sK27,sK22,sK21) | ~spl33_180),
  inference(subsumption_resolution,[],[f8504,f8019])).
fof(f8019,plain,(
  ( ! [X8,X7] : (midp(X7,X8,X8)) ) | ~spl33_180),
  inference(subsumption_resolution,[],[f7975,f3579])).
fof(f3579,plain,(
  ( ! [X4,X3] : (coll(X3,X4,X4)) ) | ~spl33_180),
  inference(resolution,[],[f2061,f425])).
fof(f425,plain,(
  ( ! [X2,X0,X1] : (~para(X0,X1,X0,X2) | coll(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f267])).
fof(f267,plain,(
  ! [X0,X1,X2] : (coll(X0,X1,X2) | ~para(X0,X1,X0,X2))),
  inference(ennf_transformation,[],[f67])).
fof(f67,axiom,(
  ! [X0,X1,X2] : (para(X0,X1,X0,X2) => coll(X0,X1,X2))),
  file('./7_Book/00EE/08/E059-58.p',ruleD66)).
fof(f2061,plain,(
  ( ! [X2,X3] : (para(X2,X3,X2,X3)) ) | ~spl33_180),
  inference(avatar_component_clause,[],[f2060])).
fof(f2060,plain,(
  spl33_180 <=> ! [X3,X2] : para(X2,X3,X2,X3)),
  introduced(avatar_definition,[new_symbols(naming,[spl33_180])])).
fof(f7975,plain,(
  ( ! [X8,X7] : (~coll(X7,X8,X8) | midp(X7,X8,X8)) ) | ~spl33_180),
  inference(resolution,[],[f7880,f426])).
fof(f426,plain,(
  ( ! [X2,X0,X1] : (~cong(X0,X1,X0,X2) | ~coll(X0,X1,X2) | midp(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f269])).
fof(f269,plain,(
  ! [X0,X1,X2] : (midp(X0,X1,X2) | ~coll(X0,X1,X2) | ~cong(X0,X1,X0,X2))),
  inference(flattening,[],[f268])).
fof(f268,plain,(
  ! [X0,X1,X2] : (midp(X0,X1,X2) | (~coll(X0,X1,X2) | ~cong(X0,X1,X0,X2)))),
  inference(ennf_transformation,[],[f68])).
fof(f68,axiom,(
  ! [X0,X1,X2] : ((coll(X0,X1,X2) & cong(X0,X1,X0,X2)) => midp(X0,X1,X2))),
  file('./7_Book/00EE/08/E059-58.p',ruleD67)).
fof(f7880,plain,(
  ( ! [X35,X36] : (cong(X35,X36,X35,X36)) ) | ~spl33_180),
  inference(subsumption_resolution,[],[f7879,f7790])).
fof(f7790,plain,(
  ( ! [X4,X5,X3] : (cyclic(X3,X4,X3,X5)) ) | ~spl33_180),
  inference(resolution,[],[f7781,f373])).
fof(f373,plain,(
  ( ! [X2,X0,X3,X1] : (~cyclic(X0,X1,X2,X3) | cyclic(X0,X2,X1,X3)) )),
  inference(cnf_transformation,[],[f188])).
fof(f188,plain,(
  ! [X0,X1,X2,X3] : (cyclic(X0,X2,X1,X3) | ~cyclic(X0,X1,X2,X3))),
  inference(ennf_transformation,[],[f15])).
fof(f15,axiom,(
  ! [X0,X1,X2,X3] : (cyclic(X0,X1,X2,X3) => cyclic(X0,X2,X1,X3))),
  file('./7_Book/00EE/08/E059-58.p',ruleD15)).
fof(f7781,plain,(
  ( ! [X6,X8,X7] : (cyclic(X6,X6,X7,X8)) ) | ~spl33_180),
  inference(resolution,[],[f3208,f2061])).
fof(f3208,plain,(
  ( ! [X68,X66,X67] : (~para(X66,X67,X66,X67) | cyclic(X67,X67,X66,X68)) )),
  inference(subsumption_resolution,[],[f2224,f2225])).
fof(f2225,plain,(
  ( ! [X70,X71,X69] : (~para(X69,X70,X69,X70) | coll(X69,X71,X70) | cyclic(X70,X70,X69,X71)) )),
  inference(resolution,[],[f858,f400])).
fof(f400,plain,(
  ( ! [X2,X0,X3,X1] : (~eqangle(X2,X0,X2,X1,X3,X0,X3,X1) | coll(X2,X3,X0) | cyclic(X0,X1,X2,X3)) )),
  inference(cnf_transformation,[],[f222])).
fof(f222,plain,(
  ! [X0,X1,X2,X3] : (cyclic(X0,X1,X2,X3) | coll(X2,X3,X0) | ~eqangle(X2,X0,X2,X1,X3,X0,X3,X1))),
  inference(flattening,[],[f221])).
fof(f221,plain,(
  ! [X0,X1,X2,X3] : (cyclic(X0,X1,X2,X3) | (coll(X2,X3,X0) | ~eqangle(X2,X0,X2,X1,X3,X0,X3,X1)))),
  inference(ennf_transformation,[],[f121])).
fof(f121,plain,(
  ! [X0,X1,X2,X3] : ((~coll(X2,X3,X0) & eqangle(X2,X0,X2,X1,X3,X0,X3,X1)) => cyclic(X0,X1,X2,X3))),
  inference(rectify,[],[f42])).
fof(f42,axiom,(
  ! [X0,X1,X8,X9] : ((~coll(X8,X9,X0) & eqangle(X8,X0,X8,X1,X9,X0,X9,X1)) => cyclic(X0,X1,X8,X9))),
  file('./7_Book/00EE/08/E059-58.p',ruleD42a)).
fof(f858,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (eqangle(X0,X1,X2,X3,X4,X5,X4,X5) | ~para(X0,X1,X2,X3)) )),
  inference(resolution,[],[f379,f398])).
fof(f398,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (eqangle(X0,X1,X4,X5,X2,X3,X4,X5) | ~para(X0,X1,X2,X3)) )),
  inference(cnf_transformation,[],[f219])).
fof(f219,plain,(
  ! [X0,X1,X2,X3,X4,X5] : (eqangle(X0,X1,X4,X5,X2,X3,X4,X5) | ~para(X0,X1,X2,X3))),
  inference(ennf_transformation,[],[f119])).
fof(f119,plain,(
  ! [X0,X1,X2,X3,X4,X5] : (para(X0,X1,X2,X3) => eqangle(X0,X1,X4,X5,X2,X3,X4,X5))),
  inference(rectify,[],[f40])).
fof(f40,axiom,(
  ! [X0,X1,X2,X3,X8,X9] : (para(X0,X1,X2,X3) => eqangle(X0,X1,X8,X9,X2,X3,X8,X9))),
  file('./7_Book/00EE/08/E059-58.p',ruleD40)).
fof(f379,plain,(
  ( ! [X6,X4,X2,X0,X7,X5,X3,X1] : (~eqangle(X0,X1,X2,X3,X4,X5,X6,X7) | eqangle(X0,X1,X4,X5,X2,X3,X6,X7)) )),
  inference(cnf_transformation,[],[f195])).
fof(f195,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6,X7] : (eqangle(X0,X1,X4,X5,X2,X3,X6,X7) | ~eqangle(X0,X1,X2,X3,X4,X5,X6,X7))),
  inference(ennf_transformation,[],[f103])).
fof(f103,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6,X7] : (eqangle(X0,X1,X2,X3,X4,X5,X6,X7) => eqangle(X0,X1,X4,X5,X2,X3,X6,X7))),
  inference(rectify,[],[f21])).
fof(f21,axiom,(
  ! [X0,X1,X2,X3,X8,X9,X10,X11] : (eqangle(X0,X1,X2,X3,X8,X9,X10,X11) => eqangle(X0,X1,X8,X9,X2,X3,X10,X11))),
  file('./7_Book/00EE/08/E059-58.p',ruleD21)).
fof(f2224,plain,(
  ( ! [X68,X66,X67] : (~para(X66,X67,X66,X67) | ~coll(X66,X68,X67) | cyclic(X67,X67,X66,X68)) )),
  inference(resolution,[],[f858,f401])).
fof(f401,plain,(
  ( ! [X2,X0,X3,X1] : (~eqangle(X2,X0,X2,X1,X3,X0,X3,X1) | ~coll(X2,X3,X1) | cyclic(X0,X1,X2,X3)) )),
  inference(cnf_transformation,[],[f224])).
fof(f224,plain,(
  ! [X0,X1,X2,X3] : (cyclic(X0,X1,X2,X3) | ~coll(X2,X3,X1) | ~eqangle(X2,X0,X2,X1,X3,X0,X3,X1))),
  inference(flattening,[],[f223])).
fof(f223,plain,(
  ! [X0,X1,X2,X3] : (cyclic(X0,X1,X2,X3) | (~coll(X2,X3,X1) | ~eqangle(X2,X0,X2,X1,X3,X0,X3,X1)))),
  inference(ennf_transformation,[],[f122])).
fof(f122,plain,(
  ! [X0,X1,X2,X3] : ((coll(X2,X3,X1) & eqangle(X2,X0,X2,X1,X3,X0,X3,X1)) => cyclic(X0,X1,X2,X3))),
  inference(rectify,[],[f43])).
fof(f43,axiom,(
  ! [X0,X1,X8,X9] : ((coll(X8,X9,X1) & eqangle(X8,X0,X8,X1,X9,X0,X9,X1)) => cyclic(X0,X1,X8,X9))),
  file('./7_Book/00EE/08/E059-58.p',ruleD42b)).
fof(f7879,plain,(
  ( ! [X35,X36] : (cong(X35,X36,X35,X36) | ~cyclic(X35,X36,X35,X35)) ) | ~spl33_180),
  inference(subsumption_resolution,[],[f7865,f7790])).
fof(f7865,plain,(
  ( ! [X35,X36] : (cong(X35,X36,X35,X36) | ~cyclic(X35,X36,X35,X36) | ~cyclic(X35,X36,X35,X35)) ) | ~spl33_180),
  inference(resolution,[],[f7790,f1354])).
fof(f1354,plain,(
  ( ! [X10,X8,X7,X9] : (~cyclic(X7,X8,X9,X10) | cong(X7,X8,X7,X8) | ~cyclic(X7,X8,X9,X8) | ~cyclic(X7,X8,X9,X7)) )),
  inference(duplicate_literal_removal,[],[f1348])).
fof(f1348,plain,(
  ( ! [X10,X8,X7,X9] : (cong(X7,X8,X7,X8) | ~cyclic(X7,X8,X9,X10) | ~cyclic(X7,X8,X9,X8) | ~cyclic(X7,X8,X9,X7) | ~cyclic(X7,X8,X9,X10)) )),
  inference(resolution,[],[f402,f399])).
fof(f399,plain,(
  ( ! [X2,X0,X3,X1] : (eqangle(X2,X0,X2,X1,X3,X0,X3,X1) | ~cyclic(X0,X1,X2,X3)) )),
  inference(cnf_transformation,[],[f220])).
fof(f220,plain,(
  ! [X0,X1,X2,X3] : (eqangle(X2,X0,X2,X1,X3,X0,X3,X1) | ~cyclic(X0,X1,X2,X3))),
  inference(ennf_transformation,[],[f120])).
fof(f120,plain,(
  ! [X0,X1,X2,X3] : (cyclic(X0,X1,X2,X3) => eqangle(X2,X0,X2,X1,X3,X0,X3,X1))),
  inference(rectify,[],[f41])).
fof(f41,axiom,(
  ! [X0,X1,X8,X9] : (cyclic(X0,X1,X8,X9) => eqangle(X8,X0,X8,X1,X9,X0,X9,X1))),
  file('./7_Book/00EE/08/E059-58.p',ruleD41)).
fof(f402,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (~eqangle(X2,X0,X2,X1,X5,X3,X5,X4) | cong(X0,X1,X3,X4) | ~cyclic(X0,X1,X2,X5) | ~cyclic(X0,X1,X2,X4) | ~cyclic(X0,X1,X2,X3)) )),
  inference(cnf_transformation,[],[f226])).
fof(f226,plain,(
  ! [X0,X1,X2,X3,X4,X5] : (cong(X0,X1,X3,X4) | ~eqangle(X2,X0,X2,X1,X5,X3,X5,X4) | ~cyclic(X0,X1,X2,X5) | ~cyclic(X0,X1,X2,X4) | ~cyclic(X0,X1,X2,X3))),
  inference(flattening,[],[f225])).
fof(f225,plain,(
  ! [X0,X1,X2,X3,X4,X5] : (cong(X0,X1,X3,X4) | (~eqangle(X2,X0,X2,X1,X5,X3,X5,X4) | ~cyclic(X0,X1,X2,X5) | ~cyclic(X0,X1,X2,X4) | ~cyclic(X0,X1,X2,X3)))),
  inference(ennf_transformation,[],[f123])).
fof(f123,plain,(
  ! [X0,X1,X2,X3,X4,X5] : ((eqangle(X2,X0,X2,X1,X5,X3,X5,X4) & cyclic(X0,X1,X2,X5) & cyclic(X0,X1,X2,X4) & cyclic(X0,X1,X2,X3)) => cong(X0,X1,X3,X4))),
  inference(rectify,[],[f44])).
fof(f44,axiom,(
  ! [X0,X1,X2,X8,X9,X14] : ((eqangle(X2,X0,X2,X1,X14,X8,X14,X9) & cyclic(X0,X1,X2,X14) & cyclic(X0,X1,X2,X9) & cyclic(X0,X1,X2,X8)) => cong(X0,X1,X8,X9))),
  file('./7_Book/00EE/08/E059-58.p',ruleD43)).
fof(f8504,plain,(
  ~midp(sK27,sK22,sK21) | ~midp(sK22,sK27,sK27)),
  inference(resolution,[],[f1933,f483])).
fof(f483,plain,(
  ~cong(sK22,sK27,sK27,sK21)),
  inference(cnf_transformation,[],[f358])).
fof(f358,plain,(
  ~cong(sK22,sK27,sK27,sK21) & coll(sK27,sK25,sK26) & coll(sK27,sK21,sK22) & perp(sK26,sK20,sK26,sK24) & circle(sK20,sK21,sK26,sK32) & perp(sK25,sK20,sK25,sK23) & circle(sK20,sK21,sK25,sK31) & cong(sK24,sK22,sK21,sK23) & coll(sK24,sK21,sK22) & coll(sK23,sK21,sK22) & circle(sK20,sK21,sK22,sK30) & circle(sK20,sK21,sK28,sK29)),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK20,sK21,sK22,sK23,sK24,sK25,sK26,sK27,sK28,sK29,sK30,sK31,sK32])],[f320,f357])).
fof(f357,plain,(
  ? [X0,X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12] : (~cong(X2,X7,X7,X1) & coll(X7,X5,X6) & coll(X7,X1,X2) & perp(X6,X0,X6,X4) & circle(X0,X1,X6,X12) & perp(X5,X0,X5,X3) & circle(X0,X1,X5,X11) & cong(X4,X2,X1,X3) & coll(X4,X1,X2) & coll(X3,X1,X2) & circle(X0,X1,X2,X10) & circle(X0,X1,X8,X9)) => (~cong(sK22,sK27,sK27,sK21) & coll(sK27,sK25,sK26) & coll(sK27,sK21,sK22) & perp(sK26,sK20,sK26,sK24) & circle(sK20,sK21,sK26,sK32) & perp(sK25,sK20,sK25,sK23) & circle(sK20,sK21,sK25,sK31) & cong(sK24,sK22,sK21,sK23) & coll(sK24,sK21,sK22) & coll(sK23,sK21,sK22) & circle(sK20,sK21,sK22,sK30) & circle(sK20,sK21,sK28,sK29))),
  introduced(choice_axiom,[])).
fof(f320,plain,(
  ? [X0,X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12] : (~cong(X2,X7,X7,X1) & coll(X7,X5,X6) & coll(X7,X1,X2) & perp(X6,X0,X6,X4) & circle(X0,X1,X6,X12) & perp(X5,X0,X5,X3) & circle(X0,X1,X5,X11) & cong(X4,X2,X1,X3) & coll(X4,X1,X2) & coll(X3,X1,X2) & circle(X0,X1,X2,X10) & circle(X0,X1,X8,X9))),
  inference(flattening,[],[f319])).
fof(f319,plain,(
  ? [X0,X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12] : (~cong(X2,X7,X7,X1) & (coll(X7,X5,X6) & coll(X7,X1,X2) & perp(X6,X0,X6,X4) & circle(X0,X1,X6,X12) & perp(X5,X0,X5,X3) & circle(X0,X1,X5,X11) & cong(X4,X2,X1,X3) & coll(X4,X1,X2) & coll(X3,X1,X2) & circle(X0,X1,X2,X10) & circle(X0,X1,X8,X9)))),
  inference(ennf_transformation,[],[f167])).
fof(f167,plain,(
  ~! [X0,X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12] : ((coll(X7,X5,X6) & coll(X7,X1,X2) & perp(X6,X0,X6,X4) & circle(X0,X1,X6,X12) & perp(X5,X0,X5,X3) & circle(X0,X1,X5,X11) & cong(X4,X2,X1,X3) & coll(X4,X1,X2) & coll(X3,X1,X2) & circle(X0,X1,X2,X10) & circle(X0,X1,X8,X9)) => cong(X2,X7,X7,X1))),
  inference(rectify,[],[f96])).
fof(f96,negated_conjecture,(
  ~! [X0,X1,X2,X3,X4,X5,X12,X13,X18,X19,X20,X21,X22] : ((coll(X13,X5,X12) & coll(X13,X1,X2) & perp(X12,X0,X12,X4) & circle(X0,X1,X12,X22) & perp(X5,X0,X5,X3) & circle(X0,X1,X5,X21) & cong(X4,X2,X1,X3) & coll(X4,X1,X2) & coll(X3,X1,X2) & circle(X0,X1,X2,X20) & circle(X0,X1,X18,X19)) => cong(X2,X13,X13,X1))),
  inference(negated_conjecture,[],[f95])).
fof(f95,conjecture,(
  ! [X0,X1,X2,X3,X4,X5,X12,X13,X18,X19,X20,X21,X22] : ((coll(X13,X5,X12) & coll(X13,X1,X2) & perp(X12,X0,X12,X4) & circle(X0,X1,X12,X22) & perp(X5,X0,X5,X3) & circle(X0,X1,X5,X21) & cong(X4,X2,X1,X3) & coll(X4,X1,X2) & coll(X3,X1,X2) & circle(X0,X1,X2,X20) & circle(X0,X1,X18,X19)) => cong(X2,X13,X13,X1))),
  file('./7_Book/00EE/08/E059-58.p',exemplo7Book00EE08E05958)).
fof(f1933,plain,(
  ( ! [X14,X12,X15,X13] : (cong(X12,X13,X14,X15) | ~midp(X14,X12,X15) | ~midp(X12,X13,X14)) )),
  inference(resolution,[],[f815,f521])).
fof(f521,plain,(
  ( ! [X2,X0,X1] : (cong(X0,X1,X2,X0) | ~midp(X0,X1,X2)) )),
  inference(resolution,[],[f381,f427])).
fof(f427,plain,(
  ( ! [X2,X0,X1] : (cong(X0,X1,X0,X2) | ~midp(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f270])).
fof(f270,plain,(
  ! [X0,X1,X2] : (cong(X0,X1,X0,X2) | ~midp(X0,X1,X2))),
  inference(ennf_transformation,[],[f69])).
fof(f69,axiom,(
  ! [X0,X1,X2] : (midp(X0,X1,X2) => cong(X0,X1,X0,X2))),
  file('./7_Book/00EE/08/E059-58.p',ruleD68)).
fof(f381,plain,(
  ( ! [X2,X0,X3,X1] : (~cong(X0,X1,X2,X3) | cong(X0,X1,X3,X2)) )),
  inference(cnf_transformation,[],[f198])).
fof(f198,plain,(
  ! [X0,X1,X2,X3] : (cong(X0,X1,X3,X2) | ~cong(X0,X1,X2,X3))),
  inference(ennf_transformation,[],[f23])).
fof(f23,axiom,(
  ! [X0,X1,X2,X3] : (cong(X0,X1,X2,X3) => cong(X0,X1,X3,X2))),
  file('./7_Book/00EE/08/E059-58.p',ruleD23)).
fof(f815,plain,(
  ( ! [X6,X4,X8,X7,X5] : (~cong(X4,X5,X6,X8) | cong(X4,X5,X6,X7) | ~midp(X6,X8,X7)) )),
  inference(resolution,[],[f383,f427])).
fof(f383,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (~cong(X2,X3,X4,X5) | cong(X0,X1,X4,X5) | ~cong(X0,X1,X2,X3)) )),
  inference(cnf_transformation,[],[f201])).
fof(f201,plain,(
  ! [X0,X1,X2,X3,X4,X5] : (cong(X0,X1,X4,X5) | ~cong(X2,X3,X4,X5) | ~cong(X0,X1,X2,X3))),
  inference(flattening,[],[f200])).
fof(f200,plain,(
  ! [X0,X1,X2,X3,X4,X5] : (cong(X0,X1,X4,X5) | (~cong(X2,X3,X4,X5) | ~cong(X0,X1,X2,X3)))),
  inference(ennf_transformation,[],[f25])).
fof(f25,axiom,(
  ! [X0,X1,X2,X3,X4,X5] : ((cong(X2,X3,X4,X5) & cong(X0,X1,X2,X3)) => cong(X0,X1,X4,X5))),
  file('./7_Book/00EE/08/E059-58.p',ruleD25)).
fof(f14115,plain,(
  ( ! [X4] : (midp(X4,sK22,sK21)) ) | (~spl33_136 | ~spl33_180 | ~spl33_624)),
  inference(subsumption_resolution,[],[f14114,f8019])).
fof(f14114,plain,(
  ( ! [X4] : (midp(X4,sK22,sK21) | ~midp(X4,sK30,sK30)) ) | (~spl33_136 | ~spl33_624)),
  inference(subsumption_resolution,[],[f14103,f8878])).
fof(f8878,plain,(
  para(sK30,sK22,sK30,sK21) | ~spl33_624),
  inference(avatar_component_clause,[],[f8877])).
fof(f8877,plain,(
  spl33_624 <=> para(sK30,sK22,sK30,sK21)),
  introduced(avatar_definition,[new_symbols(naming,[spl33_624])])).
fof(f14103,plain,(
  ( ! [X4] : (midp(X4,sK22,sK21) | ~para(sK30,sK22,sK30,sK21) | ~midp(X4,sK30,sK30)) ) | ~spl33_136),
  inference(resolution,[],[f1732,f423])).
fof(f423,plain,(
  ( ! [X4,X2,X0,X3,X1] : (~para(X0,X3,X1,X2) | midp(X4,X2,X3) | ~para(X0,X2,X1,X3) | ~midp(X4,X0,X1)) )),
  inference(cnf_transformation,[],[f264])).
fof(f264,plain,(
  ! [X0,X1,X2,X3,X4] : (midp(X4,X2,X3) | ~para(X0,X3,X1,X2) | ~para(X0,X2,X1,X3) | ~midp(X4,X0,X1))),
  inference(flattening,[],[f263])).
fof(f263,plain,(
  ! [X0,X1,X2,X3,X4] : (midp(X4,X2,X3) | (~para(X0,X3,X1,X2) | ~para(X0,X2,X1,X3) | ~midp(X4,X0,X1)))),
  inference(ennf_transformation,[],[f143])).
fof(f143,plain,(
  ! [X0,X1,X2,X3,X4] : ((para(X0,X3,X1,X2) & para(X0,X2,X1,X3) & midp(X4,X0,X1)) => midp(X4,X2,X3))),
  inference(rectify,[],[f65])).
fof(f65,axiom,(
  ! [X0,X1,X2,X3,X6] : ((para(X0,X3,X1,X2) & para(X0,X2,X1,X3) & midp(X6,X0,X1)) => midp(X6,X2,X3))),
  file('./7_Book/00EE/08/E059-58.p',ruleD64)).
fof(f1732,plain,(
  para(sK30,sK21,sK30,sK22) | ~spl33_136),
  inference(avatar_component_clause,[],[f1731])).
fof(f1731,plain,(
  spl33_136 <=> para(sK30,sK21,sK30,sK22)),
  introduced(avatar_definition,[new_symbols(naming,[spl33_136])])).
fof(f13976,plain,(
  ~spl33_180 | spl33_565),
  inference(avatar_contradiction_clause,[],[f13975])).
fof(f13975,plain,(
  $false | (~spl33_180 | ~spl33_565)),
  inference(subsumption_resolution,[],[f7632,f8065])).
fof(f8065,plain,(
  ( ! [X33,X32] : (perp(X32,X33,X33,X33)) ) | ~spl33_180),
  inference(resolution,[],[f7944,f366])).
fof(f366,plain,(
  ( ! [X2,X0,X3,X1] : (~perp(X0,X1,X2,X3) | perp(X2,X3,X0,X1)) )),
  inference(cnf_transformation,[],[f177])).
fof(f177,plain,(
  ! [X0,X1,X2,X3] : (perp(X2,X3,X0,X1) | ~perp(X0,X1,X2,X3))),
  inference(ennf_transformation,[],[f8])).
fof(f8,axiom,(
  ! [X0,X1,X2,X3] : (perp(X0,X1,X2,X3) => perp(X2,X3,X0,X1))),
  file('./7_Book/00EE/08/E059-58.p',ruleD8)).
fof(f7944,plain,(
  ( ! [X47,X46] : (perp(X46,X46,X47,X46)) ) | ~spl33_180),
  inference(subsumption_resolution,[],[f7936,f7880])).
fof(f7936,plain,(
  ( ! [X47,X46] : (perp(X46,X46,X47,X46) | ~cong(X46,X47,X46,X47)) ) | ~spl33_180),
  inference(resolution,[],[f7816,f415])).
fof(f415,plain,(
  ( ! [X2,X0,X3,X1] : (~cong(X0,X3,X1,X3) | perp(X0,X1,X2,X3) | ~cong(X0,X2,X1,X2)) )),
  inference(cnf_transformation,[],[f251])).
fof(f251,plain,(
  ! [X0,X1,X2,X3] : (perp(X0,X1,X2,X3) | ~cong(X0,X3,X1,X3) | ~cong(X0,X2,X1,X2))),
  inference(flattening,[],[f250])).
fof(f250,plain,(
  ! [X0,X1,X2,X3] : (perp(X0,X1,X2,X3) | (~cong(X0,X3,X1,X3) | ~cong(X0,X2,X1,X2)))),
  inference(ennf_transformation,[],[f135])).
fof(f135,plain,(
  ! [X0,X1,X2,X3] : ((cong(X0,X3,X1,X3) & cong(X0,X2,X1,X2)) => perp(X0,X1,X2,X3))),
  inference(rectify,[],[f57])).
fof(f57,axiom,(
  ! [X0,X1,X8,X9] : ((cong(X0,X9,X1,X9) & cong(X0,X8,X1,X8)) => perp(X0,X1,X8,X9))),
  file('./7_Book/00EE/08/E059-58.p',ruleD56)).
fof(f7816,plain,(
  ( ! [X35] : (cong(X35,X35,X35,X35)) ) | ~spl33_180),
  inference(subsumption_resolution,[],[f7813,f7781])).
fof(f7813,plain,(
  ( ! [X35,X36] : (cong(X35,X35,X35,X35) | ~cyclic(X35,X35,X36,X35)) ) | ~spl33_180),
  inference(duplicate_literal_removal,[],[f7800])).
fof(f7800,plain,(
  ( ! [X35,X36] : (cong(X35,X35,X35,X35) | ~cyclic(X35,X35,X36,X35) | ~cyclic(X35,X35,X36,X35)) ) | ~spl33_180),
  inference(resolution,[],[f7781,f1354])).
fof(f7632,plain,(
  ~perp(sK22,sK24,sK24,sK24) | ~spl33_565),
  inference(avatar_component_clause,[],[f7631])).
fof(f7631,plain,(
  spl33_565 <=> ~perp(sK22,sK24,sK24,sK24)),
  introduced(avatar_definition,[new_symbols(naming,[spl33_565])])).
fof(f13911,plain,(
  spl33_137 | ~spl33_180 | ~spl33_186 | ~spl33_564),
  inference(avatar_contradiction_clause,[],[f13894])).
fof(f13894,plain,(
  $false | (~spl33_137 | ~spl33_180 | ~spl33_186 | ~spl33_564)),
  inference(resolution,[],[f13036,f1735])).
fof(f1735,plain,(
  ~para(sK30,sK21,sK30,sK22) | ~spl33_137),
  inference(avatar_component_clause,[],[f1734])).
fof(f1734,plain,(
  spl33_137 <=> ~para(sK30,sK21,sK30,sK22)),
  introduced(avatar_definition,[new_symbols(naming,[spl33_137])])).
fof(f13036,plain,(
  ( ! [X39,X37,X38,X36] : (para(X38,X39,X36,X37)) ) | (~spl33_180 | ~spl33_186 | ~spl33_564)),
  inference(subsumption_resolution,[],[f12954,f12708])).
fof(f12708,plain,(
  ( ! [X216,X215] : (para(X216,X215,sK24,sK26)) ) | (~spl33_180 | ~spl33_186 | ~spl33_564)),
  inference(subsumption_resolution,[],[f12680,f11827])).
fof(f11827,plain,(
  ( ! [X0,X1] : (para(X0,X1,sK26,sK20)) ) | (~spl33_180 | ~spl33_186 | ~spl33_564)),
  inference(subsumption_resolution,[],[f11822,f11577])).
fof(f11577,plain,(
  ( ! [X12,X11] : (para(X11,X12,sK24,sK24)) ) | (~spl33_180 | ~spl33_186 | ~spl33_564)),
  inference(subsumption_resolution,[],[f11573,f11446])).
fof(f11446,plain,(
  ( ! [X88,X87,X89,X86] : (perp(X86,X87,X88,X89)) ) | (~spl33_180 | ~spl33_186)),
  inference(subsumption_resolution,[],[f11420,f9486])).
fof(f9486,plain,(
  ( ! [X19,X17,X18] : (cong(X18,X17,X19,X17)) ) | (~spl33_180 | ~spl33_186)),
  inference(subsumption_resolution,[],[f9453,f8019])).
fof(f9453,plain,(
  ( ! [X19,X17,X18] : (~midp(X17,X18,X18) | cong(X18,X17,X19,X17)) ) | (~spl33_180 | ~spl33_186)),
  inference(resolution,[],[f8088,f411])).
fof(f411,plain,(
  ( ! [X2,X0,X3,X1] : (~perp(X0,X1,X1,X2) | ~midp(X3,X0,X2) | cong(X0,X3,X1,X3)) )),
  inference(cnf_transformation,[],[f243])).
fof(f243,plain,(
  ! [X0,X1,X2,X3] : (cong(X0,X3,X1,X3) | ~midp(X3,X0,X2) | ~perp(X0,X1,X1,X2))),
  inference(flattening,[],[f242])).
fof(f242,plain,(
  ! [X0,X1,X2,X3] : (cong(X0,X3,X1,X3) | (~midp(X3,X0,X2) | ~perp(X0,X1,X1,X2)))),
  inference(ennf_transformation,[],[f132])).
fof(f132,plain,(
  ! [X0,X1,X2,X3] : ((midp(X3,X0,X2) & perp(X0,X1,X1,X2)) => cong(X0,X3,X1,X3))),
  inference(rectify,[],[f53])).
fof(f53,axiom,(
  ! [X0,X1,X2,X6] : ((midp(X6,X0,X2) & perp(X0,X1,X1,X2)) => cong(X0,X6,X1,X6))),
  file('./7_Book/00EE/08/E059-58.p',ruleD52)).
fof(f8088,plain,(
  ( ! [X21,X20] : (perp(X20,X21,X21,X20)) ) | (~spl33_180 | ~spl33_186)),
  inference(resolution,[],[f2240,f3586])).
fof(f3586,plain,(
  ( ! [X24,X23] : (para(X23,X24,X24,X23)) ) | ~spl33_180),
  inference(resolution,[],[f2061,f362])).
fof(f362,plain,(
  ( ! [X2,X0,X3,X1] : (~para(X0,X1,X2,X3) | para(X0,X1,X3,X2)) )),
  inference(cnf_transformation,[],[f172])).
fof(f172,plain,(
  ! [X0,X1,X2,X3] : (para(X0,X1,X3,X2) | ~para(X0,X1,X2,X3))),
  inference(ennf_transformation,[],[f4])).
fof(f4,axiom,(
  ! [X0,X1,X2,X3] : (para(X0,X1,X2,X3) => para(X0,X1,X3,X2))),
  file('./7_Book/00EE/08/E059-58.p',ruleD4)).
fof(f2240,plain,(
  ( ! [X47,X45,X46,X44] : (~para(X44,X45,X46,X47) | perp(X44,X45,X46,X47)) ) | ~spl33_186),
  inference(avatar_component_clause,[],[f2239])).
fof(f2239,plain,(
  spl33_186 <=> ! [X44,X46,X45,X47] : (~para(X44,X45,X46,X47) | perp(X44,X45,X46,X47))),
  introduced(avatar_definition,[new_symbols(naming,[spl33_186])])).
fof(f11420,plain,(
  ( ! [X88,X87,X89,X86] : (perp(X86,X87,X88,X89) | ~cong(X86,X88,X87,X88)) ) | (~spl33_180 | ~spl33_186)),
  inference(resolution,[],[f9486,f415])).
fof(f11573,plain,(
  ( ! [X12,X11] : (para(X11,X12,sK24,sK24) | ~perp(X11,X12,sK22,sK24)) ) | ~spl33_564),
  inference(resolution,[],[f7635,f367])).
fof(f367,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (~perp(X2,X3,X4,X5) | para(X0,X1,X4,X5) | ~perp(X0,X1,X2,X3)) )),
  inference(cnf_transformation,[],[f179])).
fof(f179,plain,(
  ! [X0,X1,X2,X3,X4,X5] : (para(X0,X1,X4,X5) | ~perp(X2,X3,X4,X5) | ~perp(X0,X1,X2,X3))),
  inference(flattening,[],[f178])).
fof(f178,plain,(
  ! [X0,X1,X2,X3,X4,X5] : (para(X0,X1,X4,X5) | (~perp(X2,X3,X4,X5) | ~perp(X0,X1,X2,X3)))),
  inference(ennf_transformation,[],[f9])).
fof(f9,axiom,(
  ! [X0,X1,X2,X3,X4,X5] : ((perp(X2,X3,X4,X5) & perp(X0,X1,X2,X3)) => para(X0,X1,X4,X5))),
  file('./7_Book/00EE/08/E059-58.p',ruleD9)).
fof(f7635,plain,(
  perp(sK22,sK24,sK24,sK24) | ~spl33_564),
  inference(avatar_component_clause,[],[f7634])).
fof(f7634,plain,(
  spl33_564 <=> perp(sK22,sK24,sK24,sK24)),
  introduced(avatar_definition,[new_symbols(naming,[spl33_564])])).
fof(f11822,plain,(
  ( ! [X0,X1] : (para(X0,X1,sK26,sK20) | ~para(X0,X1,sK24,sK24)) ) | ~spl33_180),
  inference(resolution,[],[f8079,f364])).
fof(f364,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (~para(X2,X3,X4,X5) | para(X0,X1,X4,X5) | ~para(X0,X1,X2,X3)) )),
  inference(cnf_transformation,[],[f175])).
fof(f175,plain,(
  ! [X0,X1,X2,X3,X4,X5] : (para(X0,X1,X4,X5) | ~para(X2,X3,X4,X5) | ~para(X0,X1,X2,X3))),
  inference(flattening,[],[f174])).
fof(f174,plain,(
  ! [X0,X1,X2,X3,X4,X5] : (para(X0,X1,X4,X5) | (~para(X2,X3,X4,X5) | ~para(X0,X1,X2,X3)))),
  inference(ennf_transformation,[],[f6])).
fof(f6,axiom,(
  ! [X0,X1,X2,X3,X4,X5] : ((para(X2,X3,X4,X5) & para(X0,X1,X2,X3)) => para(X0,X1,X4,X5))),
  file('./7_Book/00EE/08/E059-58.p',ruleD6)).
fof(f8079,plain,(
  para(sK24,sK24,sK26,sK20) | ~spl33_180),
  inference(resolution,[],[f7944,f797])).
fof(f797,plain,(
  ( ! [X19,X18] : (~perp(X18,X19,sK26,sK24) | para(X18,X19,sK26,sK20)) )),
  inference(resolution,[],[f367,f511])).
fof(f511,plain,(
  perp(sK26,sK24,sK26,sK20)),
  inference(resolution,[],[f366,f480])).
fof(f480,plain,(
  perp(sK26,sK20,sK26,sK24)),
  inference(cnf_transformation,[],[f358])).
fof(f12680,plain,(
  ( ! [X216,X215] : (~para(X215,X216,sK26,sK20) | para(X216,X215,sK24,sK26)) ) | (~spl33_180 | ~spl33_186)),
  inference(resolution,[],[f1960,f8136])).
fof(f8136,plain,(
  para(sK26,sK20,sK24,sK26) | (~spl33_180 | ~spl33_186)),
  inference(resolution,[],[f8087,f794])).
fof(f794,plain,(
  ( ! [X12,X13] : (~perp(X12,X13,sK26,sK20) | para(X12,X13,sK24,sK26)) )),
  inference(resolution,[],[f367,f507])).
fof(f507,plain,(
  perp(sK26,sK20,sK24,sK26)),
  inference(resolution,[],[f365,f480])).
fof(f365,plain,(
  ( ! [X2,X0,X3,X1] : (~perp(X0,X1,X2,X3) | perp(X0,X1,X3,X2)) )),
  inference(cnf_transformation,[],[f176])).
fof(f176,plain,(
  ! [X0,X1,X2,X3] : (perp(X0,X1,X3,X2) | ~perp(X0,X1,X2,X3))),
  inference(ennf_transformation,[],[f7])).
fof(f7,axiom,(
  ! [X0,X1,X2,X3] : (perp(X0,X1,X2,X3) => perp(X0,X1,X3,X2))),
  file('./7_Book/00EE/08/E059-58.p',ruleD7)).
fof(f8087,plain,(
  ( ! [X19,X18] : (perp(X18,X19,X18,X19)) ) | (~spl33_180 | ~spl33_186)),
  inference(resolution,[],[f2240,f2061])).
fof(f1960,plain,(
  ( ! [X39,X43,X41,X38,X42,X40] : (~para(X40,X41,X42,X43) | ~para(X38,X39,X40,X41) | para(X39,X38,X42,X43)) )),
  inference(resolution,[],[f825,f432])).
fof(f432,plain,(
  ( ! [X6,X4,X2,X0,X7,X5,X3,X1] : (~eqangle(X0,X1,X2,X3,X4,X5,X6,X7) | ~para(X4,X5,X6,X7) | para(X0,X1,X2,X3)) )),
  inference(cnf_transformation,[],[f279])).
fof(f279,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6,X7] : (para(X0,X1,X2,X3) | ~para(X4,X5,X6,X7) | ~eqangle(X0,X1,X2,X3,X4,X5,X6,X7))),
  inference(flattening,[],[f278])).
fof(f278,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6,X7] : (para(X0,X1,X2,X3) | (~para(X4,X5,X6,X7) | ~eqangle(X0,X1,X2,X3,X4,X5,X6,X7)))),
  inference(ennf_transformation,[],[f146])).
fof(f146,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6,X7] : ((para(X4,X5,X6,X7) & eqangle(X0,X1,X2,X3,X4,X5,X6,X7)) => para(X0,X1,X2,X3))),
  inference(rectify,[],[f74])).
fof(f74,axiom,(
  ! [X0,X1,X2,X3,X8,X9,X10,X11] : ((para(X8,X9,X10,X11) & eqangle(X0,X1,X2,X3,X8,X9,X10,X11)) => para(X0,X1,X2,X3))),
  file('./7_Book/00EE/08/E059-58.p',ruleD73)).
fof(f825,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (eqangle(X0,X1,X2,X3,X4,X5,X2,X3) | ~para(X1,X0,X4,X5)) )),
  inference(resolution,[],[f376,f398])).
fof(f376,plain,(
  ( ! [X6,X4,X2,X0,X7,X5,X3,X1] : (~eqangle(X0,X1,X2,X3,X4,X5,X6,X7) | eqangle(X1,X0,X2,X3,X4,X5,X6,X7)) )),
  inference(cnf_transformation,[],[f192])).
fof(f192,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6,X7] : (eqangle(X1,X0,X2,X3,X4,X5,X6,X7) | ~eqangle(X0,X1,X2,X3,X4,X5,X6,X7))),
  inference(ennf_transformation,[],[f100])).
fof(f100,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6,X7] : (eqangle(X0,X1,X2,X3,X4,X5,X6,X7) => eqangle(X1,X0,X2,X3,X4,X5,X6,X7))),
  inference(rectify,[],[f18])).
fof(f18,axiom,(
  ! [X0,X1,X2,X3,X8,X9,X10,X11] : (eqangle(X0,X1,X2,X3,X8,X9,X10,X11) => eqangle(X1,X0,X2,X3,X8,X9,X10,X11))),
  file('./7_Book/00EE/08/E059-58.p',ruleD18)).
fof(f12954,plain,(
  ( ! [X39,X37,X38,X36] : (~para(X36,X37,sK24,sK26) | para(X38,X39,X36,X37)) ) | (~spl33_180 | ~spl33_186 | ~spl33_564)),
  inference(resolution,[],[f2042,f12708])).
fof(f2042,plain,(
  ( ! [X39,X43,X41,X38,X42,X40] : (~para(X42,X43,X40,X41) | ~para(X38,X39,X40,X41) | para(X42,X43,X38,X39)) )),
  inference(resolution,[],[f832,f432])).
fof(f832,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (eqangle(X0,X1,X2,X3,X0,X1,X4,X5) | ~para(X2,X3,X4,X5)) )),
  inference(resolution,[],[f377,f398])).
fof(f377,plain,(
  ( ! [X6,X4,X2,X0,X7,X5,X3,X1] : (~eqangle(X0,X1,X2,X3,X4,X5,X6,X7) | eqangle(X2,X3,X0,X1,X6,X7,X4,X5)) )),
  inference(cnf_transformation,[],[f193])).
fof(f193,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6,X7] : (eqangle(X2,X3,X0,X1,X6,X7,X4,X5) | ~eqangle(X0,X1,X2,X3,X4,X5,X6,X7))),
  inference(ennf_transformation,[],[f101])).
fof(f101,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6,X7] : (eqangle(X0,X1,X2,X3,X4,X5,X6,X7) => eqangle(X2,X3,X0,X1,X6,X7,X4,X5))),
  inference(rectify,[],[f19])).
fof(f19,axiom,(
  ! [X0,X1,X2,X3,X8,X9,X10,X11] : (eqangle(X0,X1,X2,X3,X8,X9,X10,X11) => eqangle(X2,X3,X0,X1,X10,X11,X8,X9))),
  file('./7_Book/00EE/08/E059-58.p',ruleD19)).
fof(f13909,plain,(
  ~spl33_180 | ~spl33_186 | ~spl33_564 | spl33_625),
  inference(avatar_contradiction_clause,[],[f13896])).
fof(f13896,plain,(
  $false | (~spl33_180 | ~spl33_186 | ~spl33_564 | ~spl33_625)),
  inference(resolution,[],[f13036,f8881])).
fof(f8881,plain,(
  ~para(sK30,sK22,sK30,sK21) | ~spl33_625),
  inference(avatar_component_clause,[],[f8880])).
fof(f8880,plain,(
  spl33_625 <=> ~para(sK30,sK22,sK30,sK21)),
  introduced(avatar_definition,[new_symbols(naming,[spl33_625])])).
fof(f8080,plain,(
  ~spl33_180 | ~spl33_184),
  inference(avatar_contradiction_clause,[],[f8053])).
fof(f8053,plain,(
  $false | (~spl33_180 | ~spl33_184)),
  inference(resolution,[],[f7944,f2237])).
fof(f2237,plain,(
  ( ! [X48,X49] : (~perp(X48,X49,X48,X49)) ) | ~spl33_184),
  inference(avatar_component_clause,[],[f2236])).
fof(f2236,plain,(
  spl33_184 <=> ! [X49,X48] : ~perp(X48,X49,X48,X49)),
  introduced(avatar_definition,[new_symbols(naming,[spl33_184])])).
fof(f3575,plain,(
  ~spl33_182),
  inference(avatar_contradiction_clause,[],[f3572])).
fof(f3572,plain,(
  $false | ~spl33_182),
  inference(resolution,[],[f2064,f1448])).
fof(f1448,plain,(
  para(sK25,sK23,sK25,sK23)),
  inference(resolution,[],[f791,f510])).
fof(f510,plain,(
  perp(sK25,sK23,sK25,sK20)),
  inference(resolution,[],[f366,f478])).
fof(f478,plain,(
  perp(sK25,sK20,sK25,sK23)),
  inference(cnf_transformation,[],[f358])).
fof(f791,plain,(
  ( ! [X6,X7] : (~perp(X6,X7,sK25,sK20) | para(X6,X7,sK25,sK23)) )),
  inference(resolution,[],[f367,f478])).
fof(f2064,plain,(
  ( ! [X0,X1] : (~para(X0,X1,X0,X1)) ) | ~spl33_182),
  inference(avatar_component_clause,[],[f2063])).
fof(f2063,plain,(
  spl33_182 <=> ! [X1,X0] : ~para(X0,X1,X0,X1)),
  introduced(avatar_definition,[new_symbols(naming,[spl33_182])])).
fof(f2241,plain,(
  spl33_184 | spl33_186),
  inference(avatar_split_clause,[],[f2218,f2239,f2236])).
fof(f2218,plain,(
  ( ! [X47,X45,X48,X46,X44,X49] : (~para(X44,X45,X46,X47) | ~perp(X48,X49,X48,X49) | perp(X44,X45,X46,X47)) )),
  inference(resolution,[],[f858,f433])).
fof(f433,plain,(
  ( ! [X6,X4,X2,X0,X7,X5,X3,X1] : (~eqangle(X0,X1,X2,X3,X4,X5,X6,X7) | ~perp(X4,X5,X6,X7) | perp(X0,X1,X2,X3)) )),
  inference(cnf_transformation,[],[f281])).
fof(f281,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6,X7] : (perp(X0,X1,X2,X3) | ~perp(X4,X5,X6,X7) | ~eqangle(X0,X1,X2,X3,X4,X5,X6,X7))),
  inference(flattening,[],[f280])).
fof(f280,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6,X7] : (perp(X0,X1,X2,X3) | (~perp(X4,X5,X6,X7) | ~eqangle(X0,X1,X2,X3,X4,X5,X6,X7)))),
  inference(ennf_transformation,[],[f147])).
fof(f147,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6,X7] : ((perp(X4,X5,X6,X7) & eqangle(X0,X1,X2,X3,X4,X5,X6,X7)) => perp(X0,X1,X2,X3))),
  inference(rectify,[],[f75])).
fof(f75,axiom,(
  ! [X0,X1,X2,X3,X8,X9,X10,X11] : ((perp(X8,X9,X10,X11) & eqangle(X0,X1,X2,X3,X8,X9,X10,X11)) => perp(X0,X1,X2,X3))),
  file('./7_Book/00EE/08/E059-58.p',ruleD74)).
fof(f2065,plain,(
  spl33_180 | spl33_182),
  inference(avatar_split_clause,[],[f2036,f2063,f2060])).
fof(f2036,plain,(
  ( ! [X2,X0,X3,X1] : (~para(X0,X1,X0,X1) | para(X2,X3,X2,X3)) )),
  inference(resolution,[],[f832,f397])).
fof(f397,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (~eqangle(X0,X1,X4,X5,X2,X3,X4,X5) | para(X0,X1,X2,X3)) )),
  inference(cnf_transformation,[],[f218])).
fof(f218,plain,(
  ! [X0,X1,X2,X3,X4,X5] : (para(X0,X1,X2,X3) | ~eqangle(X0,X1,X4,X5,X2,X3,X4,X5))),
  inference(ennf_transformation,[],[f118])).
fof(f118,plain,(
  ! [X0,X1,X2,X3,X4,X5] : (eqangle(X0,X1,X4,X5,X2,X3,X4,X5) => para(X0,X1,X2,X3))),
  inference(rectify,[],[f39])).
fof(f39,axiom,(
  ! [X0,X1,X2,X3,X8,X9] : (eqangle(X0,X1,X8,X9,X2,X3,X8,X9) => para(X0,X1,X2,X3))),
  file('./7_Book/00EE/08/E059-58.p',ruleD39)).
% SZS output end Proof for E059-58
% ------------------------------
% Version: Vampire 4.4.0 (commit 7916d27 on 2019-08-23 08:50:16 +0100)
% Termination reason: Refutation

% Memory used [KB]: 12537
% Time elapsed: 0.266 s
% ------------------------------
% ------------------------------
