% Refutation found. Thanks to Tanya!
% SZS status Theorem for E061-61
% SZS output start Proof for E061-61
19. ! [X0,X1,X2,X3,X8,X9,X10,X11] : (eqangle(X0,X1,X2,X3,X8,X9,X10,X11) => eqangle(X2,X3,X0,X1,X10,X11,X8,X9)) [input]
23. ! [X0,X1,X2,X3] : (cong(X0,X1,X2,X3) => cong(X0,X1,X3,X2)) [input]
24. ! [X0,X1,X2,X3] : (cong(X0,X1,X2,X3) => cong(X2,X3,X0,X1)) [input]
48. ! [X0,X1,X2,X7,X15] : ((perp(X7,X0,X0,X15) & circle(X7,X0,X1,X2)) => eqangle(X0,X15,X0,X1,X2,X0,X2,X1)) [input]
51. ! [X0,X1,X2,X7,X6] : ((eqangle(X0,X1,X0,X2,X7,X1,X7,X6) & coll(X6,X1,X2) & circle(X7,X0,X1,X2)) => midp(X6,X1,X2)) [input]
68. ! [X0,X1,X2] : (midp(X0,X1,X2) => cong(X0,X1,X0,X2)) [input]
94. ! [X0,X1,X2,X3,X4,X5,X12,X13,X18,X19,X20,X21,X22] : (coll(X18,X2,X13) & coll(X18,X1,X4) & coll(X13,X1,X3) & para(X1,X5,X13,X2) & coll(X12,X1,X5) & coll(X12,X3,X4) & coll(X5,X2,X4) & perp(X4,X3,X3,X0) & circle(X0,X1,X3,X22) & perp(X4,X2,X2,X0) & circle(X0,X1,X2,X21) & perp(X5,X1,X1,X0) & circle(X0,X1,X19,X20)) => cong(X2,X18,X18,X13) [input]
95. ~! [X13,X18,X2] : (! [X0,X1,X2,X3,X4,X5,X12,X13,X18,X19,X20,X21,X22] : (coll(X18,X2,X13) & coll(X18,X1,X4) & coll(X13,X1,X3) & para(X1,X5,X13,X2) & coll(X12,X1,X5) & coll(X12,X3,X4) & coll(X5,X2,X4) & perp(X4,X3,X3,X0) & circle(X0,X1,X3,X22) & perp(X4,X2,X2,X0) & circle(X0,X1,X2,X21) & perp(X5,X1,X1,X0) & circle(X0,X1,X19,X20)) => cong(X2,X18,X18,X13)) [negated conjecture 94]
100. ! [X0,X1,X2,X3,X4,X5,X6,X7] : (eqangle(X0,X1,X2,X3,X4,X5,X6,X7) => eqangle(X2,X3,X0,X1,X6,X7,X4,X5)) [rectify 19]
126. ! [X0,X1,X2,X3,X4] : ((perp(X3,X0,X0,X4) & circle(X3,X0,X1,X2)) => eqangle(X0,X4,X0,X1,X2,X0,X2,X1)) [rectify 48]
129. ! [X0,X1,X2,X3,X4] : ((eqangle(X0,X1,X0,X2,X3,X1,X3,X4) & coll(X4,X1,X2) & circle(X3,X0,X1,X2)) => midp(X4,X1,X2)) [rectify 51]
165. ~! [X0,X1,X2] : (! [X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15] : (coll(X11,X5,X10) & coll(X11,X4,X7) & coll(X10,X4,X6) & para(X4,X8,X10,X5) & coll(X9,X4,X8) & coll(X9,X6,X7) & coll(X8,X5,X7) & perp(X7,X6,X6,X3) & circle(X3,X4,X6,X15) & perp(X7,X5,X5,X3) & circle(X3,X4,X5,X14) & perp(X8,X4,X4,X3) & circle(X3,X4,X12,X13)) => cong(X2,X1,X1,X0)) [rectify 95]
191. ! [X0,X1,X2,X3,X4,X5,X6,X7] : (eqangle(X2,X3,X0,X1,X6,X7,X4,X5) | ~eqangle(X0,X1,X2,X3,X4,X5,X6,X7)) [ennf transformation 100]
196. ! [X0,X1,X2,X3] : (cong(X0,X1,X3,X2) | ~cong(X0,X1,X2,X3)) [ennf transformation 23]
197. ! [X0,X1,X2,X3] : (cong(X2,X3,X0,X1) | ~cong(X0,X1,X2,X3)) [ennf transformation 24]
230. ! [X0,X1,X2,X3,X4] : (eqangle(X0,X4,X0,X1,X2,X0,X2,X1) | (~perp(X3,X0,X0,X4) | ~circle(X3,X0,X1,X2))) [ennf transformation 126]
231. ! [X0,X1,X2,X3,X4] : (eqangle(X0,X4,X0,X1,X2,X0,X2,X1) | ~perp(X3,X0,X0,X4) | ~circle(X3,X0,X1,X2)) [flattening 230]
236. ! [X0,X1,X2,X3,X4] : (midp(X4,X1,X2) | (~eqangle(X0,X1,X0,X2,X3,X1,X3,X4) | ~coll(X4,X1,X2) | ~circle(X3,X0,X1,X2))) [ennf transformation 129]
237. ! [X0,X1,X2,X3,X4] : (midp(X4,X1,X2) | ~eqangle(X0,X1,X0,X2,X3,X1,X3,X4) | ~coll(X4,X1,X2) | ~circle(X3,X0,X1,X2)) [flattening 236]
266. ! [X0,X1,X2] : (cong(X0,X1,X0,X2) | ~midp(X0,X1,X2)) [ennf transformation 68]
315. ? [X0,X1,X2] : (~cong(X2,X1,X1,X0) & ! [X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15] : (coll(X11,X5,X10) & coll(X11,X4,X7) & coll(X10,X4,X6) & para(X4,X8,X10,X5) & coll(X9,X4,X8) & coll(X9,X6,X7) & coll(X8,X5,X7) & perp(X7,X6,X6,X3) & circle(X3,X4,X6,X15) & perp(X7,X5,X5,X3) & circle(X3,X4,X5,X14) & perp(X8,X4,X4,X3) & circle(X3,X4,X12,X13))) [ennf transformation 165]
352. ? [X0,X1,X2] : (~cong(X2,X1,X1,X0) & ! [X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15] : (coll(X11,X5,X10) & coll(X11,X4,X7) & coll(X10,X4,X6) & para(X4,X8,X10,X5) & coll(X9,X4,X8) & coll(X9,X6,X7) & coll(X8,X5,X7) & perp(X7,X6,X6,X3) & circle(X3,X4,X6,X15) & perp(X7,X5,X5,X3) & circle(X3,X4,X5,X14) & perp(X8,X4,X4,X3) & circle(X3,X4,X12,X13))) => (~cong(sK22,sK21,sK21,sK20) & ! [X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15] : (coll(X11,X5,X10) & coll(X11,X4,X7) & coll(X10,X4,X6) & para(X4,X8,X10,X5) & coll(X9,X4,X8) & coll(X9,X6,X7) & coll(X8,X5,X7) & perp(X7,X6,X6,X3) & circle(X3,X4,X6,X15) & perp(X7,X5,X5,X3) & circle(X3,X4,X5,X14) & perp(X8,X4,X4,X3) & circle(X3,X4,X12,X13))) [choice axiom]
353. ~cong(sK22,sK21,sK21,sK20) & ! [X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15] : (coll(X11,X5,X10) & coll(X11,X4,X7) & coll(X10,X4,X6) & para(X4,X8,X10,X5) & coll(X9,X4,X8) & coll(X9,X6,X7) & coll(X8,X5,X7) & perp(X7,X6,X6,X3) & circle(X3,X4,X6,X15) & perp(X7,X5,X5,X3) & circle(X3,X4,X5,X14) & perp(X8,X4,X4,X3) & circle(X3,X4,X12,X13)) [skolemisation 315,352]
372. ~eqangle(X0,X1,X2,X3,X4,X5,X6,X7) | eqangle(X2,X3,X0,X1,X6,X7,X4,X5) [cnf transformation 191]
376. ~cong(X0,X1,X2,X3) | cong(X0,X1,X3,X2) [cnf transformation 196]
377. ~cong(X0,X1,X2,X3) | cong(X2,X3,X0,X1) [cnf transformation 197]
402. ~circle(X3,X0,X1,X2) | ~perp(X3,X0,X0,X4) | eqangle(X0,X4,X0,X1,X2,X0,X2,X1) [cnf transformation 231]
405. ~eqangle(X0,X1,X0,X2,X3,X1,X3,X4) | midp(X4,X1,X2) | ~coll(X4,X1,X2) | ~circle(X3,X0,X1,X2) [cnf transformation 237]
422. cong(X0,X1,X0,X2) | ~midp(X0,X1,X2) [cnf transformation 266]
471. circle(X3,X4,X6,X15) [cnf transformation 353]
472. perp(X7,X6,X6,X3) [cnf transformation 353]
479. coll(X11,X5,X10) [cnf transformation 353]
480. ~cong(sK22,sK21,sK21,sK20) [cnf transformation 353]
496. cong(X0,X1,X0,X2) | ~midp(X0,X2,X1) [resolution 377,422]
499. cong(X3,X5,X4,X3) | ~midp(X3,X4,X5) [resolution 496,376]
504. cong(X1,X0,X0,X2) | ~midp(X0,X1,X2) [resolution 499,377]
515. ~midp(sK21,sK22,sK20) [resolution 504,480]
863. ~perp(X0,X1,X1,X2) | eqangle(X1,X2,X1,X3,X4,X1,X4,X3) [resolution 402,471]
864. eqangle(X1,X2,X1,X3,X4,X1,X4,X3) [subsumption resolution 863,472]
868. eqangle(X10,X11,X10,X12,X13,X11,X13,X10) [resolution 864,372]
1226. midp(X58,X59,X60) | ~coll(X58,X59,X60) | ~circle(X61,X58,X59,X60) [resolution 405,868]
1321. midp(X58,X59,X60) | ~circle(X61,X58,X59,X60) [subsumption resolution 1226,479]
1322. midp(X58,X59,X60) [subsumption resolution 1321,471]
1397. $false [resolution 1322,515]
% SZS output end Proof for E061-61
% ------------------------------
% Version: Vampire 4.4.0 (commit 7916d27 on 2019-08-23 08:50:16 +0100)
% Termination reason: Refutation

% Memory used [KB]: 5373
% Time elapsed: 0.038 s
% ------------------------------
% ------------------------------
