% Refutation found. Thanks to Tanya!
% SZS status Theorem for E059-52
% SZS output start Proof for E059-52
fof(f1633,plain,(
  $false),
  inference(resolution,[],[f1618,f483])).
fof(f483,plain,(
  ~cong(sK23,sK22,sK21,sK20)),
  inference(cnf_transformation,[],[f357])).
fof(f357,plain,(
  ~cong(sK23,sK22,sK21,sK20) & ! [X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15] : (coll(X11,X5,X7) & coll(X11,X6,X8) & coll(X10,X7,X9) & coll(X10,X5,X8) & coll(X9,X6,X8) & perp(X9,X7,X7,X4) & perp(X4,X5,X5,X8) & perp(X4,X6,X6,X8) & circle(X4,X5,X7,X15) & circle(X4,X5,X6,X14) & perp(X4,X5,X4,X6) & circle(X4,X5,X12,X13))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK20,sK21,sK22,sK23])],[f319,f356])).
fof(f356,plain,(
  ? [X0,X1,X2,X3] : (~cong(X3,X2,X1,X0) & ! [X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15] : (coll(X11,X5,X7) & coll(X11,X6,X8) & coll(X10,X7,X9) & coll(X10,X5,X8) & coll(X9,X6,X8) & perp(X9,X7,X7,X4) & perp(X4,X5,X5,X8) & perp(X4,X6,X6,X8) & circle(X4,X5,X7,X15) & circle(X4,X5,X6,X14) & perp(X4,X5,X4,X6) & circle(X4,X5,X12,X13))) => (~cong(sK23,sK22,sK21,sK20) & ! [X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15] : (coll(X11,X5,X7) & coll(X11,X6,X8) & coll(X10,X7,X9) & coll(X10,X5,X8) & coll(X9,X6,X8) & perp(X9,X7,X7,X4) & perp(X4,X5,X5,X8) & perp(X4,X6,X6,X8) & circle(X4,X5,X7,X15) & circle(X4,X5,X6,X14) & perp(X4,X5,X4,X6) & circle(X4,X5,X12,X13)))),
  introduced(choice_axiom,[])).
fof(f319,plain,(
  ? [X0,X1,X2,X3] : (~cong(X3,X2,X1,X0) & ! [X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15] : (coll(X11,X5,X7) & coll(X11,X6,X8) & coll(X10,X7,X9) & coll(X10,X5,X8) & coll(X9,X6,X8) & perp(X9,X7,X7,X4) & perp(X4,X5,X5,X8) & perp(X4,X6,X6,X8) & circle(X4,X5,X7,X15) & circle(X4,X5,X6,X14) & perp(X4,X5,X4,X6) & circle(X4,X5,X12,X13)))),
  inference(ennf_transformation,[],[f167])).
fof(f167,plain,(
  ~! [X0,X1,X2,X3] : (! [X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15] : (coll(X11,X5,X7) & coll(X11,X6,X8) & coll(X10,X7,X9) & coll(X10,X5,X8) & coll(X9,X6,X8) & perp(X9,X7,X7,X4) & perp(X4,X5,X5,X8) & perp(X4,X6,X6,X8) & circle(X4,X5,X7,X15) & circle(X4,X5,X6,X14) & perp(X4,X5,X4,X6) & circle(X4,X5,X12,X13)) => cong(X3,X2,X1,X0))),
  inference(rectify,[],[f96])).
fof(f96,negated_conjecture,(
  ~! [X12,X3,X4,X13] : (! [X0,X1,X2,X3,X4,X5,X12,X13,X18,X19,X20,X21] : (coll(X13,X1,X3) & coll(X13,X2,X4) & coll(X12,X3,X5) & coll(X12,X1,X4) & coll(X5,X2,X4) & perp(X5,X3,X3,X0) & perp(X0,X1,X1,X4) & perp(X0,X2,X2,X4) & circle(X0,X1,X3,X21) & circle(X0,X1,X2,X20) & perp(X0,X1,X0,X2) & circle(X0,X1,X18,X19)) => cong(X13,X4,X3,X12))),
  inference(negated_conjecture,[],[f95])).
fof(f95,conjecture,(
  ( ! [X4,X12,X3,X13] : (! [X0,X1,X2,X3,X4,X5,X12,X13,X18,X19,X20,X21] : (coll(X13,X1,X3) & coll(X13,X2,X4) & coll(X12,X3,X5) & coll(X12,X1,X4) & coll(X5,X2,X4) & perp(X5,X3,X3,X0) & perp(X0,X1,X1,X4) & perp(X0,X2,X2,X4) & circle(X0,X1,X3,X21) & circle(X0,X1,X2,X20) & perp(X0,X1,X0,X2) & circle(X0,X1,X18,X19)) => cong(X13,X4,X3,X12)) )),
  file('./7_Book/00EE/07/E059-52.p',exemplo7Book00EE07E05952)).
fof(f1618,plain,(
  ( ! [X12,X10,X13,X11] : (cong(X10,X12,X11,X13)) )),
  inference(subsumption_resolution,[],[f1613,f1242])).
fof(f1242,plain,(
  ( ! [X43,X41,X42] : (cong(X41,X42,X41,X43)) )),
  inference(resolution,[],[f1214,f539])).
fof(f539,plain,(
  ( ! [X10,X11,X9] : (~midp(X11,X10,X11) | cong(X9,X10,X9,X11)) )),
  inference(resolution,[],[f413,f487])).
fof(f487,plain,(
  ( ! [X4,X5,X3] : (perp(X3,X4,X5,X4)) )),
  inference(resolution,[],[f364,f477])).
fof(f477,plain,(
  ( ! [X4,X7,X9] : (perp(X9,X7,X7,X4)) )),
  inference(cnf_transformation,[],[f357])).
fof(f364,plain,(
  ( ! [X2,X0,X3,X1] : (~perp(X0,X1,X2,X3) | perp(X0,X1,X3,X2)) )),
  inference(cnf_transformation,[],[f176])).
fof(f176,plain,(
  ! [X0,X1,X2,X3] : (perp(X0,X1,X3,X2) | ~perp(X0,X1,X2,X3))),
  inference(ennf_transformation,[],[f7])).
fof(f7,axiom,(
  ! [X0,X1,X2,X3] : (perp(X0,X1,X2,X3) => perp(X0,X1,X3,X2))),
  file('./7_Book/00EE/07/E059-52.p',ruleD7)).
fof(f413,plain,(
  ( ! [X2,X0,X3,X1] : (~perp(X3,X2,X0,X1) | cong(X3,X0,X3,X1) | ~midp(X2,X0,X1)) )),
  inference(cnf_transformation,[],[f249])).
fof(f249,plain,(
  ! [X0,X1,X2,X3] : (cong(X3,X0,X3,X1) | ~perp(X3,X2,X0,X1) | ~midp(X2,X0,X1))),
  inference(flattening,[],[f248])).
fof(f248,plain,(
  ! [X0,X1,X2,X3] : (cong(X3,X0,X3,X1) | (~perp(X3,X2,X0,X1) | ~midp(X2,X0,X1)))),
  inference(ennf_transformation,[],[f134])).
fof(f134,plain,(
  ! [X0,X1,X2,X3] : ((perp(X3,X2,X0,X1) & midp(X2,X0,X1)) => cong(X3,X0,X3,X1))),
  inference(rectify,[],[f56])).
fof(f56,axiom,(
  ! [X0,X1,X6,X7] : ((perp(X7,X6,X0,X1) & midp(X6,X0,X1)) => cong(X7,X0,X7,X1))),
  file('./7_Book/00EE/07/E059-52.p',ruleD55)).
fof(f1214,plain,(
  ( ! [X0,X1] : (midp(X0,X1,X0)) )),
  inference(subsumption_resolution,[],[f1213,f474])).
fof(f474,plain,(
  ( ! [X4,X7,X5,X15] : (circle(X4,X5,X7,X15)) )),
  inference(cnf_transformation,[],[f357])).
fof(f1213,plain,(
  ( ! [X2,X0,X1] : (midp(X0,X1,X0) | ~circle(X2,X2,X1,X0)) )),
  inference(subsumption_resolution,[],[f1186,f482])).
fof(f482,plain,(
  ( ! [X7,X5,X11] : (coll(X11,X5,X7)) )),
  inference(cnf_transformation,[],[f357])).
fof(f1186,plain,(
  ( ! [X2,X0,X1] : (midp(X0,X1,X0) | ~coll(X0,X1,X0) | ~circle(X2,X2,X1,X0)) )),
  inference(resolution,[],[f409,f645])).
fof(f645,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (eqangle(X0,X1,X2,X3,X4,X5,X2,X3)) )),
  inference(subsumption_resolution,[],[f641,f629])).
fof(f629,plain,(
  ( ! [X6,X8,X7,X9] : (para(X6,X7,X8,X9)) )),
  inference(resolution,[],[f544,f486])).
fof(f486,plain,(
  ( ! [X2,X0,X1] : (perp(X0,X1,X2,X0)) )),
  inference(resolution,[],[f364,f472])).
fof(f472,plain,(
  ( ! [X6,X4,X5] : (perp(X4,X5,X4,X6)) )),
  inference(cnf_transformation,[],[f357])).
fof(f544,plain,(
  ( ! [X4,X2,X0,X3,X1] : (~perp(X0,X1,X2,X4) | para(X0,X1,X2,X3)) )),
  inference(resolution,[],[f366,f472])).
fof(f366,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (~perp(X2,X3,X4,X5) | para(X0,X1,X4,X5) | ~perp(X0,X1,X2,X3)) )),
  inference(cnf_transformation,[],[f179])).
fof(f179,plain,(
  ! [X0,X1,X2,X3,X4,X5] : (para(X0,X1,X4,X5) | ~perp(X2,X3,X4,X5) | ~perp(X0,X1,X2,X3))),
  inference(flattening,[],[f178])).
fof(f178,plain,(
  ! [X0,X1,X2,X3,X4,X5] : (para(X0,X1,X4,X5) | (~perp(X2,X3,X4,X5) | ~perp(X0,X1,X2,X3)))),
  inference(ennf_transformation,[],[f9])).
fof(f9,axiom,(
  ! [X0,X1,X2,X3,X4,X5] : ((perp(X2,X3,X4,X5) & perp(X0,X1,X2,X3)) => para(X0,X1,X4,X5))),
  file('./7_Book/00EE/07/E059-52.p',ruleD9)).
fof(f641,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (eqangle(X0,X1,X2,X3,X4,X5,X2,X3) | ~para(X1,X0,X4,X5)) )),
  inference(resolution,[],[f375,f397])).
fof(f397,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (eqangle(X0,X1,X4,X5,X2,X3,X4,X5) | ~para(X0,X1,X2,X3)) )),
  inference(cnf_transformation,[],[f219])).
fof(f219,plain,(
  ! [X0,X1,X2,X3,X4,X5] : (eqangle(X0,X1,X4,X5,X2,X3,X4,X5) | ~para(X0,X1,X2,X3))),
  inference(ennf_transformation,[],[f119])).
fof(f119,plain,(
  ! [X0,X1,X2,X3,X4,X5] : (para(X0,X1,X2,X3) => eqangle(X0,X1,X4,X5,X2,X3,X4,X5))),
  inference(rectify,[],[f40])).
fof(f40,axiom,(
  ! [X0,X1,X2,X3,X8,X9] : (para(X0,X1,X2,X3) => eqangle(X0,X1,X8,X9,X2,X3,X8,X9))),
  file('./7_Book/00EE/07/E059-52.p',ruleD40)).
fof(f375,plain,(
  ( ! [X6,X4,X2,X0,X7,X5,X3,X1] : (~eqangle(X0,X1,X2,X3,X4,X5,X6,X7) | eqangle(X1,X0,X2,X3,X4,X5,X6,X7)) )),
  inference(cnf_transformation,[],[f192])).
fof(f192,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6,X7] : (eqangle(X1,X0,X2,X3,X4,X5,X6,X7) | ~eqangle(X0,X1,X2,X3,X4,X5,X6,X7))),
  inference(ennf_transformation,[],[f100])).
fof(f100,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6,X7] : (eqangle(X0,X1,X2,X3,X4,X5,X6,X7) => eqangle(X1,X0,X2,X3,X4,X5,X6,X7))),
  inference(rectify,[],[f18])).
fof(f18,axiom,(
  ! [X0,X1,X2,X3,X8,X9,X10,X11] : (eqangle(X0,X1,X2,X3,X8,X9,X10,X11) => eqangle(X1,X0,X2,X3,X8,X9,X10,X11))),
  file('./7_Book/00EE/07/E059-52.p',ruleD18)).
fof(f409,plain,(
  ( ! [X4,X2,X0,X3,X1] : (~eqangle(X0,X1,X0,X2,X3,X1,X3,X4) | midp(X4,X1,X2) | ~coll(X4,X1,X2) | ~circle(X3,X0,X1,X2)) )),
  inference(cnf_transformation,[],[f241])).
fof(f241,plain,(
  ! [X0,X1,X2,X3,X4] : (midp(X4,X1,X2) | ~eqangle(X0,X1,X0,X2,X3,X1,X3,X4) | ~coll(X4,X1,X2) | ~circle(X3,X0,X1,X2))),
  inference(flattening,[],[f240])).
fof(f240,plain,(
  ! [X0,X1,X2,X3,X4] : (midp(X4,X1,X2) | (~eqangle(X0,X1,X0,X2,X3,X1,X3,X4) | ~coll(X4,X1,X2) | ~circle(X3,X0,X1,X2)))),
  inference(ennf_transformation,[],[f131])).
fof(f131,plain,(
  ! [X0,X1,X2,X3,X4] : ((eqangle(X0,X1,X0,X2,X3,X1,X3,X4) & coll(X4,X1,X2) & circle(X3,X0,X1,X2)) => midp(X4,X1,X2))),
  inference(rectify,[],[f52])).
fof(f52,axiom,(
  ! [X0,X1,X2,X7,X6] : ((eqangle(X0,X1,X0,X2,X7,X1,X7,X6) & coll(X6,X1,X2) & circle(X7,X0,X1,X2)) => midp(X6,X1,X2))),
  file('./7_Book/00EE/07/E059-52.p',ruleD51)).
fof(f1613,plain,(
  ( ! [X12,X10,X13,X11,X9] : (~cong(X9,X10,X9,X11) | cong(X10,X12,X11,X13)) )),
  inference(resolution,[],[f1539,f433])).
fof(f433,plain,(
  ( ! [X6,X4,X2,X0,X7,X5,X3,X1] : (~eqratio(X0,X1,X2,X3,X4,X5,X6,X7) | ~cong(X4,X5,X6,X7) | cong(X0,X1,X2,X3)) )),
  inference(cnf_transformation,[],[f283])).
fof(f283,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6,X7] : (cong(X0,X1,X2,X3) | ~cong(X4,X5,X6,X7) | ~eqratio(X0,X1,X2,X3,X4,X5,X6,X7))),
  inference(flattening,[],[f282])).
fof(f282,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6,X7] : (cong(X0,X1,X2,X3) | (~cong(X4,X5,X6,X7) | ~eqratio(X0,X1,X2,X3,X4,X5,X6,X7)))),
  inference(ennf_transformation,[],[f148])).
fof(f148,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6,X7] : ((cong(X4,X5,X6,X7) & eqratio(X0,X1,X2,X3,X4,X5,X6,X7)) => cong(X0,X1,X2,X3))),
  inference(rectify,[],[f76])).
fof(f76,axiom,(
  ! [X0,X1,X2,X3,X8,X9,X10,X11] : ((cong(X8,X9,X10,X11) & eqratio(X0,X1,X2,X3,X8,X9,X10,X11)) => cong(X0,X1,X2,X3))),
  file('./7_Book/00EE/07/E059-52.p',rule75)).
fof(f1539,plain,(
  ( ! [X14,X12,X10,X13,X11] : (eqratio(X10,X11,X12,X13,X14,X10,X14,X12)) )),
  inference(resolution,[],[f1252,f385])).
fof(f385,plain,(
  ( ! [X6,X4,X2,X0,X7,X5,X3,X1] : (~eqratio(X0,X1,X2,X3,X4,X5,X6,X7) | eqratio(X4,X5,X6,X7,X0,X1,X2,X3)) )),
  inference(cnf_transformation,[],[f204])).
fof(f204,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6,X7] : (eqratio(X4,X5,X6,X7,X0,X1,X2,X3) | ~eqratio(X0,X1,X2,X3,X4,X5,X6,X7))),
  inference(ennf_transformation,[],[f107])).
fof(f107,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6,X7] : (eqratio(X0,X1,X2,X3,X4,X5,X6,X7) => eqratio(X4,X5,X6,X7,X0,X1,X2,X3))),
  inference(rectify,[],[f28])).
fof(f28,axiom,(
  ! [X0,X1,X2,X3,X8,X9,X10,X11] : (eqratio(X0,X1,X2,X3,X8,X9,X10,X11) => eqratio(X8,X9,X10,X11,X0,X1,X2,X3))),
  file('./7_Book/00EE/07/E059-52.p',ruleD28)).
fof(f1252,plain,(
  ( ! [X6,X8,X7,X5,X9] : (eqratio(X5,X8,X5,X6,X8,X9,X6,X7)) )),
  inference(subsumption_resolution,[],[f1251,f629])).
fof(f1251,plain,(
  ( ! [X6,X8,X7,X5,X9] : (~para(X8,X6,X9,X7) | eqratio(X5,X8,X5,X6,X8,X9,X6,X7)) )),
  inference(subsumption_resolution,[],[f1250,f482])).
fof(f1250,plain,(
  ( ! [X6,X8,X7,X5,X9] : (~coll(X5,X8,X9) | ~para(X8,X6,X9,X7) | eqratio(X5,X8,X5,X6,X8,X9,X6,X7)) )),
  inference(subsumption_resolution,[],[f1246,f482])).
fof(f1246,plain,(
  ( ! [X6,X8,X7,X5,X9] : (~coll(X5,X6,X7) | ~coll(X5,X8,X9) | ~para(X8,X6,X9,X7) | eqratio(X5,X8,X5,X6,X8,X9,X6,X7)) )),
  inference(resolution,[],[f423,f386])).
fof(f386,plain,(
  ( ! [X6,X4,X2,X0,X7,X5,X3,X1] : (~eqratio(X0,X1,X2,X3,X4,X5,X6,X7) | eqratio(X0,X1,X4,X5,X2,X3,X6,X7)) )),
  inference(cnf_transformation,[],[f205])).
fof(f205,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6,X7] : (eqratio(X0,X1,X4,X5,X2,X3,X6,X7) | ~eqratio(X0,X1,X2,X3,X4,X5,X6,X7))),
  inference(ennf_transformation,[],[f108])).
fof(f108,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6,X7] : (eqratio(X0,X1,X2,X3,X4,X5,X6,X7) => eqratio(X0,X1,X4,X5,X2,X3,X6,X7))),
  inference(rectify,[],[f29])).
fof(f29,axiom,(
  ! [X0,X1,X2,X3,X8,X9,X10,X11] : (eqratio(X0,X1,X2,X3,X8,X9,X10,X11) => eqratio(X0,X1,X8,X9,X2,X3,X10,X11))),
  file('./7_Book/00EE/07/E059-52.p',ruleD29)).
fof(f423,plain,(
  ( ! [X4,X2,X0,X3,X1] : (eqratio(X4,X0,X0,X2,X4,X1,X1,X3) | ~coll(X4,X1,X3) | ~coll(X4,X0,X2) | ~para(X0,X1,X2,X3)) )),
  inference(cnf_transformation,[],[f266])).
fof(f266,plain,(
  ! [X0,X1,X2,X3,X4] : (eqratio(X4,X0,X0,X2,X4,X1,X1,X3) | ~coll(X4,X1,X3) | ~coll(X4,X0,X2) | ~para(X0,X1,X2,X3))),
  inference(flattening,[],[f265])).
fof(f265,plain,(
  ! [X0,X1,X2,X3,X4] : (eqratio(X4,X0,X0,X2,X4,X1,X1,X3) | (~coll(X4,X1,X3) | ~coll(X4,X0,X2) | ~para(X0,X1,X2,X3)))),
  inference(ennf_transformation,[],[f144])).
fof(f144,plain,(
  ! [X0,X1,X2,X3,X4] : ((coll(X4,X1,X3) & coll(X4,X0,X2) & para(X0,X1,X2,X3)) => eqratio(X4,X0,X0,X2,X4,X1,X1,X3))),
  inference(rectify,[],[f66])).
fof(f66,axiom,(
  ! [X0,X1,X2,X3,X7] : ((coll(X7,X1,X3) & coll(X7,X0,X2) & para(X0,X1,X2,X3)) => eqratio(X7,X0,X0,X2,X7,X1,X1,X3))),
  file('./7_Book/00EE/07/E059-52.p',ruleD65)).
% SZS output end Proof for E059-52
% ------------------------------
% Version: Vampire 4.4.0 (commit 7916d27 on 2019-08-23 08:50:16 +0100)
% Termination reason: Refutation

% Memory used [KB]: 5500
% Time elapsed: 0.040 s
% ------------------------------
% ------------------------------
