% Refutation found. Thanks to Tanya!
% SZS status Theorem for 9point
% SZS output start Proof for 9point
23. ! [X0,X1,X2,X3] : (cong(X0,X1,X2,X3) => cong(X0,X1,X3,X2)) [input]
24. ! [X0,X1,X2,X3] : (cong(X0,X1,X2,X3) => cong(X2,X3,X0,X1)) [input]
52. ! [X0,X1,X2,X6] : ((midp(X6,X0,X2) & perp(X0,X1,X1,X2)) => cong(X0,X6,X1,X6)) [input]
67. ! [X0,X1,X2] : ((coll(X0,X1,X2) & cong(X0,X1,X0,X2)) => midp(X0,X1,X2)) [input]
68. ! [X0,X1,X2] : (midp(X0,X1,X2) => cong(X0,X1,X0,X2)) [input]
94. ! [X0,X1,X2,X3,X4,X5,X16,X6,X18] : (circle(X16,X3,X6,X18) & coll(X0,X1,X6) & perp(X5,X4,X21,X16) & midp(X21,X5,X4) & perp(X3,X4,X20,X16) & midp(X20,X3,X4) & perp(X3,X5,X19,X16) & midp(X19,X3,X5) & coll(X5,X0,X1) & perp(X5,X2,X0,X1) & coll(X4,X0,X2) & perp(X4,X1,X0,X2) & coll(X3,X1,X2) & perp(X3,X0,X1,X2)) => midp(X6,X0,X1) [input]
95. ~! [X1,X0,X6,X21,X20,X19] : (! [X0,X1,X2,X3,X4,X5,X16,X6,X18] : (circle(X16,X3,X6,X18) & coll(X0,X1,X6) & perp(X5,X4,X21,X16) & midp(X21,X5,X4) & perp(X3,X4,X20,X16) & midp(X20,X3,X4) & perp(X3,X5,X19,X16) & midp(X19,X3,X5) & coll(X5,X0,X1) & perp(X5,X2,X0,X1) & coll(X4,X0,X2) & perp(X4,X1,X0,X2) & coll(X3,X1,X2) & perp(X3,X0,X1,X2)) => midp(X6,X0,X1)) [negated conjecture 94]
130. ! [X0,X1,X2,X3] : ((midp(X3,X0,X2) & perp(X0,X1,X1,X2)) => cong(X0,X3,X1,X3)) [rectify 52]
165. ~! [X0,X1,X2,X3,X4,X5] : (! [X6,X7,X8,X9,X10,X11,X12,X13,X14] : (circle(X12,X9,X13,X14) & coll(X6,X7,X13) & perp(X11,X10,X3,X12) & midp(X3,X11,X10) & perp(X9,X10,X4,X12) & midp(X4,X9,X10) & perp(X9,X11,X5,X12) & midp(X5,X9,X11) & coll(X11,X6,X7) & perp(X11,X8,X6,X7) & coll(X10,X6,X8) & perp(X10,X7,X6,X8) & coll(X9,X7,X8) & perp(X9,X6,X7,X8)) => midp(X2,X1,X0)) [rectify 95]
196. ! [X0,X1,X2,X3] : (cong(X0,X1,X3,X2) | ~cong(X0,X1,X2,X3)) [ennf transformation 23]
197. ! [X0,X1,X2,X3] : (cong(X2,X3,X0,X1) | ~cong(X0,X1,X2,X3)) [ennf transformation 24]
238. ! [X0,X1,X2,X3] : (cong(X0,X3,X1,X3) | (~midp(X3,X0,X2) | ~perp(X0,X1,X1,X2))) [ennf transformation 130]
239. ! [X0,X1,X2,X3] : (cong(X0,X3,X1,X3) | ~midp(X3,X0,X2) | ~perp(X0,X1,X1,X2)) [flattening 238]
264. ! [X0,X1,X2] : (midp(X0,X1,X2) | (~coll(X0,X1,X2) | ~cong(X0,X1,X0,X2))) [ennf transformation 67]
265. ! [X0,X1,X2] : (midp(X0,X1,X2) | ~coll(X0,X1,X2) | ~cong(X0,X1,X0,X2)) [flattening 264]
266. ! [X0,X1,X2] : (cong(X0,X1,X0,X2) | ~midp(X0,X1,X2)) [ennf transformation 68]
315. ? [X0,X1,X2,X3,X4,X5] : (~midp(X2,X1,X0) & ! [X6,X7,X8,X9,X10,X11,X12,X13,X14] : (circle(X12,X9,X13,X14) & coll(X6,X7,X13) & perp(X11,X10,X3,X12) & midp(X3,X11,X10) & perp(X9,X10,X4,X12) & midp(X4,X9,X10) & perp(X9,X11,X5,X12) & midp(X5,X9,X11) & coll(X11,X6,X7) & perp(X11,X8,X6,X7) & coll(X10,X6,X8) & perp(X10,X7,X6,X8) & coll(X9,X7,X8) & perp(X9,X6,X7,X8))) [ennf transformation 165]
352. ? [X0,X1,X2,X3,X4,X5] : (~midp(X2,X1,X0) & ! [X6,X7,X8,X9,X10,X11,X12,X13,X14] : (circle(X12,X9,X13,X14) & coll(X6,X7,X13) & perp(X11,X10,X3,X12) & midp(X3,X11,X10) & perp(X9,X10,X4,X12) & midp(X4,X9,X10) & perp(X9,X11,X5,X12) & midp(X5,X9,X11) & coll(X11,X6,X7) & perp(X11,X8,X6,X7) & coll(X10,X6,X8) & perp(X10,X7,X6,X8) & coll(X9,X7,X8) & perp(X9,X6,X7,X8))) => (~midp(sK22,sK21,sK20) & ! [X14,X13,X12,X11,X10,X9,X8,X7,X6] : (circle(X12,X9,X13,X14) & coll(X6,X7,X13) & perp(X11,X10,sK23,X12) & midp(sK23,X11,X10) & perp(X9,X10,sK24,X12) & midp(sK24,X9,X10) & perp(X9,X11,sK25,X12) & midp(sK25,X9,X11) & coll(X11,X6,X7) & perp(X11,X8,X6,X7) & coll(X10,X6,X8) & perp(X10,X7,X6,X8) & coll(X9,X7,X8) & perp(X9,X6,X7,X8))) [choice axiom]
353. ~midp(sK22,sK21,sK20) & ! [X6,X7,X8,X9,X10,X11,X12,X13,X14] : (circle(X12,X9,X13,X14) & coll(X6,X7,X13) & perp(X11,X10,sK23,X12) & midp(sK23,X11,X10) & perp(X9,X10,sK24,X12) & midp(sK24,X9,X10) & perp(X9,X11,sK25,X12) & midp(sK25,X9,X11) & coll(X11,X6,X7) & perp(X11,X8,X6,X7) & coll(X10,X6,X8) & perp(X10,X7,X6,X8) & coll(X9,X7,X8) & perp(X9,X6,X7,X8)) [skolemisation 315,352]
376. ~cong(X0,X1,X2,X3) | cong(X0,X1,X3,X2) [cnf transformation 196]
377. ~cong(X0,X1,X2,X3) | cong(X2,X3,X0,X1) [cnf transformation 197]
406. ~perp(X0,X1,X1,X2) | ~midp(X3,X0,X2) | cong(X0,X3,X1,X3) [cnf transformation 239]
421. ~cong(X0,X1,X0,X2) | ~coll(X0,X1,X2) | midp(X0,X1,X2) [cnf transformation 265]
422. cong(X0,X1,X0,X2) | ~midp(X0,X1,X2) [cnf transformation 266]
471. perp(X11,X8,X6,X7) [cnf transformation 353]
477. midp(sK23,X11,X10) [cnf transformation 353]
479. coll(X6,X7,X13) [cnf transformation 353]
481. ~midp(sK22,sK21,sK20) [cnf transformation 353]
498. cong(X0,X1,X2,X0) | ~midp(X0,X1,X2) [resolution 376,422]
501. cong(X3,X4,X4,X5) | ~midp(X4,X5,X3) [resolution 377,498]
506. cong(X5,X3,X4,X3) | ~midp(X3,X4,X5) [resolution 501,376]
544. ~midp(X6,X7,X7) | ~coll(X7,X6,X6) | midp(X7,X6,X6) [resolution 506,421]
545. ~midp(X6,X7,X7) | midp(X7,X6,X6) [subsumption resolution 544,479]
548. midp(X2,sK23,sK23) [resolution 545,477]
579. ~midp(X0,X1,X2) | cong(X1,X0,X3,X0) [resolution 406,471]
586. cong(sK23,X6,X7,X6) [resolution 579,548]
713. cong(sK23,X6,X6,X7) [resolution 586,376]
749. cong(X4,X5,sK23,X4) [resolution 713,377]
830. cong(X6,X7,X6,sK23) [resolution 749,376]
998. ~coll(X16,X17,sK23) | midp(X16,X17,sK23) [resolution 830,421]
999. midp(X16,X17,sK23) [subsumption resolution 998,479]
1009. cong(X18,X19,X20,X19) [resolution 999,579]
1431. cong(X8,X9,X9,X10) [resolution 1009,376]
1536. cong(X5,X6,X7,X5) [resolution 1431,377]
1574. cong(X8,X9,X8,X10) [resolution 1536,376]
1595. ~coll(X23,X24,X25) | midp(X23,X24,X25) [resolution 1574,421]
1600. midp(X23,X24,X25) [subsumption resolution 1595,479]
1607. $false [resolution 1600,481]
% SZS output end Proof for 9point
% ------------------------------
% Version: Vampire 4.4.0 (commit 7916d27 on 2019-08-23 08:50:16 +0100)
% Termination reason: Refutation

% Memory used [KB]: 5500
% Time elapsed: 0.035 s
% ------------------------------
% ------------------------------
