% Refutation found. Thanks to Tanya!
% SZS status Theorem for C103
% SZS output start Proof for C103
fof(f12689,plain,(
  $false),
  inference(avatar_sat_refutation,[],[f1016,f1652,f1786,f2077,f2329,f2489,f2522,f4462,f7608,f11700,f12002,f12688])).
fof(f12688,plain,(
  spl28_156 | ~spl28_78),
  inference(avatar_split_clause,[],[f7378,f1451,f2849])).
fof(f2849,plain,(
  spl28_156 <=> para(sK23,sK20,sK23,sK23)),
  introduced(avatar_definition,[new_symbols(naming,[spl28_156])])).
fof(f1451,plain,(
  spl28_78 <=> cong(sK23,sK23,sK23,sK20)),
  introduced(avatar_definition,[new_symbols(naming,[spl28_78])])).
fof(f7378,plain,(
  para(sK23,sK20,sK23,sK23) | ~spl28_78),
  inference(resolution,[],[f1890,f1452])).
fof(f1452,plain,(
  cong(sK23,sK23,sK23,sK20) | ~spl28_78),
  inference(avatar_component_clause,[],[f1451])).
fof(f1890,plain,(
  ( ! [X0,X1] : (~cong(X0,X0,X0,X1) | para(X0,X1,X0,X0)) )),
  inference(resolution,[],[f770,f398])).
fof(f398,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (~eqangle(X0,X1,X4,X5,X2,X3,X4,X5) | para(X0,X1,X2,X3)) )),
  inference(cnf_transformation,[],[f219])).
fof(f219,plain,(
  ! [X0,X1,X2,X3,X4,X5] : (para(X0,X1,X2,X3) | ~eqangle(X0,X1,X4,X5,X2,X3,X4,X5))),
  inference(ennf_transformation,[],[f118])).
fof(f118,plain,(
  ! [X0,X1,X2,X3,X4,X5] : (eqangle(X0,X1,X4,X5,X2,X3,X4,X5) => para(X0,X1,X2,X3))),
  inference(rectify,[],[f39])).
fof(f39,axiom,(
  ! [X0,X1,X2,X3,X8,X9] : (eqangle(X0,X1,X8,X9,X2,X3,X8,X9) => para(X0,X1,X2,X3))),
  file('./Other/gao/C/C103.p',ruleD39)).
fof(f770,plain,(
  ( ! [X12,X10,X11] : (eqangle(X10,X11,X12,X11,X12,X10,X10,X11) | ~cong(X12,X10,X12,X11)) )),
  inference(resolution,[],[f379,f406])).
fof(f406,plain,(
  ( ! [X2,X0,X1] : (eqangle(X2,X0,X0,X1,X0,X1,X2,X1) | ~cong(X2,X0,X2,X1)) )),
  inference(cnf_transformation,[],[f232])).
fof(f232,plain,(
  ! [X0,X1,X2] : (eqangle(X2,X0,X0,X1,X0,X1,X2,X1) | ~cong(X2,X0,X2,X1))),
  inference(ennf_transformation,[],[f126])).
fof(f126,plain,(
  ! [X0,X1,X2] : (cong(X2,X0,X2,X1) => eqangle(X2,X0,X0,X1,X0,X1,X2,X1))),
  inference(rectify,[],[f47])).
fof(f47,axiom,(
  ! [X0,X1,X7] : (cong(X7,X0,X7,X1) => eqangle(X7,X0,X0,X1,X0,X1,X7,X1))),
  file('./Other/gao/C/C103.p',ruleD46)).
fof(f379,plain,(
  ( ! [X6,X4,X2,X0,X7,X5,X3,X1] : (~eqangle(X0,X1,X2,X3,X4,X5,X6,X7) | eqangle(X4,X5,X6,X7,X0,X1,X2,X3)) )),
  inference(cnf_transformation,[],[f195])).
fof(f195,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6,X7] : (eqangle(X4,X5,X6,X7,X0,X1,X2,X3) | ~eqangle(X0,X1,X2,X3,X4,X5,X6,X7))),
  inference(ennf_transformation,[],[f102])).
fof(f102,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6,X7] : (eqangle(X0,X1,X2,X3,X4,X5,X6,X7) => eqangle(X4,X5,X6,X7,X0,X1,X2,X3))),
  inference(rectify,[],[f20])).
fof(f20,axiom,(
  ! [X0,X1,X2,X3,X8,X9,X10,X11] : (eqangle(X0,X1,X2,X3,X8,X9,X10,X11) => eqangle(X8,X9,X10,X11,X0,X1,X2,X3))),
  file('./Other/gao/C/C103.p',ruleD20)).
fof(f12002,plain,(
  spl28_141 | ~spl28_156),
  inference(avatar_contradiction_clause,[],[f12001])).
fof(f12001,plain,(
  $false | (~spl28_141 | ~spl28_156)),
  inference(subsumption_resolution,[],[f7282,f2038])).
fof(f2038,plain,(
  ~para(sK20,sK23,sK23,sK23) | ~spl28_141),
  inference(avatar_component_clause,[],[f2037])).
fof(f2037,plain,(
  spl28_141 <=> ~para(sK20,sK23,sK23,sK23)),
  introduced(avatar_definition,[new_symbols(naming,[spl28_141])])).
fof(f7282,plain,(
  para(sK20,sK23,sK23,sK23) | ~spl28_156),
  inference(resolution,[],[f1681,f2850])).
fof(f2850,plain,(
  para(sK23,sK20,sK23,sK23) | ~spl28_156),
  inference(avatar_component_clause,[],[f2849])).
fof(f1681,plain,(
  ( ! [X2,X0,X3,X1] : (~para(X0,X1,X2,X3) | para(X1,X0,X2,X3)) )),
  inference(resolution,[],[f742,f398])).
fof(f742,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (eqangle(X0,X1,X2,X3,X4,X5,X2,X3) | ~para(X1,X0,X4,X5)) )),
  inference(resolution,[],[f377,f399])).
fof(f399,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (eqangle(X0,X1,X4,X5,X2,X3,X4,X5) | ~para(X0,X1,X2,X3)) )),
  inference(cnf_transformation,[],[f220])).
fof(f220,plain,(
  ! [X0,X1,X2,X3,X4,X5] : (eqangle(X0,X1,X4,X5,X2,X3,X4,X5) | ~para(X0,X1,X2,X3))),
  inference(ennf_transformation,[],[f119])).
fof(f119,plain,(
  ! [X0,X1,X2,X3,X4,X5] : (para(X0,X1,X2,X3) => eqangle(X0,X1,X4,X5,X2,X3,X4,X5))),
  inference(rectify,[],[f40])).
fof(f40,axiom,(
  ! [X0,X1,X2,X3,X8,X9] : (para(X0,X1,X2,X3) => eqangle(X0,X1,X8,X9,X2,X3,X8,X9))),
  file('./Other/gao/C/C103.p',ruleD40)).
fof(f377,plain,(
  ( ! [X6,X4,X2,X0,X7,X5,X3,X1] : (~eqangle(X0,X1,X2,X3,X4,X5,X6,X7) | eqangle(X1,X0,X2,X3,X4,X5,X6,X7)) )),
  inference(cnf_transformation,[],[f193])).
fof(f193,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6,X7] : (eqangle(X1,X0,X2,X3,X4,X5,X6,X7) | ~eqangle(X0,X1,X2,X3,X4,X5,X6,X7))),
  inference(ennf_transformation,[],[f100])).
fof(f100,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6,X7] : (eqangle(X0,X1,X2,X3,X4,X5,X6,X7) => eqangle(X1,X0,X2,X3,X4,X5,X6,X7))),
  inference(rectify,[],[f18])).
fof(f18,axiom,(
  ! [X0,X1,X2,X3,X8,X9,X10,X11] : (eqangle(X0,X1,X2,X3,X8,X9,X10,X11) => eqangle(X1,X0,X2,X3,X8,X9,X10,X11))),
  file('./Other/gao/C/C103.p',ruleD18)).
fof(f11700,plain,(
  ~spl28_64 | ~spl28_106 | ~spl28_124 | ~spl28_144),
  inference(avatar_contradiction_clause,[],[f11682])).
fof(f11682,plain,(
  $false | (~spl28_64 | ~spl28_106 | ~spl28_124 | ~spl28_144)),
  inference(resolution,[],[f11674,f6664])).
fof(f6664,plain,(
  ~midp(sK25,sK24,sK23)),
  inference(resolution,[],[f1305,f481])).
fof(f481,plain,(
  ~cong(sK23,sK25,sK25,sK24)),
  inference(cnf_transformation,[],[f359])).
fof(f359,plain,(
  ~cong(sK23,sK25,sK25,sK24) & para(sK21,sK22,sK23,sK26) & coll(sK26,sK20,sK21) & coll(sK25,sK23,sK24) & coll(sK25,sK20,sK21) & coll(sK24,sK21,sK22) & cong(sK24,sK21,sK20,sK23) & coll(sK23,sK20,sK22) & perp(sK20,sK21,sK27,sK22)),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK20,sK21,sK22,sK23,sK24,sK25,sK26,sK27])],[f321,f358])).
fof(f358,plain,(
  ? [X0,X1,X2,X3,X4,X5,X6,X7] : (~cong(X3,X5,X5,X4) & para(X1,X2,X3,X6) & coll(X6,X0,X1) & coll(X5,X3,X4) & coll(X5,X0,X1) & coll(X4,X1,X2) & cong(X4,X1,X0,X3) & coll(X3,X0,X2) & perp(X0,X1,X7,X2)) => (~cong(sK23,sK25,sK25,sK24) & para(sK21,sK22,sK23,sK26) & coll(sK26,sK20,sK21) & coll(sK25,sK23,sK24) & coll(sK25,sK20,sK21) & coll(sK24,sK21,sK22) & cong(sK24,sK21,sK20,sK23) & coll(sK23,sK20,sK22) & perp(sK20,sK21,sK27,sK22))),
  introduced(choice_axiom,[])).
fof(f321,plain,(
  ? [X0,X1,X2,X3,X4,X5,X6,X7] : (~cong(X3,X5,X5,X4) & para(X1,X2,X3,X6) & coll(X6,X0,X1) & coll(X5,X3,X4) & coll(X5,X0,X1) & coll(X4,X1,X2) & cong(X4,X1,X0,X3) & coll(X3,X0,X2) & perp(X0,X1,X7,X2))),
  inference(flattening,[],[f320])).
fof(f320,plain,(
  ? [X0,X1,X2,X3,X4,X5,X6,X7] : (~cong(X3,X5,X5,X4) & (para(X1,X2,X3,X6) & coll(X6,X0,X1) & coll(X5,X3,X4) & coll(X5,X0,X1) & coll(X4,X1,X2) & cong(X4,X1,X0,X3) & coll(X3,X0,X2) & perp(X0,X1,X7,X2)))),
  inference(ennf_transformation,[],[f168])).
fof(f168,plain,(
  ~! [X0,X1,X2,X3,X4,X5,X6,X7] : ((para(X1,X2,X3,X6) & coll(X6,X0,X1) & coll(X5,X3,X4) & coll(X5,X0,X1) & coll(X4,X1,X2) & cong(X4,X1,X0,X3) & coll(X3,X0,X2) & perp(X0,X1,X7,X2)) => cong(X3,X5,X5,X4))),
  inference(pure_predicate_removal,[],[f167])).
fof(f167,plain,(
  ~! [X0,X1,X2,X3,X4,X5,X6,X7] : ((para(X1,X2,X3,X6) & coll(X6,X0,X1) & coll(X5,X3,X4) & coll(X5,X0,X1) & coll(X4,X1,X2) & cong(X4,X1,X0,X3) & coll(X3,X0,X2) & perp(X0,X1,X7,X2) & midpoint(X7,X0,X1)) => cong(X3,X5,X5,X4))),
  inference(rectify,[],[f96])).
fof(f96,negated_conjecture,(
  ~! [X0,X1,X2,X3,X4,X5,X12,X18] : ((para(X1,X2,X3,X12) & coll(X12,X0,X1) & coll(X5,X3,X4) & coll(X5,X0,X1) & coll(X4,X1,X2) & cong(X4,X1,X0,X3) & coll(X3,X0,X2) & perp(X0,X1,X18,X2) & midpoint(X18,X0,X1)) => cong(X3,X5,X5,X4))),
  inference(negated_conjecture,[],[f95])).
fof(f95,conjecture,(
  ! [X0,X1,X2,X3,X4,X5,X12,X18] : ((para(X1,X2,X3,X12) & coll(X12,X0,X1) & coll(X5,X3,X4) & coll(X5,X0,X1) & coll(X4,X1,X2) & cong(X4,X1,X0,X3) & coll(X3,X0,X2) & perp(X0,X1,X18,X2) & midpoint(X18,X0,X1)) => cong(X3,X5,X5,X4))),
  file('./Other/gao/C/C103.p',exemploOthergaoCC103)).
fof(f1305,plain,(
  ( ! [X6,X7,X5] : (cong(X7,X5,X5,X6) | ~midp(X5,X6,X7)) )),
  inference(resolution,[],[f523,f383])).
fof(f383,plain,(
  ( ! [X2,X0,X3,X1] : (~cong(X0,X1,X2,X3) | cong(X2,X3,X0,X1)) )),
  inference(cnf_transformation,[],[f200])).
fof(f200,plain,(
  ! [X0,X1,X2,X3] : (cong(X2,X3,X0,X1) | ~cong(X0,X1,X2,X3))),
  inference(ennf_transformation,[],[f24])).
fof(f24,axiom,(
  ! [X0,X1,X2,X3] : (cong(X0,X1,X2,X3) => cong(X2,X3,X0,X1))),
  file('./Other/gao/C/C103.p',ruleD24)).
fof(f523,plain,(
  ( ! [X2,X0,X1] : (cong(X0,X1,X2,X0) | ~midp(X0,X1,X2)) )),
  inference(resolution,[],[f382,f428])).
fof(f428,plain,(
  ( ! [X2,X0,X1] : (cong(X0,X1,X0,X2) | ~midp(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f271])).
fof(f271,plain,(
  ! [X0,X1,X2] : (cong(X0,X1,X0,X2) | ~midp(X0,X1,X2))),
  inference(ennf_transformation,[],[f69])).
fof(f69,axiom,(
  ! [X0,X1,X2] : (midp(X0,X1,X2) => cong(X0,X1,X0,X2))),
  file('./Other/gao/C/C103.p',ruleD68)).
fof(f382,plain,(
  ( ! [X2,X0,X3,X1] : (~cong(X0,X1,X2,X3) | cong(X0,X1,X3,X2)) )),
  inference(cnf_transformation,[],[f199])).
fof(f199,plain,(
  ! [X0,X1,X2,X3] : (cong(X0,X1,X3,X2) | ~cong(X0,X1,X2,X3))),
  inference(ennf_transformation,[],[f23])).
fof(f23,axiom,(
  ! [X0,X1,X2,X3] : (cong(X0,X1,X2,X3) => cong(X0,X1,X3,X2))),
  file('./Other/gao/C/C103.p',ruleD23)).
fof(f11674,plain,(
  ( ! [X3] : (midp(sK25,X3,sK23)) ) | (~spl28_64 | ~spl28_106 | ~spl28_124 | ~spl28_144)),
  inference(subsumption_resolution,[],[f11673,f7335])).
fof(f7335,plain,(
  ( ! [X11] : (midp(sK23,X11,sK23)) ) | (~spl28_64 | ~spl28_124 | ~spl28_144)),
  inference(resolution,[],[f7010,f6848])).
fof(f6848,plain,(
  ( ! [X10,X11,X9] : (~midp(X9,X10,X11) | midp(X10,X9,X11)) ) | (~spl28_124 | ~spl28_144)),
  inference(subsumption_resolution,[],[f6807,f2700])).
fof(f2700,plain,(
  ( ! [X2,X0,X1] : (coll(X0,X1,X2)) ) | ~spl28_124),
  inference(subsumption_resolution,[],[f2697,f2680])).
fof(f2680,plain,(
  ( ! [X6,X5] : (coll(X5,X5,X6)) ) | ~spl28_124),
  inference(resolution,[],[f2655,f360])).
fof(f360,plain,(
  ( ! [X2,X0,X1] : (~coll(X0,X1,X2) | coll(X0,X2,X1)) )),
  inference(cnf_transformation,[],[f169])).
fof(f169,plain,(
  ! [X0,X1,X2] : (coll(X0,X2,X1) | ~coll(X0,X1,X2))),
  inference(ennf_transformation,[],[f1])).
fof(f1,axiom,(
  ! [X0,X1,X2] : (coll(X0,X1,X2) => coll(X0,X2,X1))),
  file('./Other/gao/C/C103.p',ruleD1)).
fof(f2655,plain,(
  ( ! [X4,X3] : (coll(X3,X4,X3)) ) | ~spl28_124),
  inference(resolution,[],[f2640,f361])).
fof(f361,plain,(
  ( ! [X2,X0,X1] : (~coll(X0,X1,X2) | coll(X1,X0,X2)) )),
  inference(cnf_transformation,[],[f170])).
fof(f170,plain,(
  ! [X0,X1,X2] : (coll(X1,X0,X2) | ~coll(X0,X1,X2))),
  inference(ennf_transformation,[],[f2])).
fof(f2,axiom,(
  ! [X0,X1,X2] : (coll(X0,X1,X2) => coll(X1,X0,X2))),
  file('./Other/gao/C/C103.p',ruleD2)).
fof(f2640,plain,(
  ( ! [X0,X1] : (coll(X0,X1,X1)) ) | ~spl28_124),
  inference(resolution,[],[f1782,f426])).
fof(f426,plain,(
  ( ! [X2,X0,X1] : (~para(X0,X1,X0,X2) | coll(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f268])).
fof(f268,plain,(
  ! [X0,X1,X2] : (coll(X0,X1,X2) | ~para(X0,X1,X0,X2))),
  inference(ennf_transformation,[],[f67])).
fof(f67,axiom,(
  ! [X0,X1,X2] : (para(X0,X1,X0,X2) => coll(X0,X1,X2))),
  file('./Other/gao/C/C103.p',ruleD66)).
fof(f1782,plain,(
  ( ! [X2,X3] : (para(X2,X3,X2,X3)) ) | ~spl28_124),
  inference(avatar_component_clause,[],[f1781])).
fof(f1781,plain,(
  spl28_124 <=> ! [X3,X2] : para(X2,X3,X2,X3)),
  introduced(avatar_definition,[new_symbols(naming,[spl28_124])])).
fof(f2697,plain,(
  ( ! [X2,X0,X1] : (coll(X0,X1,X2) | ~coll(X2,X2,X0)) ) | ~spl28_124),
  inference(resolution,[],[f2680,f362])).
fof(f362,plain,(
  ( ! [X2,X0,X3,X1] : (~coll(X0,X1,X3) | coll(X2,X3,X0) | ~coll(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f172])).
fof(f172,plain,(
  ! [X0,X1,X2,X3] : (coll(X2,X3,X0) | ~coll(X0,X1,X3) | ~coll(X0,X1,X2))),
  inference(flattening,[],[f171])).
fof(f171,plain,(
  ! [X0,X1,X2,X3] : (coll(X2,X3,X0) | (~coll(X0,X1,X3) | ~coll(X0,X1,X2)))),
  inference(ennf_transformation,[],[f3])).
fof(f3,axiom,(
  ! [X0,X1,X2,X3] : ((coll(X0,X1,X3) & coll(X0,X1,X2)) => coll(X2,X3,X0))),
  file('./Other/gao/C/C103.p',ruleD3)).
fof(f6807,plain,(
  ( ! [X10,X11,X9] : (~midp(X9,X10,X11) | ~coll(X10,X9,X11) | midp(X10,X9,X11)) ) | ~spl28_144),
  inference(resolution,[],[f2324,f427])).
fof(f427,plain,(
  ( ! [X2,X0,X1] : (~cong(X0,X1,X0,X2) | ~coll(X0,X1,X2) | midp(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f270])).
fof(f270,plain,(
  ! [X0,X1,X2] : (midp(X0,X1,X2) | ~coll(X0,X1,X2) | ~cong(X0,X1,X0,X2))),
  inference(flattening,[],[f269])).
fof(f269,plain,(
  ! [X0,X1,X2] : (midp(X0,X1,X2) | (~coll(X0,X1,X2) | ~cong(X0,X1,X0,X2)))),
  inference(ennf_transformation,[],[f68])).
fof(f68,axiom,(
  ! [X0,X1,X2] : ((coll(X0,X1,X2) & cong(X0,X1,X0,X2)) => midp(X0,X1,X2))),
  file('./Other/gao/C/C103.p',ruleD67)).
fof(f2324,plain,(
  ( ! [X14,X15,X13] : (cong(X14,X13,X14,X15) | ~midp(X13,X14,X15)) ) | ~spl28_144),
  inference(avatar_component_clause,[],[f2323])).
fof(f2323,plain,(
  spl28_144 <=> ! [X13,X15,X14] : (~midp(X13,X14,X15) | cong(X14,X13,X14,X15))),
  introduced(avatar_definition,[new_symbols(naming,[spl28_144])])).
fof(f7010,plain,(
  ( ! [X1] : (midp(X1,sK23,sK23)) ) | (~spl28_64 | ~spl28_124)),
  inference(resolution,[],[f6921,f4246])).
fof(f4246,plain,(
  ( ! [X8] : (midp(X8,sK23,X8)) ) | ~spl28_124),
  inference(resolution,[],[f4177,f370])).
fof(f370,plain,(
  ( ! [X2,X0,X1] : (~midp(X2,X1,X0) | midp(X2,X0,X1)) )),
  inference(cnf_transformation,[],[f183])).
fof(f183,plain,(
  ! [X0,X1,X2] : (midp(X2,X0,X1) | ~midp(X2,X1,X0))),
  inference(ennf_transformation,[],[f97])).
fof(f97,plain,(
  ! [X0,X1,X2] : (midp(X2,X1,X0) => midp(X2,X0,X1))),
  inference(rectify,[],[f11])).
fof(f11,axiom,(
  ! [X0,X1,X6] : (midp(X6,X1,X0) => midp(X6,X0,X1))),
  file('./Other/gao/C/C103.p',ruleD11)).
fof(f4177,plain,(
  ( ! [X19] : (midp(X19,X19,sK23)) ) | ~spl28_124),
  inference(resolution,[],[f2733,f3400])).
fof(f3400,plain,(
  ( ! [X24] : (midp(sK23,X24,X24)) ) | ~spl28_124),
  inference(resolution,[],[f2650,f2790])).
fof(f2790,plain,(
  midp(sK23,sK20,sK20)),
  inference(subsumption_resolution,[],[f2774,f1080])).
fof(f1080,plain,(
  coll(sK23,sK20,sK20)),
  inference(resolution,[],[f1076,f361])).
fof(f1076,plain,(
  coll(sK20,sK23,sK20)),
  inference(resolution,[],[f1070,f360])).
fof(f1070,plain,(
  coll(sK20,sK20,sK23)),
  inference(resolution,[],[f533,f482])).
fof(f482,plain,(
  coll(sK23,sK22,sK20)),
  inference(resolution,[],[f360,f474])).
fof(f474,plain,(
  coll(sK23,sK20,sK22)),
  inference(cnf_transformation,[],[f359])).
fof(f533,plain,(
  ( ! [X3] : (~coll(sK23,sK22,X3) | coll(X3,sK20,sK23)) )),
  inference(resolution,[],[f362,f482])).
fof(f2774,plain,(
  ~coll(sK23,sK20,sK20) | midp(sK23,sK20,sK20)),
  inference(resolution,[],[f2271,f427])).
fof(f2271,plain,(
  cong(sK23,sK20,sK23,sK20)),
  inference(resolution,[],[f1370,f382])).
fof(f1370,plain,(
  cong(sK23,sK20,sK20,sK23)),
  inference(resolution,[],[f696,f592])).
fof(f592,plain,(
  cong(sK23,sK20,sK24,sK21)),
  inference(resolution,[],[f522,f383])).
fof(f522,plain,(
  cong(sK24,sK21,sK23,sK20)),
  inference(resolution,[],[f382,f475])).
fof(f475,plain,(
  cong(sK24,sK21,sK20,sK23)),
  inference(cnf_transformation,[],[f359])).
fof(f696,plain,(
  ( ! [X0,X1] : (~cong(X0,X1,sK24,sK21) | cong(X0,X1,sK20,sK23)) )),
  inference(resolution,[],[f384,f475])).
fof(f384,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (~cong(X2,X3,X4,X5) | cong(X0,X1,X4,X5) | ~cong(X0,X1,X2,X3)) )),
  inference(cnf_transformation,[],[f202])).
fof(f202,plain,(
  ! [X0,X1,X2,X3,X4,X5] : (cong(X0,X1,X4,X5) | ~cong(X2,X3,X4,X5) | ~cong(X0,X1,X2,X3))),
  inference(flattening,[],[f201])).
fof(f201,plain,(
  ! [X0,X1,X2,X3,X4,X5] : (cong(X0,X1,X4,X5) | (~cong(X2,X3,X4,X5) | ~cong(X0,X1,X2,X3)))),
  inference(ennf_transformation,[],[f25])).
fof(f25,axiom,(
  ! [X0,X1,X2,X3,X4,X5] : ((cong(X2,X3,X4,X5) & cong(X0,X1,X2,X3)) => cong(X0,X1,X4,X5))),
  file('./Other/gao/C/C103.p',ruleD25)).
fof(f2650,plain,(
  ( ! [X10,X8,X9] : (~midp(X8,X10,X10) | midp(X8,X9,X9)) ) | ~spl28_124),
  inference(subsumption_resolution,[],[f2643,f1782])).
fof(f2643,plain,(
  ( ! [X10,X8,X9] : (midp(X8,X9,X9) | ~para(X10,X9,X10,X9) | ~midp(X8,X10,X10)) ) | ~spl28_124),
  inference(resolution,[],[f1782,f424])).
fof(f424,plain,(
  ( ! [X4,X2,X0,X3,X1] : (~para(X0,X3,X1,X2) | midp(X4,X2,X3) | ~para(X0,X2,X1,X3) | ~midp(X4,X0,X1)) )),
  inference(cnf_transformation,[],[f265])).
fof(f265,plain,(
  ! [X0,X1,X2,X3,X4] : (midp(X4,X2,X3) | ~para(X0,X3,X1,X2) | ~para(X0,X2,X1,X3) | ~midp(X4,X0,X1))),
  inference(flattening,[],[f264])).
fof(f264,plain,(
  ! [X0,X1,X2,X3,X4] : (midp(X4,X2,X3) | (~para(X0,X3,X1,X2) | ~para(X0,X2,X1,X3) | ~midp(X4,X0,X1)))),
  inference(ennf_transformation,[],[f143])).
fof(f143,plain,(
  ! [X0,X1,X2,X3,X4] : ((para(X0,X3,X1,X2) & para(X0,X2,X1,X3) & midp(X4,X0,X1)) => midp(X4,X2,X3))),
  inference(rectify,[],[f65])).
fof(f65,axiom,(
  ! [X0,X1,X2,X3,X6] : ((para(X0,X3,X1,X2) & para(X0,X2,X1,X3) & midp(X6,X0,X1)) => midp(X6,X2,X3))),
  file('./Other/gao/C/C103.p',ruleD64)).
fof(f2733,plain,(
  ( ! [X10,X8,X9] : (~midp(X10,X9,X8) | midp(X8,X9,X10)) ) | ~spl28_124),
  inference(subsumption_resolution,[],[f2727,f2700])).
fof(f2727,plain,(
  ( ! [X10,X8,X9] : (~coll(X8,X9,X10) | midp(X8,X9,X10) | ~midp(X10,X9,X8)) ) | ~spl28_124),
  inference(resolution,[],[f2647,f405])).
fof(f405,plain,(
  ( ! [X4,X2,X0,X3,X1] : (~para(X3,X4,X1,X2) | ~coll(X4,X0,X2) | midp(X4,X0,X2) | ~midp(X3,X0,X1)) )),
  inference(cnf_transformation,[],[f231])).
fof(f231,plain,(
  ! [X0,X1,X2,X3,X4] : (midp(X4,X0,X2) | ~coll(X4,X0,X2) | ~para(X3,X4,X1,X2) | ~midp(X3,X0,X1))),
  inference(flattening,[],[f230])).
fof(f230,plain,(
  ! [X0,X1,X2,X3,X4] : (midp(X4,X0,X2) | (~coll(X4,X0,X2) | ~para(X3,X4,X1,X2) | ~midp(X3,X0,X1)))),
  inference(ennf_transformation,[],[f125])).
fof(f125,plain,(
  ! [X0,X1,X2,X3,X4] : ((coll(X4,X0,X2) & para(X3,X4,X1,X2) & midp(X3,X0,X1)) => midp(X4,X0,X2))),
  inference(rectify,[],[f46])).
fof(f46,axiom,(
  ! [X0,X1,X2,X4,X5] : ((coll(X5,X0,X2) & para(X4,X5,X1,X2) & midp(X4,X0,X1)) => midp(X5,X0,X2))),
  file('./Other/gao/C/C103.p',ruleD45)).
fof(f2647,plain,(
  ( ! [X21,X20] : (para(X20,X21,X21,X20)) ) | ~spl28_124),
  inference(resolution,[],[f1782,f363])).
fof(f363,plain,(
  ( ! [X2,X0,X3,X1] : (~para(X0,X1,X2,X3) | para(X0,X1,X3,X2)) )),
  inference(cnf_transformation,[],[f173])).
fof(f173,plain,(
  ! [X0,X1,X2,X3] : (para(X0,X1,X3,X2) | ~para(X0,X1,X2,X3))),
  inference(ennf_transformation,[],[f4])).
fof(f4,axiom,(
  ! [X0,X1,X2,X3] : (para(X0,X1,X2,X3) => para(X0,X1,X3,X2))),
  file('./Other/gao/C/C103.p',ruleD4)).
fof(f6921,plain,(
  ( ! [X8,X7] : (~midp(X7,X8,X7) | midp(X7,X8,X8)) ) | (~spl28_64 | ~spl28_124)),
  inference(subsumption_resolution,[],[f6888,f2640])).
fof(f6888,plain,(
  ( ! [X8,X7] : (~midp(X7,X8,X7) | ~coll(X7,X8,X8) | midp(X7,X8,X8)) ) | ~spl28_64),
  inference(resolution,[],[f2538,f427])).
fof(f2538,plain,(
  ( ! [X47,X48,X46] : (cong(X46,X47,X48,X47) | ~midp(X46,X47,X48)) ) | ~spl28_64),
  inference(resolution,[],[f1012,f382])).
fof(f1012,plain,(
  ( ! [X10,X11,X9] : (cong(X9,X10,X10,X11) | ~midp(X9,X10,X11)) ) | ~spl28_64),
  inference(avatar_component_clause,[],[f1011])).
fof(f1011,plain,(
  spl28_64 <=> ! [X9,X11,X10] : (cong(X9,X10,X10,X11) | ~midp(X9,X10,X11))),
  introduced(avatar_definition,[new_symbols(naming,[spl28_64])])).
fof(f11673,plain,(
  ( ! [X3] : (midp(sK25,X3,sK23) | ~midp(sK23,X3,sK23)) ) | (~spl28_106 | ~spl28_124)),
  inference(subsumption_resolution,[],[f11664,f2700])).
fof(f11664,plain,(
  ( ! [X3] : (~coll(sK25,X3,sK23) | midp(sK25,X3,sK23) | ~midp(sK23,X3,sK23)) ) | ~spl28_106),
  inference(resolution,[],[f7630,f405])).
fof(f7630,plain,(
  para(sK23,sK25,sK23,sK23) | ~spl28_106),
  inference(resolution,[],[f1543,f1890])).
fof(f1543,plain,(
  cong(sK23,sK23,sK23,sK25) | ~spl28_106),
  inference(avatar_component_clause,[],[f1542])).
fof(f1542,plain,(
  spl28_106 <=> cong(sK23,sK23,sK23,sK25)),
  introduced(avatar_definition,[new_symbols(naming,[spl28_106])])).
fof(f7608,plain,(
  ~spl28_64 | spl28_107 | ~spl28_124),
  inference(avatar_contradiction_clause,[],[f7596])).
fof(f7596,plain,(
  $false | (~spl28_64 | ~spl28_107 | ~spl28_124)),
  inference(resolution,[],[f7333,f1888])).
fof(f1888,plain,(
  ~midp(sK23,sK23,sK25) | ~spl28_107),
  inference(resolution,[],[f1546,f428])).
fof(f1546,plain,(
  ~cong(sK23,sK23,sK23,sK25) | ~spl28_107),
  inference(avatar_component_clause,[],[f1545])).
fof(f1545,plain,(
  spl28_107 <=> ~cong(sK23,sK23,sK23,sK25)),
  introduced(avatar_definition,[new_symbols(naming,[spl28_107])])).
fof(f7333,plain,(
  ( ! [X9] : (midp(sK23,sK23,X9)) ) | (~spl28_64 | ~spl28_124)),
  inference(resolution,[],[f7010,f2733])).
fof(f4462,plain,(
  spl28_79 | ~spl28_120 | ~spl28_124),
  inference(avatar_contradiction_clause,[],[f4461])).
fof(f4461,plain,(
  $false | (~spl28_79 | ~spl28_120 | ~spl28_124)),
  inference(subsumption_resolution,[],[f4182,f2331])).
fof(f2331,plain,(
  ~midp(sK23,sK23,sK20) | ~spl28_79),
  inference(resolution,[],[f1455,f428])).
fof(f1455,plain,(
  ~cong(sK23,sK23,sK23,sK20) | ~spl28_79),
  inference(avatar_component_clause,[],[f1454])).
fof(f1454,plain,(
  spl28_79 <=> ~cong(sK23,sK23,sK23,sK20)),
  introduced(avatar_definition,[new_symbols(naming,[spl28_79])])).
fof(f4182,plain,(
  midp(sK23,sK23,sK20) | (~spl28_120 | ~spl28_124)),
  inference(resolution,[],[f2733,f1645])).
fof(f1645,plain,(
  midp(sK20,sK23,sK23) | ~spl28_120),
  inference(avatar_component_clause,[],[f1644])).
fof(f1644,plain,(
  spl28_120 <=> midp(sK20,sK23,sK23)),
  introduced(avatar_definition,[new_symbols(naming,[spl28_120])])).
fof(f2522,plain,(
  ~spl28_126),
  inference(avatar_contradiction_clause,[],[f2519])).
fof(f2519,plain,(
  $false | ~spl28_126),
  inference(resolution,[],[f1785,f1347])).
fof(f1347,plain,(
  para(sK23,sK26,sK23,sK26)),
  inference(resolution,[],[f644,f507])).
fof(f507,plain,(
  para(sK23,sK26,sK21,sK22)),
  inference(resolution,[],[f364,f480])).
fof(f480,plain,(
  para(sK21,sK22,sK23,sK26)),
  inference(cnf_transformation,[],[f359])).
fof(f364,plain,(
  ( ! [X2,X0,X3,X1] : (~para(X0,X1,X2,X3) | para(X2,X3,X0,X1)) )),
  inference(cnf_transformation,[],[f174])).
fof(f174,plain,(
  ! [X0,X1,X2,X3] : (para(X2,X3,X0,X1) | ~para(X0,X1,X2,X3))),
  inference(ennf_transformation,[],[f5])).
fof(f5,axiom,(
  ! [X0,X1,X2,X3] : (para(X0,X1,X2,X3) => para(X2,X3,X0,X1))),
  file('./Other/gao/C/C103.p',ruleD5)).
fof(f644,plain,(
  ( ! [X0,X1] : (~para(X0,X1,sK21,sK22) | para(X0,X1,sK23,sK26)) )),
  inference(resolution,[],[f365,f480])).
fof(f365,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (~para(X2,X3,X4,X5) | para(X0,X1,X4,X5) | ~para(X0,X1,X2,X3)) )),
  inference(cnf_transformation,[],[f176])).
fof(f176,plain,(
  ! [X0,X1,X2,X3,X4,X5] : (para(X0,X1,X4,X5) | ~para(X2,X3,X4,X5) | ~para(X0,X1,X2,X3))),
  inference(flattening,[],[f175])).
fof(f175,plain,(
  ! [X0,X1,X2,X3,X4,X5] : (para(X0,X1,X4,X5) | (~para(X2,X3,X4,X5) | ~para(X0,X1,X2,X3)))),
  inference(ennf_transformation,[],[f6])).
fof(f6,axiom,(
  ! [X0,X1,X2,X3,X4,X5] : ((para(X2,X3,X4,X5) & para(X0,X1,X2,X3)) => para(X0,X1,X4,X5))),
  file('./Other/gao/C/C103.p',ruleD6)).
fof(f1785,plain,(
  ( ! [X0,X1] : (~para(X0,X1,X0,X1)) ) | ~spl28_126),
  inference(avatar_component_clause,[],[f1784])).
fof(f1784,plain,(
  spl28_126 <=> ! [X1,X0] : ~para(X0,X1,X0,X1)),
  introduced(avatar_definition,[new_symbols(naming,[spl28_126])])).
fof(f2489,plain,(
  ~spl28_66 | ~spl28_120 | ~spl28_124 | ~spl28_140),
  inference(avatar_contradiction_clause,[],[f2488])).
fof(f2488,plain,(
  $false | (~spl28_66 | ~spl28_120 | ~spl28_124 | ~spl28_140)),
  inference(resolution,[],[f2461,f1645])).
fof(f2461,plain,(
  ( ! [X1] : (~midp(sK20,X1,sK23)) ) | (~spl28_66 | ~spl28_124 | ~spl28_140)),
  inference(subsumption_resolution,[],[f2460,f2338])).
fof(f2338,plain,(
  ( ! [X0,X1] : (~midp(X0,X1,X0)) ) | ~spl28_66),
  inference(subsumption_resolution,[],[f2332,f370])).
fof(f2332,plain,(
  ( ! [X0,X1] : (~midp(X0,X0,X1) | ~midp(X0,X1,X0)) ) | ~spl28_66),
  inference(resolution,[],[f1015,f527])).
fof(f527,plain,(
  ( ! [X2,X0,X1] : (cong(X0,X1,X0,X2) | ~midp(X0,X2,X1)) )),
  inference(resolution,[],[f383,f428])).
fof(f1015,plain,(
  ( ! [X6,X8,X7] : (~cong(X6,X7,X7,X8) | ~midp(X6,X7,X8)) ) | ~spl28_66),
  inference(avatar_component_clause,[],[f1014])).
fof(f1014,plain,(
  spl28_66 <=> ! [X7,X8,X6] : (~cong(X6,X7,X7,X8) | ~midp(X6,X7,X8))),
  introduced(avatar_definition,[new_symbols(naming,[spl28_66])])).
fof(f2460,plain,(
  ( ! [X1] : (midp(sK23,X1,sK23) | ~midp(sK20,X1,sK23)) ) | (~spl28_124 | ~spl28_140)),
  inference(subsumption_resolution,[],[f2455,f2061])).
fof(f2061,plain,(
  ( ! [X4,X3] : (coll(X3,X4,X3)) ) | ~spl28_124),
  inference(resolution,[],[f2049,f361])).
fof(f2049,plain,(
  ( ! [X0,X1] : (coll(X0,X1,X1)) ) | ~spl28_124),
  inference(resolution,[],[f1782,f426])).
fof(f2455,plain,(
  ( ! [X1] : (~coll(sK23,X1,sK23) | midp(sK23,X1,sK23) | ~midp(sK20,X1,sK23)) ) | ~spl28_140),
  inference(resolution,[],[f2041,f405])).
fof(f2041,plain,(
  para(sK20,sK23,sK23,sK23) | ~spl28_140),
  inference(avatar_component_clause,[],[f2040])).
fof(f2040,plain,(
  spl28_140 <=> para(sK20,sK23,sK23,sK23)),
  introduced(avatar_definition,[new_symbols(naming,[spl28_140])])).
fof(f2329,plain,(
  spl28_144 | spl28_66),
  inference(avatar_split_clause,[],[f2316,f1014,f2323])).
fof(f2316,plain,(
  ( ! [X14,X12,X10,X15,X13,X11] : (~midp(X10,X11,X12) | ~midp(X13,X14,X15) | ~cong(X10,X11,X11,X12) | cong(X14,X13,X14,X15)) )),
  inference(resolution,[],[f783,f435])).
fof(f435,plain,(
  ( ! [X6,X4,X2,X0,X7,X5,X3,X1] : (~eqratio(X0,X1,X2,X3,X4,X5,X6,X7) | ~cong(X4,X5,X6,X7) | cong(X0,X1,X2,X3)) )),
  inference(cnf_transformation,[],[f284])).
fof(f284,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6,X7] : (cong(X0,X1,X2,X3) | ~cong(X4,X5,X6,X7) | ~eqratio(X0,X1,X2,X3,X4,X5,X6,X7))),
  inference(flattening,[],[f283])).
fof(f283,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6,X7] : (cong(X0,X1,X2,X3) | (~cong(X4,X5,X6,X7) | ~eqratio(X0,X1,X2,X3,X4,X5,X6,X7)))),
  inference(ennf_transformation,[],[f148])).
fof(f148,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6,X7] : ((cong(X4,X5,X6,X7) & eqratio(X0,X1,X2,X3,X4,X5,X6,X7)) => cong(X0,X1,X2,X3))),
  inference(rectify,[],[f76])).
fof(f76,axiom,(
  ! [X0,X1,X2,X3,X8,X9,X10,X11] : ((cong(X8,X9,X10,X11) & eqratio(X0,X1,X2,X3,X8,X9,X10,X11)) => cong(X0,X1,X2,X3))),
  file('./Other/gao/C/C103.p',ruleD75)).
fof(f783,plain,(
  ( ! [X6,X10,X8,X7,X11,X9] : (eqratio(X6,X7,X6,X8,X9,X10,X10,X11) | ~midp(X9,X10,X11) | ~midp(X7,X6,X8)) )),
  inference(resolution,[],[f385,f430])).
fof(f430,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (eqratio(X4,X0,X0,X1,X5,X2,X2,X3) | ~midp(X5,X2,X3) | ~midp(X4,X0,X1)) )),
  inference(cnf_transformation,[],[f274])).
fof(f274,plain,(
  ! [X0,X1,X2,X3,X4,X5] : (eqratio(X4,X0,X0,X1,X5,X2,X2,X3) | ~midp(X5,X2,X3) | ~midp(X4,X0,X1))),
  inference(flattening,[],[f273])).
fof(f273,plain,(
  ! [X0,X1,X2,X3,X4,X5] : (eqratio(X4,X0,X0,X1,X5,X2,X2,X3) | (~midp(X5,X2,X3) | ~midp(X4,X0,X1)))),
  inference(ennf_transformation,[],[f145])).
fof(f145,plain,(
  ! [X0,X1,X2,X3,X4,X5] : ((midp(X5,X2,X3) & midp(X4,X0,X1)) => eqratio(X4,X0,X0,X1,X5,X2,X2,X3))),
  inference(rectify,[],[f71])).
fof(f71,axiom,(
  ! [X0,X1,X2,X3,X6,X16] : ((midp(X16,X2,X3) & midp(X6,X0,X1)) => eqratio(X6,X0,X0,X1,X16,X2,X2,X3))),
  file('./Other/gao/C/C103.p',ruleD70)).
fof(f385,plain,(
  ( ! [X6,X4,X2,X0,X7,X5,X3,X1] : (~eqratio(X0,X1,X2,X3,X4,X5,X6,X7) | eqratio(X1,X0,X2,X3,X4,X5,X6,X7)) )),
  inference(cnf_transformation,[],[f203])).
fof(f203,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6,X7] : (eqratio(X1,X0,X2,X3,X4,X5,X6,X7) | ~eqratio(X0,X1,X2,X3,X4,X5,X6,X7))),
  inference(ennf_transformation,[],[f105])).
fof(f105,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6,X7] : (eqratio(X0,X1,X2,X3,X4,X5,X6,X7) => eqratio(X1,X0,X2,X3,X4,X5,X6,X7))),
  inference(rectify,[],[f26])).
fof(f26,axiom,(
  ! [X0,X1,X2,X3,X8,X9,X10,X11] : (eqratio(X0,X1,X2,X3,X8,X9,X10,X11) => eqratio(X1,X0,X2,X3,X8,X9,X10,X11))),
  file('./Other/gao/C/C103.p',ruleD26)).
fof(f2077,plain,(
  spl28_123 | ~spl28_124),
  inference(avatar_contradiction_clause,[],[f2066])).
fof(f2066,plain,(
  $false | (~spl28_123 | ~spl28_124)),
  inference(resolution,[],[f2049,f1651])).
fof(f1651,plain,(
  ~coll(sK20,sK23,sK23) | ~spl28_123),
  inference(avatar_component_clause,[],[f1650])).
fof(f1650,plain,(
  spl28_123 <=> ~coll(sK20,sK23,sK23)),
  introduced(avatar_definition,[new_symbols(naming,[spl28_123])])).
fof(f1786,plain,(
  spl28_124 | spl28_126),
  inference(avatar_split_clause,[],[f1757,f1784,f1781])).
fof(f1757,plain,(
  ( ! [X2,X0,X3,X1] : (~para(X0,X1,X0,X1) | para(X2,X3,X2,X3)) )),
  inference(resolution,[],[f764,f398])).
fof(f764,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (eqangle(X0,X1,X2,X3,X0,X1,X4,X5) | ~para(X2,X3,X4,X5)) )),
  inference(resolution,[],[f378,f399])).
fof(f378,plain,(
  ( ! [X6,X4,X2,X0,X7,X5,X3,X1] : (~eqangle(X0,X1,X2,X3,X4,X5,X6,X7) | eqangle(X2,X3,X0,X1,X6,X7,X4,X5)) )),
  inference(cnf_transformation,[],[f194])).
fof(f194,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6,X7] : (eqangle(X2,X3,X0,X1,X6,X7,X4,X5) | ~eqangle(X0,X1,X2,X3,X4,X5,X6,X7))),
  inference(ennf_transformation,[],[f101])).
fof(f101,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6,X7] : (eqangle(X0,X1,X2,X3,X4,X5,X6,X7) => eqangle(X2,X3,X0,X1,X6,X7,X4,X5))),
  inference(rectify,[],[f19])).
fof(f19,axiom,(
  ! [X0,X1,X2,X3,X8,X9,X10,X11] : (eqangle(X0,X1,X2,X3,X8,X9,X10,X11) => eqangle(X2,X3,X0,X1,X10,X11,X8,X9))),
  file('./Other/gao/C/C103.p',ruleD19)).
fof(f1652,plain,(
  spl28_120 | ~spl28_123),
  inference(avatar_split_clause,[],[f1635,f1650,f1644])).
fof(f1635,plain,(
  ~coll(sK20,sK23,sK23) | midp(sK20,sK23,sK23)),
  inference(resolution,[],[f1369,f427])).
fof(f1369,plain,(
  cong(sK20,sK23,sK20,sK23)),
  inference(resolution,[],[f696,f526])).
fof(f526,plain,(
  cong(sK20,sK23,sK24,sK21)),
  inference(resolution,[],[f383,f475])).
fof(f1016,plain,(
  spl28_64 | spl28_66),
  inference(avatar_split_clause,[],[f1009,f1014,f1011])).
fof(f1009,plain,(
  ( ! [X6,X10,X8,X7,X11,X9] : (~cong(X6,X7,X7,X8) | cong(X9,X10,X10,X11) | ~midp(X6,X7,X8) | ~midp(X9,X10,X11)) )),
  inference(resolution,[],[f435,f430])).
% SZS output end Proof for C103
% ------------------------------
% Version: Vampire 4.4.0 (commit 7916d27 on 2019-08-23 08:50:16 +0100)
% Termination reason: Refutation

% Memory used [KB]: 11513
% Time elapsed: 0.249 s
% ------------------------------
% ------------------------------
