% Refutation found. Thanks to Tanya!
% SZS status Theorem for 61
% SZS output start Proof for 61
18. ! [X0,X1,X2,X3,X8,X9,X10,X11] : (eqangle(X0,X1,X2,X3,X8,X9,X10,X11) => eqangle(X1,X0,X2,X3,X8,X9,X10,X11)) [input]
19. ! [X0,X1,X2,X3,X8,X9,X10,X11] : (eqangle(X0,X1,X2,X3,X8,X9,X10,X11) => eqangle(X2,X3,X0,X1,X10,X11,X8,X9)) [input]
21. ! [X0,X1,X2,X3,X8,X9,X10,X11] : (eqangle(X0,X1,X2,X3,X8,X9,X10,X11) => eqangle(X0,X1,X8,X9,X2,X3,X10,X11)) [input]
22. ! [X0,X1,X2,X3,X8,X9,X10,X11,X4,X5,X12,X13] : ((eqangle(X8,X9,X10,X11,X4,X5,X12,X13) & eqangle(X0,X1,X2,X3,X8,X9,X10,X11)) => eqangle(X0,X1,X2,X3,X4,X5,X12,X13)) [input]
40. ! [X0,X1,X2,X3,X8,X9] : (para(X0,X1,X2,X3) => eqangle(X0,X1,X8,X9,X2,X3,X8,X9)) [input]
44. ! [X0,X1,X2,X4,X5] : ((midp(X5,X0,X2) & midp(X4,X0,X1)) => para(X4,X5,X1,X2)) [input]
94. ! [X0,X1,X2,X3,X4,X5,X12,X18,X19,X20] : (circle(X5,X0,X12,X20) & coll(X12,X1,X2) & circle(X5,X0,X18,X19) & coll(X5,X3,X4) & midp(X4,X1,X0) & circle(X3,X0,X1,X2)) => (eqangle(X0,X3,X3,X5,X0,X2,X2,X12) & (eqangle(X3,X0,X0,X5,X0,X12,X12,X2) | eqangle(X0,X3,X3,X5,X2,X0,X0,X12)) & (eqangle(X3,X0,X0,X5,X0,X12,X12,X2) | eqangle(X0,X3,X3,X5,X0,X12,X12,X2)) & (eqangle(X3,X0,X0,X5,X0,X2,X2,X12) | eqangle(X0,X3,X3,X5,X2,X0,X0,X12)) & (eqangle(X3,X0,X0,X5,X0,X2,X2,X12) | eqangle(X0,X3,X3,X5,X0,X12,X12,X2)) & (eqangle(X3,X0,X0,X5,X2,X0,X0,X12) | eqangle(X0,X3,X3,X5,X0,X2,X2,X12)) & eqangle(X3,X0,X0,X5,X2,X0,X0,X12)) [input]
95. ~! [X12,X2,X0,X5,X3] : (! [X0,X1,X2,X3,X4,X5,X12,X18,X19,X20] : (circle(X5,X0,X12,X20) & coll(X12,X1,X2) & circle(X5,X0,X18,X19) & coll(X5,X3,X4) & midp(X4,X1,X0) & circle(X3,X0,X1,X2)) => (eqangle(X0,X3,X3,X5,X0,X2,X2,X12) & (eqangle(X3,X0,X0,X5,X0,X12,X12,X2) | eqangle(X0,X3,X3,X5,X2,X0,X0,X12)) & (eqangle(X3,X0,X0,X5,X0,X12,X12,X2) | eqangle(X0,X3,X3,X5,X0,X12,X12,X2)) & (eqangle(X3,X0,X0,X5,X0,X2,X2,X12) | eqangle(X0,X3,X3,X5,X2,X0,X0,X12)) & (eqangle(X3,X0,X0,X5,X0,X2,X2,X12) | eqangle(X0,X3,X3,X5,X0,X12,X12,X2)) & (eqangle(X3,X0,X0,X5,X2,X0,X0,X12) | eqangle(X0,X3,X3,X5,X0,X2,X2,X12)) & eqangle(X3,X0,X0,X5,X2,X0,X0,X12))) [negated conjecture 94]
99. ! [X0,X1,X2,X3,X4,X5,X6,X7] : (eqangle(X0,X1,X2,X3,X4,X5,X6,X7) => eqangle(X1,X0,X2,X3,X4,X5,X6,X7)) [rectify 18]
100. ! [X0,X1,X2,X3,X4,X5,X6,X7] : (eqangle(X0,X1,X2,X3,X4,X5,X6,X7) => eqangle(X2,X3,X0,X1,X6,X7,X4,X5)) [rectify 19]
102. ! [X0,X1,X2,X3,X4,X5,X6,X7] : (eqangle(X0,X1,X2,X3,X4,X5,X6,X7) => eqangle(X0,X1,X4,X5,X2,X3,X6,X7)) [rectify 21]
103. ! [X0,X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11] : ((eqangle(X4,X5,X6,X7,X8,X9,X10,X11) & eqangle(X0,X1,X2,X3,X4,X5,X6,X7)) => eqangle(X0,X1,X2,X3,X8,X9,X10,X11)) [rectify 22]
118. ! [X0,X1,X2,X3,X4,X5] : (para(X0,X1,X2,X3) => eqangle(X0,X1,X4,X5,X2,X3,X4,X5)) [rectify 40]
122. ! [X0,X1,X2,X3,X4] : ((midp(X4,X0,X2) & midp(X3,X0,X1)) => para(X3,X4,X1,X2)) [rectify 44]
165. ~! [X0,X1,X2,X3,X4] : (! [X5,X6,X7,X8,X9,X10,X11,X12,X13,X14] : (circle(X10,X5,X11,X14) & coll(X11,X6,X7) & circle(X10,X5,X12,X13) & coll(X10,X8,X9) & midp(X9,X6,X5) & circle(X8,X5,X6,X7)) => (eqangle(X2,X4,X4,X3,X2,X1,X1,X0) & (eqangle(X4,X2,X2,X3,X2,X0,X0,X1) | eqangle(X2,X4,X4,X3,X1,X2,X2,X0)) & (eqangle(X4,X2,X2,X3,X2,X0,X0,X1) | eqangle(X2,X4,X4,X3,X2,X0,X0,X1)) & (eqangle(X4,X2,X2,X3,X2,X1,X1,X0) | eqangle(X2,X4,X4,X3,X1,X2,X2,X0)) & (eqangle(X4,X2,X2,X3,X2,X1,X1,X0) | eqangle(X2,X4,X4,X3,X2,X0,X0,X1)) & (eqangle(X4,X2,X2,X3,X1,X2,X2,X0) | eqangle(X2,X4,X4,X3,X2,X1,X1,X0)) & eqangle(X4,X2,X2,X3,X1,X2,X2,X0))) [rectify 95]
190. ! [X0,X1,X2,X3,X4,X5,X6,X7] : (eqangle(X1,X0,X2,X3,X4,X5,X6,X7) | ~eqangle(X0,X1,X2,X3,X4,X5,X6,X7)) [ennf transformation 99]
191. ! [X0,X1,X2,X3,X4,X5,X6,X7] : (eqangle(X2,X3,X0,X1,X6,X7,X4,X5) | ~eqangle(X0,X1,X2,X3,X4,X5,X6,X7)) [ennf transformation 100]
193. ! [X0,X1,X2,X3,X4,X5,X6,X7] : (eqangle(X0,X1,X4,X5,X2,X3,X6,X7) | ~eqangle(X0,X1,X2,X3,X4,X5,X6,X7)) [ennf transformation 102]
194. ! [X0,X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11] : (eqangle(X0,X1,X2,X3,X8,X9,X10,X11) | (~eqangle(X4,X5,X6,X7,X8,X9,X10,X11) | ~eqangle(X0,X1,X2,X3,X4,X5,X6,X7))) [ennf transformation 103]
195. ! [X0,X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11] : (eqangle(X0,X1,X2,X3,X8,X9,X10,X11) | ~eqangle(X4,X5,X6,X7,X8,X9,X10,X11) | ~eqangle(X0,X1,X2,X3,X4,X5,X6,X7)) [flattening 194]
217. ! [X0,X1,X2,X3,X4,X5] : (eqangle(X0,X1,X4,X5,X2,X3,X4,X5) | ~para(X0,X1,X2,X3)) [ennf transformation 118]
223. ! [X0,X1,X2,X3,X4] : (para(X3,X4,X1,X2) | (~midp(X4,X0,X2) | ~midp(X3,X0,X1))) [ennf transformation 122]
224. ! [X0,X1,X2,X3,X4] : (para(X3,X4,X1,X2) | ~midp(X4,X0,X2) | ~midp(X3,X0,X1)) [flattening 223]
315. ? [X0,X1,X2,X3,X4] : ((~eqangle(X2,X4,X4,X3,X2,X1,X1,X0) | (~eqangle(X4,X2,X2,X3,X2,X0,X0,X1) & ~eqangle(X2,X4,X4,X3,X1,X2,X2,X0)) | (~eqangle(X4,X2,X2,X3,X2,X0,X0,X1) & ~eqangle(X2,X4,X4,X3,X2,X0,X0,X1)) | (~eqangle(X4,X2,X2,X3,X2,X1,X1,X0) & ~eqangle(X2,X4,X4,X3,X1,X2,X2,X0)) | (~eqangle(X4,X2,X2,X3,X2,X1,X1,X0) & ~eqangle(X2,X4,X4,X3,X2,X0,X0,X1)) | (~eqangle(X4,X2,X2,X3,X1,X2,X2,X0) & ~eqangle(X2,X4,X4,X3,X2,X1,X1,X0)) | ~eqangle(X4,X2,X2,X3,X1,X2,X2,X0)) & ! [X5,X6,X7,X8,X9,X10,X11,X12,X13,X14] : (circle(X10,X5,X11,X14) & coll(X11,X6,X7) & circle(X10,X5,X12,X13) & coll(X10,X8,X9) & midp(X9,X6,X5) & circle(X8,X5,X6,X7))) [ennf transformation 165]
316. ! [X0,X2,X1,X3,X4] : ((~eqangle(X4,X2,X2,X3,X1,X2,X2,X0) & ~eqangle(X2,X4,X4,X3,X2,X1,X1,X0)) | ~sP0(X0,X2,X1,X3,X4)) [predicate definition introduction]
317. ! [X0,X1,X2,X3,X4] : ((~eqangle(X4,X2,X2,X3,X2,X1,X1,X0) & ~eqangle(X2,X4,X4,X3,X2,X0,X0,X1)) | ~sP1(X0,X1,X2,X3,X4)) [predicate definition introduction]
318. ? [X0,X1,X2,X3,X4] : ((~eqangle(X2,X4,X4,X3,X2,X1,X1,X0) | (~eqangle(X4,X2,X2,X3,X2,X0,X0,X1) & ~eqangle(X2,X4,X4,X3,X1,X2,X2,X0)) | (~eqangle(X4,X2,X2,X3,X2,X0,X0,X1) & ~eqangle(X2,X4,X4,X3,X2,X0,X0,X1)) | (~eqangle(X4,X2,X2,X3,X2,X1,X1,X0) & ~eqangle(X2,X4,X4,X3,X1,X2,X2,X0)) | sP1(X0,X1,X2,X3,X4) | sP0(X0,X2,X1,X3,X4) | ~eqangle(X4,X2,X2,X3,X1,X2,X2,X0)) & ! [X5,X6,X7,X8,X9,X10,X11,X12,X13,X14] : (circle(X10,X5,X11,X14) & coll(X11,X6,X7) & circle(X10,X5,X12,X13) & coll(X10,X8,X9) & midp(X9,X6,X5) & circle(X8,X5,X6,X7))) [definition folding 315,317,316]
355. ! [X0,X1,X2,X3,X4] : ((~eqangle(X4,X2,X2,X3,X2,X1,X1,X0) & ~eqangle(X2,X4,X4,X3,X2,X0,X0,X1)) | ~sP1(X0,X1,X2,X3,X4)) [nnf transformation 317]
356. ! [X0,X2,X1,X3,X4] : ((~eqangle(X4,X2,X2,X3,X1,X2,X2,X0) & ~eqangle(X2,X4,X4,X3,X2,X1,X1,X0)) | ~sP0(X0,X2,X1,X3,X4)) [nnf transformation 316]
357. ! [X0,X1,X2,X3,X4] : ((~eqangle(X4,X1,X1,X3,X2,X1,X1,X0) & ~eqangle(X1,X4,X4,X3,X1,X2,X2,X0)) | ~sP0(X0,X1,X2,X3,X4)) [rectify 356]
358. ? [X0,X1,X2,X3,X4] : ((~eqangle(X2,X4,X4,X3,X2,X1,X1,X0) | (~eqangle(X4,X2,X2,X3,X2,X0,X0,X1) & ~eqangle(X2,X4,X4,X3,X1,X2,X2,X0)) | (~eqangle(X4,X2,X2,X3,X2,X0,X0,X1) & ~eqangle(X2,X4,X4,X3,X2,X0,X0,X1)) | (~eqangle(X4,X2,X2,X3,X2,X1,X1,X0) & ~eqangle(X2,X4,X4,X3,X1,X2,X2,X0)) | sP1(X0,X1,X2,X3,X4) | sP0(X0,X2,X1,X3,X4) | ~eqangle(X4,X2,X2,X3,X1,X2,X2,X0)) & ! [X5,X6,X7,X8,X9,X10,X11,X12,X13,X14] : (circle(X10,X5,X11,X14) & coll(X11,X6,X7) & circle(X10,X5,X12,X13) & coll(X10,X8,X9) & midp(X9,X6,X5) & circle(X8,X5,X6,X7))) => ((~eqangle(sK24,sK26,sK26,sK25,sK24,sK23,sK23,sK22) | (~eqangle(sK26,sK24,sK24,sK25,sK24,sK22,sK22,sK23) & ~eqangle(sK24,sK26,sK26,sK25,sK23,sK24,sK24,sK22)) | (~eqangle(sK26,sK24,sK24,sK25,sK24,sK22,sK22,sK23) & ~eqangle(sK24,sK26,sK26,sK25,sK24,sK22,sK22,sK23)) | (~eqangle(sK26,sK24,sK24,sK25,sK24,sK23,sK23,sK22) & ~eqangle(sK24,sK26,sK26,sK25,sK23,sK24,sK24,sK22)) | sP1(sK22,sK23,sK24,sK25,sK26) | sP0(sK22,sK24,sK23,sK25,sK26) | ~eqangle(sK26,sK24,sK24,sK25,sK23,sK24,sK24,sK22)) & ! [X5,X6,X7,X8,X9,X10,X11,X12,X13,X14] : (circle(X10,X5,X11,X14) & coll(X11,X6,X7) & circle(X10,X5,X12,X13) & coll(X10,X8,X9) & midp(X9,X6,X5) & circle(X8,X5,X6,X7))) [choice axiom]
359. (~eqangle(sK24,sK26,sK26,sK25,sK24,sK23,sK23,sK22) | (~eqangle(sK26,sK24,sK24,sK25,sK24,sK22,sK22,sK23) & ~eqangle(sK24,sK26,sK26,sK25,sK23,sK24,sK24,sK22)) | (~eqangle(sK26,sK24,sK24,sK25,sK24,sK22,sK22,sK23) & ~eqangle(sK24,sK26,sK26,sK25,sK24,sK22,sK22,sK23)) | (~eqangle(sK26,sK24,sK24,sK25,sK24,sK23,sK23,sK22) & ~eqangle(sK24,sK26,sK26,sK25,sK23,sK24,sK24,sK22)) | sP1(sK22,sK23,sK24,sK25,sK26) | sP0(sK22,sK24,sK23,sK25,sK26) | ~eqangle(sK26,sK24,sK24,sK25,sK23,sK24,sK24,sK22)) & ! [X5,X6,X7,X8,X9,X10,X11,X12,X13,X14] : (circle(X10,X5,X11,X14) & coll(X11,X6,X7) & circle(X10,X5,X12,X13) & coll(X10,X8,X9) & midp(X9,X6,X5) & circle(X8,X5,X6,X7)) [skolemisation 318,358]
377. ~eqangle(X0,X1,X2,X3,X4,X5,X6,X7) | eqangle(X1,X0,X2,X3,X4,X5,X6,X7) [cnf transformation 190]
378. ~eqangle(X0,X1,X2,X3,X4,X5,X6,X7) | eqangle(X2,X3,X0,X1,X6,X7,X4,X5) [cnf transformation 191]
380. ~eqangle(X0,X1,X2,X3,X4,X5,X6,X7) | eqangle(X0,X1,X4,X5,X2,X3,X6,X7) [cnf transformation 193]
381. ~eqangle(X4,X5,X6,X7,X8,X9,X10,X11) | eqangle(X0,X1,X2,X3,X8,X9,X10,X11) | ~eqangle(X0,X1,X2,X3,X4,X5,X6,X7) [cnf transformation 195]
399. eqangle(X0,X1,X4,X5,X2,X3,X4,X5) | ~para(X0,X1,X2,X3) [cnf transformation 217]
404. ~midp(X4,X0,X2) | para(X3,X4,X1,X2) | ~midp(X3,X0,X1) [cnf transformation 224]
473. ~eqangle(X2,X4,X4,X3,X2,X0,X0,X1) | ~sP1(X0,X1,X2,X3,X4) [cnf transformation 355]
476. ~eqangle(X4,X1,X1,X3,X2,X1,X1,X0) | ~sP0(X0,X1,X2,X3,X4) [cnf transformation 357]
478. midp(X9,X6,X5) [cnf transformation 359]
483. ~eqangle(sK24,sK26,sK26,sK25,sK24,sK23,sK23,sK22) | ~eqangle(sK24,sK26,sK26,sK25,sK23,sK24,sK24,sK22) | ~eqangle(sK24,sK26,sK26,sK25,sK24,sK22,sK22,sK23) | ~eqangle(sK24,sK26,sK26,sK25,sK23,sK24,sK24,sK22) | sP1(sK22,sK23,sK24,sK25,sK26) | sP0(sK22,sK24,sK23,sK25,sK26) | ~eqangle(sK26,sK24,sK24,sK25,sK23,sK24,sK24,sK22) [cnf transformation 359]
491. ~eqangle(sK24,sK26,sK26,sK25,sK24,sK23,sK23,sK22) | ~eqangle(sK24,sK26,sK26,sK25,sK23,sK24,sK24,sK22) | ~eqangle(sK24,sK26,sK26,sK25,sK24,sK22,sK22,sK23) | sP1(sK22,sK23,sK24,sK25,sK26) | sP0(sK22,sK24,sK23,sK25,sK26) | ~eqangle(sK26,sK24,sK24,sK25,sK23,sK24,sK24,sK22) [duplicate literal removal 483]
501. 1 <=> ~eqangle(sK26,sK24,sK24,sK25,sK23,sK24,sK24,sK22) [avatar definition]
502. ~eqangle(sK26,sK24,sK24,sK25,sK23,sK24,sK24,sK22) <- (1) [avatar component clause 501]
519. 7 <=> ~eqangle(sK24,sK26,sK26,sK25,sK24,sK23,sK23,sK22) [avatar definition]
520. ~eqangle(sK24,sK26,sK26,sK25,sK24,sK23,sK23,sK22) <- (7) [avatar component clause 519]
533. 11 <=> ~eqangle(sK24,sK26,sK26,sK25,sK23,sK24,sK24,sK22) [avatar definition]
534. ~eqangle(sK24,sK26,sK26,sK25,sK23,sK24,sK24,sK22) <- (11) [avatar component clause 533]
542. 13 <=> ~eqangle(sK24,sK26,sK26,sK25,sK24,sK22,sK22,sK23) [avatar definition]
543. ~eqangle(sK24,sK26,sK26,sK25,sK24,sK22,sK22,sK23) <- (13) [avatar component clause 542]
556. ~eqangle(sK24,sK26,sK26,sK25,sK24,sK23,sK23,sK22) | ~eqangle(sK24,sK26,sK26,sK25,sK23,sK24,sK24,sK22) | ~eqangle(sK24,sK26,sK26,sK25,sK24,sK22,sK22,sK23) | sP1(sK22,sK23,sK24,sK25,sK26) | ~eqangle(sK26,sK24,sK24,sK25,sK23,sK24,sK24,sK22) [subsumption resolution 491,476]
557. ~eqangle(sK24,sK26,sK26,sK25,sK24,sK23,sK23,sK22) | ~eqangle(sK24,sK26,sK26,sK25,sK23,sK24,sK24,sK22) | ~eqangle(sK24,sK26,sK26,sK25,sK24,sK22,sK22,sK23) | ~eqangle(sK26,sK24,sK24,sK25,sK23,sK24,sK24,sK22) [subsumption resolution 556,473]
558. ~1 | ~13 | ~11 | ~7 [avatar split clause 557,519,533,542,501]
575. para(X0,X1,X2,X3) | ~midp(X0,X4,X2) [resolution 404,478]
576. para(X0,X1,X2,X3) [subsumption resolution 575,478]
700. eqangle(X0,X1,X2,X3,X4,X5,X2,X3) | ~para(X1,X0,X4,X5) [resolution 377,399]
704. eqangle(X0,X1,X2,X3,X4,X5,X2,X3) [subsumption resolution 700,576]
712. eqangle(X0,X1,X2,X3,X0,X1,X4,X5) [resolution 378,704]
740. eqangle(X0,X1,X2,X3,X4,X5,X4,X5) [resolution 380,704]
791. eqangle(X2,X3,X2,X3,X4,X5,X6,X7) [resolution 712,380]
1748. eqangle(X222,X223,X224,X225,X226,X227,X228,X229) | ~eqangle(X222,X223,X224,X225,X230,X231,X230,X231) [resolution 381,791]
1759. eqangle(X222,X223,X224,X225,X226,X227,X228,X229) [subsumption resolution 1748,740]
2754. $false <- (1) [resolution 1759,502]
2782. 1 [avatar contradiction clause 2754]
2783. $false <- (7) [subsumption resolution 520,1759]
2784. 7 [avatar contradiction clause 2783]
2785. $false <- (11) [subsumption resolution 534,1759]
2786. 11 [avatar contradiction clause 2785]
2789. $false <- (13) [subsumption resolution 543,1759]
2790. 13 [avatar contradiction clause 2789]
2795. $false [avatar sat refutation 558,2782,2784,2786,2790]
% SZS output end Proof for 61
% ------------------------------
% Version: Vampire 4.4.0 (commit 7916d27 on 2019-08-23 08:50:16 +0100)
% Termination reason: Refutation

% Memory used [KB]: 5756
% Time elapsed: 0.044 s
% ------------------------------
% ------------------------------
