% Refutation found. Thanks to Tanya!
% SZS status Theorem for 64
% SZS output start Proof for 64
9. ! [X0,X1,X2,X3,X4,X5] : ((perp(X2,X3,X4,X5) & perp(X0,X1,X2,X3)) => para(X0,X1,X4,X5)) [input]
10. ! [X0,X1,X2,X3,X4,X5] : ((perp(X2,X3,X4,X5) & para(X0,X1,X2,X3)) => perp(X0,X1,X4,X5)) [input]
53. ! [X0,X1,X2,X7] : ((coll(X7,X0,X2) & circle(X7,X0,X1,X2)) => perp(X0,X1,X1,X2)) [input]
94. ! [X0,X1,X4,X5,X3,X6,X18,X19,X20,X21,X22,X23,X24,X25,X26,X27,X28] : (circle(X0,X4,X19,X28) & coll(X19,X5,X6) & coll(X18,X4,X6) & circle(X0,X4,X18,X27) & circle(X1,X4,X6,X26) & coll(X3,X4,X5) & coll(X3,X0,X1) & circle(X0,X4,X5,X25) & circle(X1,X4,X5,X24) & circle(X0,X4,X22,X23) & circle(X1,X4,X20,X21)) => perp(X18,X19,X6,X1) [input]
95. ~! [X1,X6,X19,X18] : (! [X0,X1,X4,X5,X3,X6,X18,X19,X20,X21,X22,X23,X24,X25,X26,X27,X28] : (circle(X0,X4,X19,X28) & coll(X19,X5,X6) & coll(X18,X4,X6) & circle(X0,X4,X18,X27) & circle(X1,X4,X6,X26) & coll(X3,X4,X5) & coll(X3,X0,X1) & circle(X0,X4,X5,X25) & circle(X1,X4,X5,X24) & circle(X0,X4,X22,X23) & circle(X1,X4,X20,X21)) => perp(X18,X19,X6,X1)) [negated conjecture 94]
131. ! [X0,X1,X2,X3] : ((coll(X3,X0,X2) & circle(X3,X0,X1,X2)) => perp(X0,X1,X1,X2)) [rectify 53]
165. ~! [X0,X1,X2,X3] : (! [X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20] : (circle(X4,X6,X11,X20) & coll(X11,X7,X9) & coll(X10,X6,X9) & circle(X4,X6,X10,X19) & circle(X5,X6,X9,X18) & coll(X8,X6,X7) & coll(X8,X4,X5) & circle(X4,X6,X7,X17) & circle(X5,X6,X7,X16) & circle(X4,X6,X14,X15) & circle(X5,X6,X12,X13)) => perp(X3,X2,X1,X0)) [rectify 95]
176. ! [X0,X1,X2,X3,X4,X5] : (para(X0,X1,X4,X5) | (~perp(X2,X3,X4,X5) | ~perp(X0,X1,X2,X3))) [ennf transformation 9]
177. ! [X0,X1,X2,X3,X4,X5] : (para(X0,X1,X4,X5) | ~perp(X2,X3,X4,X5) | ~perp(X0,X1,X2,X3)) [flattening 176]
178. ! [X0,X1,X2,X3,X4,X5] : (perp(X0,X1,X4,X5) | (~perp(X2,X3,X4,X5) | ~para(X0,X1,X2,X3))) [ennf transformation 10]
179. ! [X0,X1,X2,X3,X4,X5] : (perp(X0,X1,X4,X5) | ~perp(X2,X3,X4,X5) | ~para(X0,X1,X2,X3)) [flattening 178]
240. ! [X0,X1,X2,X3] : (perp(X0,X1,X1,X2) | (~coll(X3,X0,X2) | ~circle(X3,X0,X1,X2))) [ennf transformation 131]
241. ! [X0,X1,X2,X3] : (perp(X0,X1,X1,X2) | ~coll(X3,X0,X2) | ~circle(X3,X0,X1,X2)) [flattening 240]
315. ? [X0,X1,X2,X3] : (~perp(X3,X2,X1,X0) & ! [X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20] : (circle(X4,X6,X11,X20) & coll(X11,X7,X9) & coll(X10,X6,X9) & circle(X4,X6,X10,X19) & circle(X5,X6,X9,X18) & coll(X8,X6,X7) & coll(X8,X4,X5) & circle(X4,X6,X7,X17) & circle(X5,X6,X7,X16) & circle(X4,X6,X14,X15) & circle(X5,X6,X12,X13))) [ennf transformation 165]
352. ? [X0,X1,X2,X3] : (~perp(X3,X2,X1,X0) & ! [X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20] : (circle(X4,X6,X11,X20) & coll(X11,X7,X9) & coll(X10,X6,X9) & circle(X4,X6,X10,X19) & circle(X5,X6,X9,X18) & coll(X8,X6,X7) & coll(X8,X4,X5) & circle(X4,X6,X7,X17) & circle(X5,X6,X7,X16) & circle(X4,X6,X14,X15) & circle(X5,X6,X12,X13))) => (~perp(sK23,sK22,sK21,sK20) & ! [X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20] : (circle(X4,X6,X11,X20) & coll(X11,X7,X9) & coll(X10,X6,X9) & circle(X4,X6,X10,X19) & circle(X5,X6,X9,X18) & coll(X8,X6,X7) & coll(X8,X4,X5) & circle(X4,X6,X7,X17) & circle(X5,X6,X7,X16) & circle(X4,X6,X14,X15) & circle(X5,X6,X12,X13))) [choice axiom]
353. ~perp(sK23,sK22,sK21,sK20) & ! [X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20] : (circle(X4,X6,X11,X20) & coll(X11,X7,X9) & coll(X10,X6,X9) & circle(X4,X6,X10,X19) & circle(X5,X6,X9,X18) & coll(X8,X6,X7) & coll(X8,X4,X5) & circle(X4,X6,X7,X17) & circle(X5,X6,X7,X16) & circle(X4,X6,X14,X15) & circle(X5,X6,X12,X13)) [skolemisation 315,352]
362. ~perp(X2,X3,X4,X5) | para(X0,X1,X4,X5) | ~perp(X0,X1,X2,X3) [cnf transformation 177]
363. ~perp(X2,X3,X4,X5) | perp(X0,X1,X4,X5) | ~para(X0,X1,X2,X3) [cnf transformation 179]
407. ~circle(X3,X0,X1,X2) | ~coll(X3,X0,X2) | perp(X0,X1,X1,X2) [cnf transformation 241]
476. coll(X11,X7,X9) [cnf transformation 353]
477. circle(X4,X6,X11,X20) [cnf transformation 353]
478. ~perp(sK23,sK22,sK21,sK20) [cnf transformation 353]
521. ~coll(X0,X1,X2) | perp(X1,X3,X3,X2) [resolution 407,477]
522. perp(X1,X3,X3,X2) [subsumption resolution 521,476]
544. ~perp(X0,X1,X4,X2) | para(X0,X1,X2,X3) [resolution 362,522]
549. ~para(X0,X1,X4,X2) | perp(X0,X1,X2,X3) [resolution 363,522]
634. para(X0,X1,X2,X3) [resolution 544,522]
794. perp(X24,X25,X26,X27) [resolution 549,634]
798. $false [resolution 794,478]
% SZS output end Proof for 64
% ------------------------------
% Version: Vampire 4.4.0 (commit 7916d27 on 2019-08-23 08:50:16 +0100)
% Termination reason: Refutation

% Memory used [KB]: 5245
% Time elapsed: 0.030 s
% ------------------------------
% ------------------------------
