% Refutation found. Thanks to Tanya!
% SZS status Theorem for 67
% SZS output start Proof for 67
fof(f43847,plain,(
  $false),
  inference(avatar_sat_refutation,[],[f2516,f3740,f3747,f5724,f5814,f6206,f12427,f14449,f14482,f43824])).
fof(f43824,plain,(
  ~spl31_98 | ~spl31_244),
  inference(avatar_contradiction_clause,[],[f43819])).
fof(f43819,plain,(
  $false | (~spl31_98 | ~spl31_244)),
  inference(resolution,[],[f20788,f487])).
fof(f487,plain,(
  ~perp(sK29,sK24,sK24,sK30)),
  inference(cnf_transformation,[],[f359])).
fof(f359,plain,(
  ~perp(sK29,sK24,sK24,sK30) & perp(sK26,sK25,sK20,sK30) & midp(sK20,sK26,sK25) & perp(sK24,sK25,sK21,sK30) & midp(sK21,sK24,sK25) & perp(sK24,sK26,sK22,sK30) & midp(sK22,sK24,sK26) & midp(sK29,sK28,sK27) & coll(sK28,sK24,sK26) & coll(sK27,sK25,sK26) & perp(sK25,sK26,sK25,sK28) & perp(sK24,sK26,sK24,sK27) & eqangle(sK26,sK25,sK25,sK23,sK26,sK25,sK25,sK24) & eqangle(sK26,sK24,sK24,sK25,sK26,sK24,sK24,sK23) & eqangle(sK26,sK23,sK23,sK24,sK26,sK23,sK23,sK25)),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK20,sK21,sK22,sK23,sK24,sK25,sK26,sK27,sK28,sK29,sK30])],[f321,f358])).
fof(f358,plain,(
  ? [X0,X1,X2,X3,X4,X5,X6,X7,X8,X9,X10] : (~perp(X9,X4,X4,X10) & perp(X6,X5,X0,X10) & midp(X0,X6,X5) & perp(X4,X5,X1,X10) & midp(X1,X4,X5) & perp(X4,X6,X2,X10) & midp(X2,X4,X6) & midp(X9,X8,X7) & coll(X8,X4,X6) & coll(X7,X5,X6) & perp(X5,X6,X5,X8) & perp(X4,X6,X4,X7) & eqangle(X6,X5,X5,X3,X6,X5,X5,X4) & eqangle(X6,X4,X4,X5,X6,X4,X4,X3) & eqangle(X6,X3,X3,X4,X6,X3,X3,X5)) => (~perp(sK29,sK24,sK24,sK30) & perp(sK26,sK25,sK20,sK30) & midp(sK20,sK26,sK25) & perp(sK24,sK25,sK21,sK30) & midp(sK21,sK24,sK25) & perp(sK24,sK26,sK22,sK30) & midp(sK22,sK24,sK26) & midp(sK29,sK28,sK27) & coll(sK28,sK24,sK26) & coll(sK27,sK25,sK26) & perp(sK25,sK26,sK25,sK28) & perp(sK24,sK26,sK24,sK27) & eqangle(sK26,sK25,sK25,sK23,sK26,sK25,sK25,sK24) & eqangle(sK26,sK24,sK24,sK25,sK26,sK24,sK24,sK23) & eqangle(sK26,sK23,sK23,sK24,sK26,sK23,sK23,sK25))),
  introduced(choice_axiom,[])).
fof(f321,plain,(
  ? [X0,X1,X2,X3,X4,X5,X6,X7,X8,X9,X10] : (~perp(X9,X4,X4,X10) & perp(X6,X5,X0,X10) & midp(X0,X6,X5) & perp(X4,X5,X1,X10) & midp(X1,X4,X5) & perp(X4,X6,X2,X10) & midp(X2,X4,X6) & midp(X9,X8,X7) & coll(X8,X4,X6) & coll(X7,X5,X6) & perp(X5,X6,X5,X8) & perp(X4,X6,X4,X7) & eqangle(X6,X5,X5,X3,X6,X5,X5,X4) & eqangle(X6,X4,X4,X5,X6,X4,X4,X3) & eqangle(X6,X3,X3,X4,X6,X3,X3,X5))),
  inference(flattening,[],[f320])).
fof(f320,plain,(
  ? [X0,X1,X2,X3,X4,X5,X6,X7,X8,X9,X10] : (~perp(X9,X4,X4,X10) & (perp(X6,X5,X0,X10) & midp(X0,X6,X5) & perp(X4,X5,X1,X10) & midp(X1,X4,X5) & perp(X4,X6,X2,X10) & midp(X2,X4,X6) & midp(X9,X8,X7) & coll(X8,X4,X6) & coll(X7,X5,X6) & perp(X5,X6,X5,X8) & perp(X4,X6,X4,X7) & eqangle(X6,X5,X5,X3,X6,X5,X5,X4) & eqangle(X6,X4,X4,X5,X6,X4,X4,X3) & eqangle(X6,X3,X3,X4,X6,X3,X3,X5)))),
  inference(ennf_transformation,[],[f168])).
fof(f168,plain,(
  ~! [X0,X1,X2,X3,X4,X5,X6,X7,X8,X9,X10] : ((perp(X6,X5,X0,X10) & midp(X0,X6,X5) & perp(X4,X5,X1,X10) & midp(X1,X4,X5) & perp(X4,X6,X2,X10) & midp(X2,X4,X6) & midp(X9,X8,X7) & coll(X8,X4,X6) & coll(X7,X5,X6) & perp(X5,X6,X5,X8) & perp(X4,X6,X4,X7) & eqangle(X6,X5,X5,X3,X6,X5,X5,X4) & eqangle(X6,X4,X4,X5,X6,X4,X4,X3) & eqangle(X6,X3,X3,X4,X6,X3,X3,X5)) => perp(X9,X4,X4,X10))),
  inference(flattening,[],[f167])).
fof(f167,plain,(
  ~! [X0,X1,X2] : ! [X3,X4,X5,X6,X7,X8,X9,X10] : ((perp(X6,X5,X0,X10) & midp(X0,X6,X5) & perp(X4,X5,X1,X10) & midp(X1,X4,X5) & perp(X4,X6,X2,X10) & midp(X2,X4,X6) & midp(X9,X8,X7) & coll(X8,X4,X6) & coll(X7,X5,X6) & perp(X5,X6,X5,X8) & perp(X4,X6,X4,X7) & eqangle(X6,X5,X5,X3,X6,X5,X5,X4) & eqangle(X6,X4,X4,X5,X6,X4,X4,X3) & eqangle(X6,X3,X3,X4,X6,X3,X3,X5)) => perp(X9,X4,X4,X10))),
  inference(rectify,[],[f96])).
fof(f96,negated_conjecture,(
  ~! [X23,X22,X21] : ! [X6,X1,X2,X18,X19,X20,X0,X7] : ((perp(X18,X2,X23,X7) & midp(X23,X18,X2) & perp(X1,X2,X22,X7) & midp(X22,X1,X2) & perp(X1,X18,X21,X7) & midp(X21,X1,X18) & midp(X0,X20,X19) & coll(X20,X1,X18) & coll(X19,X2,X18) & perp(X2,X18,X2,X20) & perp(X1,X18,X1,X19) & eqangle(X18,X2,X2,X6,X18,X2,X2,X1) & eqangle(X18,X1,X1,X2,X18,X1,X1,X6) & eqangle(X18,X6,X6,X1,X18,X6,X6,X2)) => perp(X0,X1,X1,X7))),
  inference(negated_conjecture,[],[f95])).
fof(f95,conjecture,(
  ( ! [X23,X21,X22] : (! [X6,X1,X2,X18,X19,X20,X0,X7] : ((perp(X18,X2,X23,X7) & midp(X23,X18,X2) & perp(X1,X2,X22,X7) & midp(X22,X1,X2) & perp(X1,X18,X21,X7) & midp(X21,X1,X18) & midp(X0,X20,X19) & coll(X20,X1,X18) & coll(X19,X2,X18) & perp(X2,X18,X2,X20) & perp(X1,X18,X1,X19) & eqangle(X18,X2,X2,X6,X18,X2,X2,X1) & eqangle(X18,X1,X1,X2,X18,X1,X1,X6) & eqangle(X18,X6,X6,X1,X18,X6,X6,X2)) => perp(X0,X1,X1,X7))) )),
  file('./6_GDD_FULL/61-80/67.p',exemplo6GDDFULL618067)).
fof(f20788,plain,(
  ( ! [X30,X31,X29,X32] : (perp(X29,X30,X31,X32)) ) | (~spl31_98 | ~spl31_244)),
  inference(subsumption_resolution,[],[f20754,f20787])).
fof(f20787,plain,(
  ( ! [X28,X26,X27,X25] : (para(X25,X26,X27,X28)) ) | (~spl31_98 | ~spl31_244)),
  inference(subsumption_resolution,[],[f20753,f20752])).
fof(f20752,plain,(
  ( ! [X24,X23] : (perp(X23,X24,sK26,sK25)) ) | (~spl31_98 | ~spl31_244)),
  inference(resolution,[],[f15889,f367])).
fof(f367,plain,(
  ( ! [X2,X0,X3,X1] : (~perp(X0,X1,X2,X3) | perp(X2,X3,X0,X1)) )),
  inference(cnf_transformation,[],[f178])).
fof(f178,plain,(
  ! [X0,X1,X2,X3] : (perp(X2,X3,X0,X1) | ~perp(X0,X1,X2,X3))),
  inference(ennf_transformation,[],[f8])).
fof(f8,axiom,(
  ! [X0,X1,X2,X3] : (perp(X0,X1,X2,X3) => perp(X2,X3,X0,X1))),
  file('./6_GDD_FULL/61-80/67.p',ruleD8)).
fof(f15889,plain,(
  ( ! [X12,X13] : (perp(sK26,sK25,X12,X13)) ) | (~spl31_98 | ~spl31_244)),
  inference(subsumption_resolution,[],[f15886,f14854])).
fof(f14854,plain,(
  ( ! [X6] : (cong(sK26,X6,sK25,X6)) ) | (~spl31_98 | ~spl31_244)),
  inference(subsumption_resolution,[],[f14847,f14623])).
fof(f14623,plain,(
  ( ! [X92] : (midp(X92,sK26,sK24)) ) | ~spl31_244),
  inference(subsumption_resolution,[],[f14598,f14452])).
fof(f14452,plain,(
  ( ! [X30,X29] : (midp(X29,X30,X30)) ) | ~spl31_244),
  inference(subsumption_resolution,[],[f14418,f6321])).
fof(f6321,plain,(
  ( ! [X4,X3] : (coll(X3,X4,X4)) ) | ~spl31_244),
  inference(resolution,[],[f2512,f426])).
fof(f426,plain,(
  ( ! [X2,X0,X1] : (~para(X0,X1,X0,X2) | coll(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f268])).
fof(f268,plain,(
  ! [X0,X1,X2] : (coll(X0,X1,X2) | ~para(X0,X1,X0,X2))),
  inference(ennf_transformation,[],[f67])).
fof(f67,axiom,(
  ! [X0,X1,X2] : (para(X0,X1,X0,X2) => coll(X0,X1,X2))),
  file('./6_GDD_FULL/61-80/67.p',ruleD66)).
fof(f2512,plain,(
  ( ! [X2,X3] : (para(X2,X3,X2,X3)) ) | ~spl31_244),
  inference(avatar_component_clause,[],[f2511])).
fof(f2511,plain,(
  spl31_244 <=> ! [X3,X2] : para(X2,X3,X2,X3)),
  introduced(avatar_definition,[new_symbols(naming,[spl31_244])])).
fof(f14418,plain,(
  ( ! [X30,X29] : (~coll(X29,X30,X30) | midp(X29,X30,X30)) ) | ~spl31_244),
  inference(resolution,[],[f14407,f427])).
fof(f427,plain,(
  ( ! [X2,X0,X1] : (~cong(X0,X1,X0,X2) | ~coll(X0,X1,X2) | midp(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f270])).
fof(f270,plain,(
  ! [X0,X1,X2] : (midp(X0,X1,X2) | ~coll(X0,X1,X2) | ~cong(X0,X1,X0,X2))),
  inference(flattening,[],[f269])).
fof(f269,plain,(
  ! [X0,X1,X2] : (midp(X0,X1,X2) | (~coll(X0,X1,X2) | ~cong(X0,X1,X0,X2)))),
  inference(ennf_transformation,[],[f68])).
fof(f68,axiom,(
  ! [X0,X1,X2] : ((coll(X0,X1,X2) & cong(X0,X1,X0,X2)) => midp(X0,X1,X2))),
  file('./6_GDD_FULL/61-80/67.p',ruleD67)).
fof(f14407,plain,(
  ( ! [X87,X86] : (cong(X86,X87,X86,X87)) ) | ~spl31_244),
  inference(subsumption_resolution,[],[f14406,f10777])).
fof(f10777,plain,(
  ( ! [X12,X10,X11,X9] : (cyclic(X9,X10,X11,X12)) ) | ~spl31_244),
  inference(subsumption_resolution,[],[f10755,f10721])).
fof(f10721,plain,(
  ( ! [X2,X0,X1] : (cyclic(X0,X0,X1,X2)) ) | ~spl31_244),
  inference(resolution,[],[f5047,f2512])).
fof(f5047,plain,(
  ( ! [X68,X66,X67] : (~para(X66,X67,X66,X67) | cyclic(X67,X67,X66,X68)) )),
  inference(subsumption_resolution,[],[f2691,f2692])).
fof(f2692,plain,(
  ( ! [X70,X71,X69] : (~para(X69,X70,X69,X70) | coll(X69,X71,X70) | cyclic(X70,X70,X69,X71)) )),
  inference(resolution,[],[f833,f401])).
fof(f401,plain,(
  ( ! [X2,X0,X3,X1] : (~eqangle(X2,X0,X2,X1,X3,X0,X3,X1) | coll(X2,X3,X0) | cyclic(X0,X1,X2,X3)) )),
  inference(cnf_transformation,[],[f223])).
fof(f223,plain,(
  ! [X0,X1,X2,X3] : (cyclic(X0,X1,X2,X3) | coll(X2,X3,X0) | ~eqangle(X2,X0,X2,X1,X3,X0,X3,X1))),
  inference(flattening,[],[f222])).
fof(f222,plain,(
  ! [X0,X1,X2,X3] : (cyclic(X0,X1,X2,X3) | (coll(X2,X3,X0) | ~eqangle(X2,X0,X2,X1,X3,X0,X3,X1)))),
  inference(ennf_transformation,[],[f121])).
fof(f121,plain,(
  ! [X0,X1,X2,X3] : ((~coll(X2,X3,X0) & eqangle(X2,X0,X2,X1,X3,X0,X3,X1)) => cyclic(X0,X1,X2,X3))),
  inference(rectify,[],[f42])).
fof(f42,axiom,(
  ! [X0,X1,X8,X9] : ((~coll(X8,X9,X0) & eqangle(X8,X0,X8,X1,X9,X0,X9,X1)) => cyclic(X0,X1,X8,X9))),
  file('./6_GDD_FULL/61-80/67.p',ruleD42a)).
fof(f833,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (eqangle(X0,X1,X2,X3,X4,X5,X4,X5) | ~para(X0,X1,X2,X3)) )),
  inference(resolution,[],[f380,f399])).
fof(f399,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (eqangle(X0,X1,X4,X5,X2,X3,X4,X5) | ~para(X0,X1,X2,X3)) )),
  inference(cnf_transformation,[],[f220])).
fof(f220,plain,(
  ! [X0,X1,X2,X3,X4,X5] : (eqangle(X0,X1,X4,X5,X2,X3,X4,X5) | ~para(X0,X1,X2,X3))),
  inference(ennf_transformation,[],[f119])).
fof(f119,plain,(
  ! [X0,X1,X2,X3,X4,X5] : (para(X0,X1,X2,X3) => eqangle(X0,X1,X4,X5,X2,X3,X4,X5))),
  inference(rectify,[],[f40])).
fof(f40,axiom,(
  ! [X0,X1,X2,X3,X8,X9] : (para(X0,X1,X2,X3) => eqangle(X0,X1,X8,X9,X2,X3,X8,X9))),
  file('./6_GDD_FULL/61-80/67.p',ruleD40)).
fof(f380,plain,(
  ( ! [X6,X4,X2,X0,X7,X5,X3,X1] : (~eqangle(X0,X1,X2,X3,X4,X5,X6,X7) | eqangle(X0,X1,X4,X5,X2,X3,X6,X7)) )),
  inference(cnf_transformation,[],[f196])).
fof(f196,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6,X7] : (eqangle(X0,X1,X4,X5,X2,X3,X6,X7) | ~eqangle(X0,X1,X2,X3,X4,X5,X6,X7))),
  inference(ennf_transformation,[],[f103])).
fof(f103,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6,X7] : (eqangle(X0,X1,X2,X3,X4,X5,X6,X7) => eqangle(X0,X1,X4,X5,X2,X3,X6,X7))),
  inference(rectify,[],[f21])).
fof(f21,axiom,(
  ! [X0,X1,X2,X3,X8,X9,X10,X11] : (eqangle(X0,X1,X2,X3,X8,X9,X10,X11) => eqangle(X0,X1,X8,X9,X2,X3,X10,X11))),
  file('./6_GDD_FULL/61-80/67.p',ruleD21)).
fof(f2691,plain,(
  ( ! [X68,X66,X67] : (~para(X66,X67,X66,X67) | ~coll(X66,X68,X67) | cyclic(X67,X67,X66,X68)) )),
  inference(resolution,[],[f833,f402])).
fof(f402,plain,(
  ( ! [X2,X0,X3,X1] : (~eqangle(X2,X0,X2,X1,X3,X0,X3,X1) | ~coll(X2,X3,X1) | cyclic(X0,X1,X2,X3)) )),
  inference(cnf_transformation,[],[f225])).
fof(f225,plain,(
  ! [X0,X1,X2,X3] : (cyclic(X0,X1,X2,X3) | ~coll(X2,X3,X1) | ~eqangle(X2,X0,X2,X1,X3,X0,X3,X1))),
  inference(flattening,[],[f224])).
fof(f224,plain,(
  ! [X0,X1,X2,X3] : (cyclic(X0,X1,X2,X3) | (~coll(X2,X3,X1) | ~eqangle(X2,X0,X2,X1,X3,X0,X3,X1)))),
  inference(ennf_transformation,[],[f122])).
fof(f122,plain,(
  ! [X0,X1,X2,X3] : ((coll(X2,X3,X1) & eqangle(X2,X0,X2,X1,X3,X0,X3,X1)) => cyclic(X0,X1,X2,X3))),
  inference(rectify,[],[f43])).
fof(f43,axiom,(
  ! [X0,X1,X8,X9] : ((coll(X8,X9,X1) & eqangle(X8,X0,X8,X1,X9,X0,X9,X1)) => cyclic(X0,X1,X8,X9))),
  file('./6_GDD_FULL/61-80/67.p',ruleD42b)).
fof(f10755,plain,(
  ( ! [X12,X10,X11,X9] : (cyclic(X9,X10,X11,X12) | ~cyclic(X9,X9,X10,X11)) ) | ~spl31_244),
  inference(resolution,[],[f10721,f376])).
fof(f376,plain,(
  ( ! [X4,X2,X0,X3,X1] : (~cyclic(X0,X1,X2,X4) | cyclic(X1,X2,X3,X4) | ~cyclic(X0,X1,X2,X3)) )),
  inference(cnf_transformation,[],[f192])).
fof(f192,plain,(
  ! [X0,X1,X2,X3,X4] : (cyclic(X1,X2,X3,X4) | ~cyclic(X0,X1,X2,X4) | ~cyclic(X0,X1,X2,X3))),
  inference(flattening,[],[f191])).
fof(f191,plain,(
  ! [X0,X1,X2,X3,X4] : (cyclic(X1,X2,X3,X4) | (~cyclic(X0,X1,X2,X4) | ~cyclic(X0,X1,X2,X3)))),
  inference(ennf_transformation,[],[f17])).
fof(f17,axiom,(
  ! [X0,X1,X2,X3,X4] : ((cyclic(X0,X1,X2,X4) & cyclic(X0,X1,X2,X3)) => cyclic(X1,X2,X3,X4))),
  file('./6_GDD_FULL/61-80/67.p',ruleD17)).
fof(f14406,plain,(
  ( ! [X88,X87,X86] : (cong(X86,X87,X86,X87) | ~cyclic(X86,X87,X88,X87)) ) | ~spl31_244),
  inference(subsumption_resolution,[],[f14405,f10777])).
fof(f14405,plain,(
  ( ! [X88,X87,X86] : (cong(X86,X87,X86,X87) | ~cyclic(X86,X87,X88,X86) | ~cyclic(X86,X87,X88,X87)) ) | ~spl31_244),
  inference(subsumption_resolution,[],[f14391,f10753])).
fof(f10753,plain,(
  ( ! [X4,X5,X3] : (cyclic(X3,X4,X3,X5)) ) | ~spl31_244),
  inference(resolution,[],[f10721,f374])).
fof(f374,plain,(
  ( ! [X2,X0,X3,X1] : (~cyclic(X0,X1,X2,X3) | cyclic(X0,X2,X1,X3)) )),
  inference(cnf_transformation,[],[f189])).
fof(f189,plain,(
  ! [X0,X1,X2,X3] : (cyclic(X0,X2,X1,X3) | ~cyclic(X0,X1,X2,X3))),
  inference(ennf_transformation,[],[f15])).
fof(f15,axiom,(
  ! [X0,X1,X2,X3] : (cyclic(X0,X1,X2,X3) => cyclic(X0,X2,X1,X3))),
  file('./6_GDD_FULL/61-80/67.p',ruleD15)).
fof(f14391,plain,(
  ( ! [X88,X87,X86] : (~cyclic(X86,X87,X86,X88) | cong(X86,X87,X86,X87) | ~cyclic(X86,X87,X88,X86) | ~cyclic(X86,X87,X88,X87)) )),
  inference(duplicate_literal_removal,[],[f14381])).
fof(f14381,plain,(
  ( ! [X88,X87,X86] : (~cyclic(X86,X87,X86,X88) | cong(X86,X87,X86,X87) | ~cyclic(X86,X87,X88,X86) | ~cyclic(X86,X87,X88,X87) | ~cyclic(X86,X87,X88,X86)) )),
  inference(resolution,[],[f2422,f403])).
fof(f403,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (~eqangle(X2,X0,X2,X1,X5,X3,X5,X4) | cong(X0,X1,X3,X4) | ~cyclic(X0,X1,X2,X5) | ~cyclic(X0,X1,X2,X4) | ~cyclic(X0,X1,X2,X3)) )),
  inference(cnf_transformation,[],[f227])).
fof(f227,plain,(
  ! [X0,X1,X2,X3,X4,X5] : (cong(X0,X1,X3,X4) | ~eqangle(X2,X0,X2,X1,X5,X3,X5,X4) | ~cyclic(X0,X1,X2,X5) | ~cyclic(X0,X1,X2,X4) | ~cyclic(X0,X1,X2,X3))),
  inference(flattening,[],[f226])).
fof(f226,plain,(
  ! [X0,X1,X2,X3,X4,X5] : (cong(X0,X1,X3,X4) | (~eqangle(X2,X0,X2,X1,X5,X3,X5,X4) | ~cyclic(X0,X1,X2,X5) | ~cyclic(X0,X1,X2,X4) | ~cyclic(X0,X1,X2,X3)))),
  inference(ennf_transformation,[],[f123])).
fof(f123,plain,(
  ! [X0,X1,X2,X3,X4,X5] : ((eqangle(X2,X0,X2,X1,X5,X3,X5,X4) & cyclic(X0,X1,X2,X5) & cyclic(X0,X1,X2,X4) & cyclic(X0,X1,X2,X3)) => cong(X0,X1,X3,X4))),
  inference(rectify,[],[f44])).
fof(f44,axiom,(
  ! [X0,X1,X2,X8,X9,X14] : ((eqangle(X2,X0,X2,X1,X14,X8,X14,X9) & cyclic(X0,X1,X2,X14) & cyclic(X0,X1,X2,X9) & cyclic(X0,X1,X2,X8)) => cong(X0,X1,X8,X9))),
  file('./6_GDD_FULL/61-80/67.p',ruleD43)).
fof(f2422,plain,(
  ( ! [X14,X12,X13,X11] : (eqangle(X14,X11,X14,X12,X11,X13,X13,X12) | ~cyclic(X11,X12,X13,X14)) )),
  inference(resolution,[],[f795,f379])).
fof(f379,plain,(
  ( ! [X6,X4,X2,X0,X7,X5,X3,X1] : (~eqangle(X0,X1,X2,X3,X4,X5,X6,X7) | eqangle(X4,X5,X6,X7,X0,X1,X2,X3)) )),
  inference(cnf_transformation,[],[f195])).
fof(f195,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6,X7] : (eqangle(X4,X5,X6,X7,X0,X1,X2,X3) | ~eqangle(X0,X1,X2,X3,X4,X5,X6,X7))),
  inference(ennf_transformation,[],[f102])).
fof(f102,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6,X7] : (eqangle(X0,X1,X2,X3,X4,X5,X6,X7) => eqangle(X4,X5,X6,X7,X0,X1,X2,X3))),
  inference(rectify,[],[f20])).
fof(f20,axiom,(
  ! [X0,X1,X2,X3,X8,X9,X10,X11] : (eqangle(X0,X1,X2,X3,X8,X9,X10,X11) => eqangle(X8,X9,X10,X11,X0,X1,X2,X3))),
  file('./6_GDD_FULL/61-80/67.p',ruleD20)).
fof(f795,plain,(
  ( ! [X6,X8,X7,X9] : (eqangle(X6,X7,X7,X8,X9,X6,X9,X8) | ~cyclic(X6,X8,X7,X9)) )),
  inference(resolution,[],[f377,f400])).
fof(f400,plain,(
  ( ! [X2,X0,X3,X1] : (eqangle(X2,X0,X2,X1,X3,X0,X3,X1) | ~cyclic(X0,X1,X2,X3)) )),
  inference(cnf_transformation,[],[f221])).
fof(f221,plain,(
  ! [X0,X1,X2,X3] : (eqangle(X2,X0,X2,X1,X3,X0,X3,X1) | ~cyclic(X0,X1,X2,X3))),
  inference(ennf_transformation,[],[f120])).
fof(f120,plain,(
  ! [X0,X1,X2,X3] : (cyclic(X0,X1,X2,X3) => eqangle(X2,X0,X2,X1,X3,X0,X3,X1))),
  inference(rectify,[],[f41])).
fof(f41,axiom,(
  ! [X0,X1,X8,X9] : (cyclic(X0,X1,X8,X9) => eqangle(X8,X0,X8,X1,X9,X0,X9,X1))),
  file('./6_GDD_FULL/61-80/67.p',ruleD41)).
fof(f377,plain,(
  ( ! [X6,X4,X2,X0,X7,X5,X3,X1] : (~eqangle(X0,X1,X2,X3,X4,X5,X6,X7) | eqangle(X1,X0,X2,X3,X4,X5,X6,X7)) )),
  inference(cnf_transformation,[],[f193])).
fof(f193,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6,X7] : (eqangle(X1,X0,X2,X3,X4,X5,X6,X7) | ~eqangle(X0,X1,X2,X3,X4,X5,X6,X7))),
  inference(ennf_transformation,[],[f100])).
fof(f100,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6,X7] : (eqangle(X0,X1,X2,X3,X4,X5,X6,X7) => eqangle(X1,X0,X2,X3,X4,X5,X6,X7))),
  inference(rectify,[],[f18])).
fof(f18,axiom,(
  ! [X0,X1,X2,X3,X8,X9,X10,X11] : (eqangle(X0,X1,X2,X3,X8,X9,X10,X11) => eqangle(X1,X0,X2,X3,X8,X9,X10,X11))),
  file('./6_GDD_FULL/61-80/67.p',ruleD18)).
fof(f14598,plain,(
  ( ! [X92,X93] : (midp(X92,sK26,sK24) | ~midp(X92,X93,X93)) ) | ~spl31_244),
  inference(resolution,[],[f14452,f1731])).
fof(f1731,plain,(
  ( ! [X2,X3,X1] : (~midp(sK22,X1,X2) | midp(X3,sK26,sK24) | ~midp(X3,X1,X2)) )),
  inference(subsumption_resolution,[],[f1726,f583])).
fof(f583,plain,(
  ( ! [X10,X11] : (para(X10,sK26,X11,sK24) | ~midp(sK22,X10,X11)) )),
  inference(resolution,[],[f423,f496])).
fof(f496,plain,(
  midp(sK22,sK26,sK24)),
  inference(resolution,[],[f370,f481])).
fof(f481,plain,(
  midp(sK22,sK24,sK26)),
  inference(cnf_transformation,[],[f359])).
fof(f370,plain,(
  ( ! [X2,X0,X1] : (~midp(X2,X1,X0) | midp(X2,X0,X1)) )),
  inference(cnf_transformation,[],[f183])).
fof(f183,plain,(
  ! [X0,X1,X2] : (midp(X2,X0,X1) | ~midp(X2,X1,X0))),
  inference(ennf_transformation,[],[f97])).
fof(f97,plain,(
  ! [X0,X1,X2] : (midp(X2,X1,X0) => midp(X2,X0,X1))),
  inference(rectify,[],[f11])).
fof(f11,axiom,(
  ! [X0,X1,X6] : (midp(X6,X1,X0) => midp(X6,X0,X1))),
  file('./6_GDD_FULL/61-80/67.p',ruleD11)).
fof(f423,plain,(
  ( ! [X4,X2,X0,X3,X1] : (~midp(X4,X2,X3) | para(X0,X2,X1,X3) | ~midp(X4,X0,X1)) )),
  inference(cnf_transformation,[],[f263])).
fof(f263,plain,(
  ! [X0,X1,X2,X3,X4] : (para(X0,X2,X1,X3) | ~midp(X4,X2,X3) | ~midp(X4,X0,X1))),
  inference(flattening,[],[f262])).
fof(f262,plain,(
  ! [X0,X1,X2,X3,X4] : (para(X0,X2,X1,X3) | (~midp(X4,X2,X3) | ~midp(X4,X0,X1)))),
  inference(ennf_transformation,[],[f142])).
fof(f142,plain,(
  ! [X0,X1,X2,X3,X4] : ((midp(X4,X2,X3) & midp(X4,X0,X1)) => para(X0,X2,X1,X3))),
  inference(rectify,[],[f64])).
fof(f64,axiom,(
  ! [X0,X1,X2,X3,X6] : ((midp(X6,X2,X3) & midp(X6,X0,X1)) => para(X0,X2,X1,X3))),
  file('./6_GDD_FULL/61-80/67.p',ruleD63)).
fof(f1726,plain,(
  ( ! [X2,X3,X1] : (~midp(sK22,X1,X2) | midp(X3,sK26,sK24) | ~para(X1,sK26,X2,sK24) | ~midp(X3,X1,X2)) )),
  inference(resolution,[],[f582,f424])).
fof(f424,plain,(
  ( ! [X4,X2,X0,X3,X1] : (~para(X0,X3,X1,X2) | midp(X4,X2,X3) | ~para(X0,X2,X1,X3) | ~midp(X4,X0,X1)) )),
  inference(cnf_transformation,[],[f265])).
fof(f265,plain,(
  ! [X0,X1,X2,X3,X4] : (midp(X4,X2,X3) | ~para(X0,X3,X1,X2) | ~para(X0,X2,X1,X3) | ~midp(X4,X0,X1))),
  inference(flattening,[],[f264])).
fof(f264,plain,(
  ! [X0,X1,X2,X3,X4] : (midp(X4,X2,X3) | (~para(X0,X3,X1,X2) | ~para(X0,X2,X1,X3) | ~midp(X4,X0,X1)))),
  inference(ennf_transformation,[],[f143])).
fof(f143,plain,(
  ! [X0,X1,X2,X3,X4] : ((para(X0,X3,X1,X2) & para(X0,X2,X1,X3) & midp(X4,X0,X1)) => midp(X4,X2,X3))),
  inference(rectify,[],[f65])).
fof(f65,axiom,(
  ! [X0,X1,X2,X3,X6] : ((para(X0,X3,X1,X2) & para(X0,X2,X1,X3) & midp(X6,X0,X1)) => midp(X6,X2,X3))),
  file('./6_GDD_FULL/61-80/67.p',ruleD64)).
fof(f582,plain,(
  ( ! [X8,X9] : (para(X8,sK24,X9,sK26) | ~midp(sK22,X8,X9)) )),
  inference(resolution,[],[f423,f481])).
fof(f14847,plain,(
  ( ! [X6] : (~midp(X6,sK26,sK24) | cong(sK26,X6,sK25,X6)) ) | ~spl31_98),
  inference(resolution,[],[f1201,f412])).
fof(f412,plain,(
  ( ! [X2,X0,X3,X1] : (~perp(X0,X1,X1,X2) | ~midp(X3,X0,X2) | cong(X0,X3,X1,X3)) )),
  inference(cnf_transformation,[],[f244])).
fof(f244,plain,(
  ! [X0,X1,X2,X3] : (cong(X0,X3,X1,X3) | ~midp(X3,X0,X2) | ~perp(X0,X1,X1,X2))),
  inference(flattening,[],[f243])).
fof(f243,plain,(
  ! [X0,X1,X2,X3] : (cong(X0,X3,X1,X3) | (~midp(X3,X0,X2) | ~perp(X0,X1,X1,X2)))),
  inference(ennf_transformation,[],[f132])).
fof(f132,plain,(
  ! [X0,X1,X2,X3] : ((midp(X3,X0,X2) & perp(X0,X1,X1,X2)) => cong(X0,X3,X1,X3))),
  inference(rectify,[],[f53])).
fof(f53,axiom,(
  ! [X0,X1,X2,X6] : ((midp(X6,X0,X2) & perp(X0,X1,X1,X2)) => cong(X0,X6,X1,X6))),
  file('./6_GDD_FULL/61-80/67.p',ruleD52)).
fof(f1201,plain,(
  perp(sK26,sK25,sK25,sK24) | ~spl31_98),
  inference(avatar_component_clause,[],[f1200])).
fof(f1200,plain,(
  spl31_98 <=> perp(sK26,sK25,sK25,sK24)),
  introduced(avatar_definition,[new_symbols(naming,[spl31_98])])).
fof(f15886,plain,(
  ( ! [X12,X13] : (perp(sK26,sK25,X12,X13) | ~cong(sK26,X12,sK25,X12)) ) | (~spl31_98 | ~spl31_244)),
  inference(resolution,[],[f14854,f416])).
fof(f416,plain,(
  ( ! [X2,X0,X3,X1] : (~cong(X0,X3,X1,X3) | perp(X0,X1,X2,X3) | ~cong(X0,X2,X1,X2)) )),
  inference(cnf_transformation,[],[f252])).
fof(f252,plain,(
  ! [X0,X1,X2,X3] : (perp(X0,X1,X2,X3) | ~cong(X0,X3,X1,X3) | ~cong(X0,X2,X1,X2))),
  inference(flattening,[],[f251])).
fof(f251,plain,(
  ! [X0,X1,X2,X3] : (perp(X0,X1,X2,X3) | (~cong(X0,X3,X1,X3) | ~cong(X0,X2,X1,X2)))),
  inference(ennf_transformation,[],[f135])).
fof(f135,plain,(
  ! [X0,X1,X2,X3] : ((cong(X0,X3,X1,X3) & cong(X0,X2,X1,X2)) => perp(X0,X1,X2,X3))),
  inference(rectify,[],[f57])).
fof(f57,axiom,(
  ! [X0,X1,X8,X9] : ((cong(X0,X9,X1,X9) & cong(X0,X8,X1,X8)) => perp(X0,X1,X8,X9))),
  file('./6_GDD_FULL/61-80/67.p',ruleD56)).
fof(f20753,plain,(
  ( ! [X28,X26,X27,X25] : (para(X25,X26,X27,X28) | ~perp(X25,X26,sK26,sK25)) ) | (~spl31_98 | ~spl31_244)),
  inference(resolution,[],[f15889,f368])).
fof(f368,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (~perp(X2,X3,X4,X5) | para(X0,X1,X4,X5) | ~perp(X0,X1,X2,X3)) )),
  inference(cnf_transformation,[],[f180])).
fof(f180,plain,(
  ! [X0,X1,X2,X3,X4,X5] : (para(X0,X1,X4,X5) | ~perp(X2,X3,X4,X5) | ~perp(X0,X1,X2,X3))),
  inference(flattening,[],[f179])).
fof(f179,plain,(
  ! [X0,X1,X2,X3,X4,X5] : (para(X0,X1,X4,X5) | (~perp(X2,X3,X4,X5) | ~perp(X0,X1,X2,X3)))),
  inference(ennf_transformation,[],[f9])).
fof(f9,axiom,(
  ! [X0,X1,X2,X3,X4,X5] : ((perp(X2,X3,X4,X5) & perp(X0,X1,X2,X3)) => para(X0,X1,X4,X5))),
  file('./6_GDD_FULL/61-80/67.p',ruleD9)).
fof(f20754,plain,(
  ( ! [X30,X31,X29,X32] : (perp(X29,X30,X31,X32) | ~para(X29,X30,sK26,sK25)) ) | (~spl31_98 | ~spl31_244)),
  inference(resolution,[],[f15889,f369])).
fof(f369,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (~perp(X2,X3,X4,X5) | perp(X0,X1,X4,X5) | ~para(X0,X1,X2,X3)) )),
  inference(cnf_transformation,[],[f182])).
fof(f182,plain,(
  ! [X0,X1,X2,X3,X4,X5] : (perp(X0,X1,X4,X5) | ~perp(X2,X3,X4,X5) | ~para(X0,X1,X2,X3))),
  inference(flattening,[],[f181])).
fof(f181,plain,(
  ! [X0,X1,X2,X3,X4,X5] : (perp(X0,X1,X4,X5) | (~perp(X2,X3,X4,X5) | ~para(X0,X1,X2,X3)))),
  inference(ennf_transformation,[],[f10])).
fof(f10,axiom,(
  ! [X0,X1,X2,X3,X4,X5] : ((perp(X2,X3,X4,X5) & para(X0,X1,X2,X3)) => perp(X0,X1,X4,X5))),
  file('./6_GDD_FULL/61-80/67.p',ruleD10)).
fof(f14482,plain,(
  spl31_98 | ~spl31_97),
  inference(avatar_split_clause,[],[f1968,f1194,f1200])).
fof(f1194,plain,(
  spl31_97 <=> ~perp(sK26,sK25,sK25,sK23)),
  introduced(avatar_definition,[new_symbols(naming,[spl31_97])])).
fof(f1968,plain,(
  ~perp(sK26,sK25,sK25,sK23) | perp(sK26,sK25,sK25,sK24)),
  inference(resolution,[],[f825,f434])).
fof(f434,plain,(
  ( ! [X6,X4,X2,X0,X7,X5,X3,X1] : (~eqangle(X0,X1,X2,X3,X4,X5,X6,X7) | ~perp(X4,X5,X6,X7) | perp(X0,X1,X2,X3)) )),
  inference(cnf_transformation,[],[f282])).
fof(f282,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6,X7] : (perp(X0,X1,X2,X3) | ~perp(X4,X5,X6,X7) | ~eqangle(X0,X1,X2,X3,X4,X5,X6,X7))),
  inference(flattening,[],[f281])).
fof(f281,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6,X7] : (perp(X0,X1,X2,X3) | (~perp(X4,X5,X6,X7) | ~eqangle(X0,X1,X2,X3,X4,X5,X6,X7)))),
  inference(ennf_transformation,[],[f147])).
fof(f147,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6,X7] : ((perp(X4,X5,X6,X7) & eqangle(X0,X1,X2,X3,X4,X5,X6,X7)) => perp(X0,X1,X2,X3))),
  inference(rectify,[],[f75])).
fof(f75,axiom,(
  ! [X0,X1,X2,X3,X8,X9,X10,X11] : ((perp(X8,X9,X10,X11) & eqangle(X0,X1,X2,X3,X8,X9,X10,X11)) => perp(X0,X1,X2,X3))),
  file('./6_GDD_FULL/61-80/67.p',ruleD74)).
fof(f825,plain,(
  eqangle(sK26,sK25,sK25,sK24,sK26,sK25,sK25,sK23)),
  inference(resolution,[],[f379,f475])).
fof(f475,plain,(
  eqangle(sK26,sK25,sK25,sK23,sK26,sK25,sK25,sK24)),
  inference(cnf_transformation,[],[f359])).
fof(f14449,plain,(
  ~spl31_244 | spl31_629),
  inference(avatar_contradiction_clause,[],[f14448])).
fof(f14448,plain,(
  $false | (~spl31_244 | ~spl31_629)),
  inference(resolution,[],[f14407,f5813])).
fof(f5813,plain,(
  ~cong(sK25,sK23,sK25,sK23) | ~spl31_629),
  inference(avatar_component_clause,[],[f5812])).
fof(f5812,plain,(
  spl31_629 <=> ~cong(sK25,sK23,sK25,sK23)),
  introduced(avatar_definition,[new_symbols(naming,[spl31_629])])).
fof(f12427,plain,(
  spl31_626 | ~spl31_220 | ~spl31_244),
  inference(avatar_split_clause,[],[f12426,f2511,f2105,f5803])).
fof(f5803,plain,(
  spl31_626 <=> cong(sK25,sK26,sK25,sK26)),
  introduced(avatar_definition,[new_symbols(naming,[spl31_626])])).
fof(f2105,plain,(
  spl31_220 <=> cyclic(sK25,sK25,sK26,sK25)),
  introduced(avatar_definition,[new_symbols(naming,[spl31_220])])).
fof(f12426,plain,(
  cong(sK25,sK26,sK25,sK26) | (~spl31_220 | ~spl31_244)),
  inference(subsumption_resolution,[],[f12425,f5940])).
fof(f5940,plain,(
  cyclic(sK25,sK26,sK25,sK25) | ~spl31_220),
  inference(resolution,[],[f2106,f374])).
fof(f2106,plain,(
  cyclic(sK25,sK25,sK26,sK25) | ~spl31_220),
  inference(avatar_component_clause,[],[f2105])).
fof(f12425,plain,(
  cong(sK25,sK26,sK25,sK26) | ~cyclic(sK25,sK26,sK25,sK25) | (~spl31_220 | ~spl31_244)),
  inference(subsumption_resolution,[],[f8280,f10753])).
fof(f8280,plain,(
  cong(sK25,sK26,sK25,sK26) | ~cyclic(sK25,sK26,sK25,sK26) | ~cyclic(sK25,sK26,sK25,sK25) | ~spl31_220),
  inference(resolution,[],[f5940,f1589])).
fof(f1589,plain,(
  ( ! [X10,X8,X7,X9] : (~cyclic(X7,X8,X9,X10) | cong(X7,X8,X7,X8) | ~cyclic(X7,X8,X9,X8) | ~cyclic(X7,X8,X9,X7)) )),
  inference(duplicate_literal_removal,[],[f1583])).
fof(f1583,plain,(
  ( ! [X10,X8,X7,X9] : (cong(X7,X8,X7,X8) | ~cyclic(X7,X8,X9,X10) | ~cyclic(X7,X8,X9,X8) | ~cyclic(X7,X8,X9,X7) | ~cyclic(X7,X8,X9,X10)) )),
  inference(resolution,[],[f403,f400])).
fof(f6206,plain,(
  ~spl31_326),
  inference(avatar_contradiction_clause,[],[f6179])).
fof(f6179,plain,(
  $false | ~spl31_326),
  inference(resolution,[],[f3216,f485])).
fof(f485,plain,(
  midp(sK20,sK26,sK25)),
  inference(cnf_transformation,[],[f359])).
fof(f3216,plain,(
  ( ! [X12,X10,X11] : (~midp(X10,X11,X12)) ) | ~spl31_326),
  inference(avatar_component_clause,[],[f3215])).
fof(f3215,plain,(
  spl31_326 <=> ! [X11,X10,X12] : ~midp(X10,X11,X12)),
  introduced(avatar_definition,[new_symbols(naming,[spl31_326])])).
fof(f5814,plain,(
  ~spl31_627 | ~spl31_629 | spl31_97 | ~spl31_216),
  inference(avatar_split_clause,[],[f5801,f2093,f1194,f5812,f5806])).
fof(f5806,plain,(
  spl31_627 <=> ~cong(sK25,sK26,sK25,sK26)),
  introduced(avatar_definition,[new_symbols(naming,[spl31_627])])).
fof(f2093,plain,(
  spl31_216 <=> cyclic(sK25,sK25,sK26,sK23)),
  introduced(avatar_definition,[new_symbols(naming,[spl31_216])])).
fof(f5801,plain,(
  ~cong(sK25,sK23,sK25,sK23) | ~cong(sK25,sK26,sK25,sK26) | (~spl31_97 | ~spl31_216)),
  inference(subsumption_resolution,[],[f5793,f1195])).
fof(f1195,plain,(
  ~perp(sK26,sK25,sK25,sK23) | ~spl31_97),
  inference(avatar_component_clause,[],[f1194])).
fof(f5793,plain,(
  perp(sK26,sK25,sK25,sK23) | ~cong(sK25,sK23,sK25,sK23) | ~cong(sK25,sK26,sK25,sK26) | ~spl31_216),
  inference(resolution,[],[f2094,f417])).
fof(f417,plain,(
  ( ! [X2,X0,X3,X1] : (~cyclic(X0,X1,X2,X3) | perp(X2,X0,X0,X3) | ~cong(X0,X3,X1,X3) | ~cong(X0,X2,X1,X2)) )),
  inference(cnf_transformation,[],[f254])).
fof(f254,plain,(
  ! [X0,X1,X2,X3] : (perp(X2,X0,X0,X3) | ~cyclic(X0,X1,X2,X3) | ~cong(X0,X3,X1,X3) | ~cong(X0,X2,X1,X2))),
  inference(flattening,[],[f253])).
fof(f253,plain,(
  ! [X0,X1,X2,X3] : (perp(X2,X0,X0,X3) | (~cyclic(X0,X1,X2,X3) | ~cong(X0,X3,X1,X3) | ~cong(X0,X2,X1,X2)))),
  inference(ennf_transformation,[],[f136])).
fof(f136,plain,(
  ! [X0,X1,X2,X3] : ((cyclic(X0,X1,X2,X3) & cong(X0,X3,X1,X3) & cong(X0,X2,X1,X2)) => perp(X2,X0,X0,X3))),
  inference(rectify,[],[f58])).
fof(f58,axiom,(
  ! [X0,X1,X8,X9] : ((cyclic(X0,X1,X8,X9) & cong(X0,X9,X1,X9) & cong(X0,X8,X1,X8)) => perp(X8,X0,X0,X9))),
  file('./6_GDD_FULL/61-80/67.p',ruleD57)).
fof(f2094,plain,(
  cyclic(sK25,sK25,sK26,sK23) | ~spl31_216),
  inference(avatar_component_clause,[],[f2093])).
fof(f5724,plain,(
  spl31_326 | spl31_326 | ~spl31_246),
  inference(avatar_split_clause,[],[f5723,f2514,f3215,f3215])).
fof(f2514,plain,(
  spl31_246 <=> ! [X1,X0] : ~para(X0,X1,X0,X1)),
  introduced(avatar_definition,[new_symbols(naming,[spl31_246])])).
fof(f5723,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (~midp(X0,X1,X2) | ~midp(X4,X5,X3)) ) | ~spl31_246),
  inference(subsumption_resolution,[],[f5722,f2515])).
fof(f2515,plain,(
  ( ! [X0,X1] : (~para(X0,X1,X0,X1)) ) | ~spl31_246),
  inference(avatar_component_clause,[],[f2514])).
fof(f5722,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (~midp(X4,X5,X3) | para(X2,X3,X2,X3) | ~midp(X0,X1,X2)) )),
  inference(duplicate_literal_removal,[],[f5705])).
fof(f5705,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (~midp(X0,X1,X2) | para(X2,X3,X2,X3) | ~midp(X4,X5,X3) | ~midp(X0,X1,X2) | ~midp(X4,X5,X3)) )),
  inference(resolution,[],[f1217,f1218])).
fof(f1218,plain,(
  ( ! [X26,X24,X23,X27,X25,X22] : (para(X27,X24,sK5(X26,X27,X23,X24,X25,X22),X25) | ~midp(X25,X26,X27) | ~midp(X22,X23,X24)) )),
  inference(resolution,[],[f447,f364])).
fof(f364,plain,(
  ( ! [X2,X0,X3,X1] : (~para(X0,X1,X2,X3) | para(X2,X3,X0,X1)) )),
  inference(cnf_transformation,[],[f174])).
fof(f174,plain,(
  ! [X0,X1,X2,X3] : (para(X2,X3,X0,X1) | ~para(X0,X1,X2,X3))),
  inference(ennf_transformation,[],[f5])).
fof(f5,axiom,(
  ! [X0,X1,X2,X3] : (para(X0,X1,X2,X3) => para(X2,X3,X0,X1))),
  file('./6_GDD_FULL/61-80/67.p',ruleD5)).
fof(f447,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (para(sK5(X0,X1,X2,X3,X4,X5),X4,X1,X3) | ~midp(X5,X2,X3) | ~midp(X4,X0,X1)) )),
  inference(cnf_transformation,[],[f333])).
fof(f333,plain,(
  ! [X0,X1,X2,X3,X4,X5] : ((para(sK5(X0,X1,X2,X3,X4,X5),X5,X0,X2) & para(sK5(X0,X1,X2,X3,X4,X5),X4,X1,X3) & midp(sK5(X0,X1,X2,X3,X4,X5),X0,X3)) | ~midp(X5,X2,X3) | ~midp(X4,X0,X1))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK5])],[f296,f332])).
fof(f332,plain,(
  ! [X5,X4,X3,X2,X1,X0] : (? [X6] : ((para(X6,X5,X0,X2) & para(X6,X4,X1,X3) & midp(X6,X0,X3)) | ~midp(X5,X2,X3) | ~midp(X4,X0,X1)) => ((para(sK5(X0,X1,X2,X3,X4,X5),X5,X0,X2) & para(sK5(X0,X1,X2,X3,X4,X5),X4,X1,X3) & midp(sK5(X0,X1,X2,X3,X4,X5),X0,X3)) | ~midp(X5,X2,X3) | ~midp(X4,X0,X1)))),
  introduced(choice_axiom,[])).
fof(f296,plain,(
  ! [X0,X1,X2,X3,X4,X5] : ? [X6] : ((para(X6,X5,X0,X2) & para(X6,X4,X1,X3) & midp(X6,X0,X3)) | ~midp(X5,X2,X3) | ~midp(X4,X0,X1))),
  inference(flattening,[],[f295])).
fof(f295,plain,(
  ! [X0,X1,X2,X3,X4,X5] : ? [X6] : ((para(X6,X5,X0,X2) & para(X6,X4,X1,X3) & midp(X6,X0,X3)) | (~midp(X5,X2,X3) | ~midp(X4,X0,X1)))),
  inference(ennf_transformation,[],[f154])).
fof(f154,plain,(
  ! [X0,X1,X2,X3,X4,X5] : ? [X6] : ((midp(X5,X2,X3) & midp(X4,X0,X1)) => (para(X6,X5,X0,X2) & para(X6,X4,X1,X3) & midp(X6,X0,X3)))),
  inference(rectify,[],[f82])).
fof(f82,axiom,(
  ! [X0,X1,X2,X3,X6,X16] : ? [X8] : ((midp(X16,X2,X3) & midp(X6,X0,X1)) => (para(X8,X16,X0,X2) & para(X8,X6,X1,X3) & midp(X8,X0,X3)))),
  file('./6_GDD_FULL/61-80/67.p',ruleX6)).
fof(f1217,plain,(
  ( ! [X14,X21,X19,X17,X15,X20,X18,X16] : (~para(X20,X21,sK5(X18,X19,X15,X16,X17,X14),X17) | ~midp(X17,X18,X19) | para(X20,X21,X19,X16) | ~midp(X14,X15,X16)) )),
  inference(resolution,[],[f447,f365])).
fof(f365,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (~para(X2,X3,X4,X5) | para(X0,X1,X4,X5) | ~para(X0,X1,X2,X3)) )),
  inference(cnf_transformation,[],[f176])).
fof(f176,plain,(
  ! [X0,X1,X2,X3,X4,X5] : (para(X0,X1,X4,X5) | ~para(X2,X3,X4,X5) | ~para(X0,X1,X2,X3))),
  inference(flattening,[],[f175])).
fof(f175,plain,(
  ! [X0,X1,X2,X3,X4,X5] : (para(X0,X1,X4,X5) | (~para(X2,X3,X4,X5) | ~para(X0,X1,X2,X3)))),
  inference(ennf_transformation,[],[f6])).
fof(f6,axiom,(
  ! [X0,X1,X2,X3,X4,X5] : ((para(X2,X3,X4,X5) & para(X0,X1,X2,X3)) => para(X0,X1,X4,X5))),
  file('./6_GDD_FULL/61-80/67.p',ruleD6)).
fof(f3747,plain,(
  spl31_217 | ~spl31_244),
  inference(avatar_contradiction_clause,[],[f3746])).
fof(f3746,plain,(
  $false | (~spl31_217 | ~spl31_244)),
  inference(subsumption_resolution,[],[f2097,f2712])).
fof(f2712,plain,(
  ( ! [X70,X71,X69] : (cyclic(X70,X70,X69,X71)) ) | ~spl31_244),
  inference(subsumption_resolution,[],[f2711,f2710])).
fof(f2710,plain,(
  ( ! [X68,X66,X67] : (~coll(X66,X68,X67) | cyclic(X67,X67,X66,X68)) ) | ~spl31_244),
  inference(subsumption_resolution,[],[f2691,f2512])).
fof(f2711,plain,(
  ( ! [X70,X71,X69] : (coll(X69,X71,X70) | cyclic(X70,X70,X69,X71)) ) | ~spl31_244),
  inference(subsumption_resolution,[],[f2692,f2512])).
fof(f2097,plain,(
  ~cyclic(sK25,sK25,sK26,sK23) | ~spl31_217),
  inference(avatar_component_clause,[],[f2096])).
fof(f2096,plain,(
  spl31_217 <=> ~cyclic(sK25,sK25,sK26,sK23)),
  introduced(avatar_definition,[new_symbols(naming,[spl31_217])])).
fof(f3740,plain,(
  spl31_221 | ~spl31_244),
  inference(avatar_contradiction_clause,[],[f3739])).
fof(f3739,plain,(
  $false | (~spl31_221 | ~spl31_244)),
  inference(resolution,[],[f2712,f2109])).
fof(f2109,plain,(
  ~cyclic(sK25,sK25,sK26,sK25) | ~spl31_221),
  inference(avatar_component_clause,[],[f2108])).
fof(f2108,plain,(
  spl31_221 <=> ~cyclic(sK25,sK25,sK26,sK25)),
  introduced(avatar_definition,[new_symbols(naming,[spl31_221])])).
fof(f2516,plain,(
  spl31_244 | spl31_246),
  inference(avatar_split_clause,[],[f2487,f2514,f2511])).
fof(f2487,plain,(
  ( ! [X2,X0,X3,X1] : (~para(X0,X1,X0,X1) | para(X2,X3,X2,X3)) )),
  inference(resolution,[],[f801,f398])).
fof(f398,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (~eqangle(X0,X1,X4,X5,X2,X3,X4,X5) | para(X0,X1,X2,X3)) )),
  inference(cnf_transformation,[],[f219])).
fof(f219,plain,(
  ! [X0,X1,X2,X3,X4,X5] : (para(X0,X1,X2,X3) | ~eqangle(X0,X1,X4,X5,X2,X3,X4,X5))),
  inference(ennf_transformation,[],[f118])).
fof(f118,plain,(
  ! [X0,X1,X2,X3,X4,X5] : (eqangle(X0,X1,X4,X5,X2,X3,X4,X5) => para(X0,X1,X2,X3))),
  inference(rectify,[],[f39])).
fof(f39,axiom,(
  ! [X0,X1,X2,X3,X8,X9] : (eqangle(X0,X1,X8,X9,X2,X3,X8,X9) => para(X0,X1,X2,X3))),
  file('./6_GDD_FULL/61-80/67.p',ruleD39)).
fof(f801,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (eqangle(X0,X1,X2,X3,X0,X1,X4,X5) | ~para(X2,X3,X4,X5)) )),
  inference(resolution,[],[f378,f399])).
fof(f378,plain,(
  ( ! [X6,X4,X2,X0,X7,X5,X3,X1] : (~eqangle(X0,X1,X2,X3,X4,X5,X6,X7) | eqangle(X2,X3,X0,X1,X6,X7,X4,X5)) )),
  inference(cnf_transformation,[],[f194])).
fof(f194,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6,X7] : (eqangle(X2,X3,X0,X1,X6,X7,X4,X5) | ~eqangle(X0,X1,X2,X3,X4,X5,X6,X7))),
  inference(ennf_transformation,[],[f101])).
fof(f101,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6,X7] : (eqangle(X0,X1,X2,X3,X4,X5,X6,X7) => eqangle(X2,X3,X0,X1,X6,X7,X4,X5))),
  inference(rectify,[],[f19])).
fof(f19,axiom,(
  ! [X0,X1,X2,X3,X8,X9,X10,X11] : (eqangle(X0,X1,X2,X3,X8,X9,X10,X11) => eqangle(X2,X3,X0,X1,X10,X11,X8,X9))),
  file('./6_GDD_FULL/61-80/67.p',ruleD19)).
% SZS output end Proof for 67
% ------------------------------
% Version: Vampire 4.4.0 (commit 7916d27 on 2019-08-23 08:50:16 +0100)
% Termination reason: Refutation

% Memory used [KB]: 27121
% Time elapsed: 0.640 s
% ------------------------------
% ------------------------------
