% Refutation found. Thanks to Tanya!
% SZS status Theorem for 72
% SZS output start Proof for 72
fof(f25164,plain,(
  $false),
  inference(avatar_sat_refutation,[],[f564,f2172,f2527,f3498,f9240,f14250,f14277,f14647,f15686,f24334,f25121,f25153,f25155,f25159])).
fof(f25159,plain,(
  spl30_13 | ~spl30_154 | ~spl30_178 | ~spl30_224 | ~spl30_456),
  inference(avatar_contradiction_clause,[],[f25158])).
fof(f25158,plain,(
  $false | (~spl30_13 | ~spl30_154 | ~spl30_178 | ~spl30_224 | ~spl30_456)),
  inference(subsumption_resolution,[],[f549,f25134])).
fof(f25134,plain,(
  ( ! [X61,X59,X57,X62,X60,X58,X56,X63] : (eqangle(X58,X59,X62,X63,X56,X57,X60,X61)) ) | (~spl30_154 | ~spl30_178 | ~spl30_224 | ~spl30_456)),
  inference(subsumption_resolution,[],[f25133,f24685])).
fof(f24685,plain,(
  ( ! [X10,X8,X11,X9] : (para(X8,X9,X10,X11)) ) | (~spl30_154 | ~spl30_178 | ~spl30_224 | ~spl30_456)),
  inference(subsumption_resolution,[],[f24662,f24482])).
fof(f24482,plain,(
  ( ! [X4,X5] : (midp(X4,sK22,X5)) ) | (~spl30_154 | ~spl30_178 | ~spl30_224 | ~spl30_456)),
  inference(resolution,[],[f24413,f376])).
fof(f376,plain,(
  ( ! [X2,X0,X1] : (~midp(X2,X1,X0) | midp(X2,X0,X1)) )),
  inference(cnf_transformation,[],[f182])).
fof(f182,plain,(
  ! [X0,X1,X2] : (midp(X2,X0,X1) | ~midp(X2,X1,X0))),
  inference(ennf_transformation,[],[f97])).
fof(f97,plain,(
  ! [X0,X1,X2] : (midp(X2,X1,X0) => midp(X2,X0,X1))),
  inference(rectify,[],[f11])).
fof(f11,axiom,(
  ! [X0,X1,X6] : (midp(X6,X1,X0) => midp(X6,X0,X1))),
  file('./6_GDD_FULL/61-80/72.p',ruleD11)).
fof(f24413,plain,(
  ( ! [X8,X9] : (midp(X8,X9,sK22)) ) | (~spl30_154 | ~spl30_178 | ~spl30_224 | ~spl30_456)),
  inference(subsumption_resolution,[],[f24412,f19267])).
fof(f19267,plain,(
  ( ! [X0] : (midp(sK22,X0,sK24)) ) | (~spl30_154 | ~spl30_178 | ~spl30_224)),
  inference(resolution,[],[f19206,f3149])).
fof(f3149,plain,(
  ( ! [X1] : (~midp(sK24,X1,sK22) | midp(sK22,X1,sK24)) ) | ~spl30_178),
  inference(subsumption_resolution,[],[f3145,f3113])).
fof(f3113,plain,(
  ( ! [X2,X0,X1] : (coll(X0,X1,X2)) ) | ~spl30_178),
  inference(subsumption_resolution,[],[f3107,f3086])).
fof(f3086,plain,(
  ( ! [X6,X5] : (coll(X5,X5,X6)) ) | ~spl30_178),
  inference(resolution,[],[f3061,f366])).
fof(f366,plain,(
  ( ! [X2,X0,X1] : (~coll(X0,X1,X2) | coll(X0,X2,X1)) )),
  inference(cnf_transformation,[],[f168])).
fof(f168,plain,(
  ! [X0,X1,X2] : (coll(X0,X2,X1) | ~coll(X0,X1,X2))),
  inference(ennf_transformation,[],[f1])).
fof(f1,axiom,(
  ! [X0,X1,X2] : (coll(X0,X1,X2) => coll(X0,X2,X1))),
  file('./6_GDD_FULL/61-80/72.p',ruleD1)).
fof(f3061,plain,(
  ( ! [X4,X3] : (coll(X3,X4,X3)) ) | ~spl30_178),
  inference(resolution,[],[f3045,f367])).
fof(f367,plain,(
  ( ! [X2,X0,X1] : (~coll(X0,X1,X2) | coll(X1,X0,X2)) )),
  inference(cnf_transformation,[],[f169])).
fof(f169,plain,(
  ! [X0,X1,X2] : (coll(X1,X0,X2) | ~coll(X0,X1,X2))),
  inference(ennf_transformation,[],[f2])).
fof(f2,axiom,(
  ! [X0,X1,X2] : (coll(X0,X1,X2) => coll(X1,X0,X2))),
  file('./6_GDD_FULL/61-80/72.p',ruleD2)).
fof(f3045,plain,(
  ( ! [X0,X1] : (coll(X0,X1,X1)) ) | ~spl30_178),
  inference(resolution,[],[f2168,f432])).
fof(f432,plain,(
  ( ! [X2,X0,X1] : (~para(X0,X1,X0,X2) | coll(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f267])).
fof(f267,plain,(
  ! [X0,X1,X2] : (coll(X0,X1,X2) | ~para(X0,X1,X0,X2))),
  inference(ennf_transformation,[],[f67])).
fof(f67,axiom,(
  ! [X0,X1,X2] : (para(X0,X1,X0,X2) => coll(X0,X1,X2))),
  file('./6_GDD_FULL/61-80/72.p',ruleD66)).
fof(f2168,plain,(
  ( ! [X2,X3] : (para(X2,X3,X2,X3)) ) | ~spl30_178),
  inference(avatar_component_clause,[],[f2167])).
fof(f2167,plain,(
  spl30_178 <=> ! [X3,X2] : para(X2,X3,X2,X3)),
  introduced(avatar_definition,[new_symbols(naming,[spl30_178])])).
fof(f3107,plain,(
  ( ! [X2,X0,X1] : (coll(X0,X1,X2) | ~coll(X2,X2,X0)) ) | ~spl30_178),
  inference(resolution,[],[f3086,f368])).
fof(f368,plain,(
  ( ! [X2,X0,X3,X1] : (~coll(X0,X1,X3) | coll(X2,X3,X0) | ~coll(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f171])).
fof(f171,plain,(
  ! [X0,X1,X2,X3] : (coll(X2,X3,X0) | ~coll(X0,X1,X3) | ~coll(X0,X1,X2))),
  inference(flattening,[],[f170])).
fof(f170,plain,(
  ! [X0,X1,X2,X3] : (coll(X2,X3,X0) | (~coll(X0,X1,X3) | ~coll(X0,X1,X2)))),
  inference(ennf_transformation,[],[f3])).
fof(f3,axiom,(
  ! [X0,X1,X2,X3] : ((coll(X0,X1,X3) & coll(X0,X1,X2)) => coll(X2,X3,X0))),
  file('./6_GDD_FULL/61-80/72.p',ruleD3)).
fof(f3145,plain,(
  ( ! [X1] : (~coll(sK22,X1,sK24) | midp(sK22,X1,sK24) | ~midp(sK24,X1,sK22)) )),
  inference(resolution,[],[f1543,f411])).
fof(f411,plain,(
  ( ! [X4,X2,X0,X3,X1] : (~para(X3,X4,X1,X2) | ~coll(X4,X0,X2) | midp(X4,X0,X2) | ~midp(X3,X0,X1)) )),
  inference(cnf_transformation,[],[f230])).
fof(f230,plain,(
  ! [X0,X1,X2,X3,X4] : (midp(X4,X0,X2) | ~coll(X4,X0,X2) | ~para(X3,X4,X1,X2) | ~midp(X3,X0,X1))),
  inference(flattening,[],[f229])).
fof(f229,plain,(
  ! [X0,X1,X2,X3,X4] : (midp(X4,X0,X2) | (~coll(X4,X0,X2) | ~para(X3,X4,X1,X2) | ~midp(X3,X0,X1)))),
  inference(ennf_transformation,[],[f125])).
fof(f125,plain,(
  ! [X0,X1,X2,X3,X4] : ((coll(X4,X0,X2) & para(X3,X4,X1,X2) & midp(X3,X0,X1)) => midp(X4,X0,X2))),
  inference(rectify,[],[f46])).
fof(f46,axiom,(
  ! [X0,X1,X2,X4,X5] : ((coll(X5,X0,X2) & para(X4,X5,X1,X2) & midp(X4,X0,X1)) => midp(X5,X0,X2))),
  file('./6_GDD_FULL/61-80/72.p',ruleD45)).
fof(f1543,plain,(
  para(sK24,sK22,sK22,sK24)),
  inference(resolution,[],[f921,f593])).
fof(f593,plain,(
  perp(sK24,sK22,sK27,sK26)),
  inference(resolution,[],[f585,f373])).
fof(f373,plain,(
  ( ! [X2,X0,X3,X1] : (~perp(X0,X1,X2,X3) | perp(X2,X3,X0,X1)) )),
  inference(cnf_transformation,[],[f177])).
fof(f177,plain,(
  ! [X0,X1,X2,X3] : (perp(X2,X3,X0,X1) | ~perp(X0,X1,X2,X3))),
  inference(ennf_transformation,[],[f8])).
fof(f8,axiom,(
  ! [X0,X1,X2,X3] : (perp(X0,X1,X2,X3) => perp(X2,X3,X0,X1))),
  file('./6_GDD_FULL/61-80/72.p',ruleD8)).
fof(f585,plain,(
  perp(sK27,sK26,sK24,sK22)),
  inference(resolution,[],[f372,f485])).
fof(f485,plain,(
  perp(sK27,sK26,sK22,sK24)),
  inference(cnf_transformation,[],[f365])).
fof(f365,plain,(
  (~eqangle(sK26,sK27,sK27,sK28,sK26,sK24,sK24,sK23) | (~eqangle(sK27,sK26,sK26,sK28,sK26,sK23,sK23,sK24) & ~eqangle(sK26,sK27,sK27,sK28,sK24,sK26,sK26,sK23)) | (~eqangle(sK27,sK26,sK26,sK28,sK26,sK23,sK23,sK24) & ~eqangle(sK26,sK27,sK27,sK28,sK26,sK23,sK23,sK24)) | (~eqangle(sK27,sK26,sK26,sK28,sK26,sK24,sK24,sK23) & ~eqangle(sK26,sK27,sK27,sK28,sK24,sK26,sK26,sK23)) | sP1(sK23,sK24,sK26,sK28,sK27) | sP0(sK23,sK26,sK24,sK28,sK27) | ~eqangle(sK27,sK26,sK26,sK28,sK24,sK26,sK26,sK23)) & coll(sK28,sK22,sK23) & perp(sK28,sK26,sK22,sK23) & coll(sK27,sK22,sK24) & perp(sK27,sK26,sK22,sK24) & circle(sK25,sK22,sK26,sK29) & circle(sK25,sK22,sK23,sK24)),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK22,sK23,sK24,sK25,sK26,sK27,sK28,sK29])],[f323,f364])).
fof(f364,plain,(
  ? [X0,X1,X2,X3,X4,X5,X6,X7] : ((~eqangle(X4,X5,X5,X6,X4,X2,X2,X1) | (~eqangle(X5,X4,X4,X6,X4,X1,X1,X2) & ~eqangle(X4,X5,X5,X6,X2,X4,X4,X1)) | (~eqangle(X5,X4,X4,X6,X4,X1,X1,X2) & ~eqangle(X4,X5,X5,X6,X4,X1,X1,X2)) | (~eqangle(X5,X4,X4,X6,X4,X2,X2,X1) & ~eqangle(X4,X5,X5,X6,X2,X4,X4,X1)) | sP1(X1,X2,X4,X6,X5) | sP0(X1,X4,X2,X6,X5) | ~eqangle(X5,X4,X4,X6,X2,X4,X4,X1)) & coll(X6,X0,X1) & perp(X6,X4,X0,X1) & coll(X5,X0,X2) & perp(X5,X4,X0,X2) & circle(X3,X0,X4,X7) & circle(X3,X0,X1,X2)) => ((~eqangle(sK26,sK27,sK27,sK28,sK26,sK24,sK24,sK23) | (~eqangle(sK27,sK26,sK26,sK28,sK26,sK23,sK23,sK24) & ~eqangle(sK26,sK27,sK27,sK28,sK24,sK26,sK26,sK23)) | (~eqangle(sK27,sK26,sK26,sK28,sK26,sK23,sK23,sK24) & ~eqangle(sK26,sK27,sK27,sK28,sK26,sK23,sK23,sK24)) | (~eqangle(sK27,sK26,sK26,sK28,sK26,sK24,sK24,sK23) & ~eqangle(sK26,sK27,sK27,sK28,sK24,sK26,sK26,sK23)) | sP1(sK23,sK24,sK26,sK28,sK27) | sP0(sK23,sK26,sK24,sK28,sK27) | ~eqangle(sK27,sK26,sK26,sK28,sK24,sK26,sK26,sK23)) & coll(sK28,sK22,sK23) & perp(sK28,sK26,sK22,sK23) & coll(sK27,sK22,sK24) & perp(sK27,sK26,sK22,sK24) & circle(sK25,sK22,sK26,sK29) & circle(sK25,sK22,sK23,sK24))),
  introduced(choice_axiom,[])).
fof(f323,plain,(
  ? [X0,X1,X2,X3,X4,X5,X6,X7] : ((~eqangle(X4,X5,X5,X6,X4,X2,X2,X1) | (~eqangle(X5,X4,X4,X6,X4,X1,X1,X2) & ~eqangle(X4,X5,X5,X6,X2,X4,X4,X1)) | (~eqangle(X5,X4,X4,X6,X4,X1,X1,X2) & ~eqangle(X4,X5,X5,X6,X4,X1,X1,X2)) | (~eqangle(X5,X4,X4,X6,X4,X2,X2,X1) & ~eqangle(X4,X5,X5,X6,X2,X4,X4,X1)) | sP1(X1,X2,X4,X6,X5) | sP0(X1,X4,X2,X6,X5) | ~eqangle(X5,X4,X4,X6,X2,X4,X4,X1)) & coll(X6,X0,X1) & perp(X6,X4,X0,X1) & coll(X5,X0,X2) & perp(X5,X4,X0,X2) & circle(X3,X0,X4,X7) & circle(X3,X0,X1,X2))),
  inference(definition_folding,[],[f320,f322,f321])).
fof(f321,plain,(
  ! [X1,X4,X2,X6,X5] : ((~eqangle(X5,X4,X4,X6,X2,X4,X4,X1) & ~eqangle(X4,X5,X5,X6,X4,X2,X2,X1)) | ~sP0(X1,X4,X2,X6,X5))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP0])])).
fof(f322,plain,(
  ! [X1,X2,X4,X6,X5] : ((~eqangle(X5,X4,X4,X6,X4,X2,X2,X1) & ~eqangle(X4,X5,X5,X6,X4,X1,X1,X2)) | ~sP1(X1,X2,X4,X6,X5))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP1])])).
fof(f320,plain,(
  ? [X0,X1,X2,X3,X4,X5,X6,X7] : ((~eqangle(X4,X5,X5,X6,X4,X2,X2,X1) | (~eqangle(X5,X4,X4,X6,X4,X1,X1,X2) & ~eqangle(X4,X5,X5,X6,X2,X4,X4,X1)) | (~eqangle(X5,X4,X4,X6,X4,X1,X1,X2) & ~eqangle(X4,X5,X5,X6,X4,X1,X1,X2)) | (~eqangle(X5,X4,X4,X6,X4,X2,X2,X1) & ~eqangle(X4,X5,X5,X6,X2,X4,X4,X1)) | (~eqangle(X5,X4,X4,X6,X4,X2,X2,X1) & ~eqangle(X4,X5,X5,X6,X4,X1,X1,X2)) | (~eqangle(X5,X4,X4,X6,X2,X4,X4,X1) & ~eqangle(X4,X5,X5,X6,X4,X2,X2,X1)) | ~eqangle(X5,X4,X4,X6,X2,X4,X4,X1)) & coll(X6,X0,X1) & perp(X6,X4,X0,X1) & coll(X5,X0,X2) & perp(X5,X4,X0,X2) & circle(X3,X0,X4,X7) & circle(X3,X0,X1,X2))),
  inference(flattening,[],[f319])).
fof(f319,plain,(
  ? [X0,X1,X2,X3,X4,X5,X6,X7] : ((~eqangle(X4,X5,X5,X6,X4,X2,X2,X1) | (~eqangle(X5,X4,X4,X6,X4,X1,X1,X2) & ~eqangle(X4,X5,X5,X6,X2,X4,X4,X1)) | (~eqangle(X5,X4,X4,X6,X4,X1,X1,X2) & ~eqangle(X4,X5,X5,X6,X4,X1,X1,X2)) | (~eqangle(X5,X4,X4,X6,X4,X2,X2,X1) & ~eqangle(X4,X5,X5,X6,X2,X4,X4,X1)) | (~eqangle(X5,X4,X4,X6,X4,X2,X2,X1) & ~eqangle(X4,X5,X5,X6,X4,X1,X1,X2)) | (~eqangle(X5,X4,X4,X6,X2,X4,X4,X1) & ~eqangle(X4,X5,X5,X6,X4,X2,X2,X1)) | ~eqangle(X5,X4,X4,X6,X2,X4,X4,X1)) & (coll(X6,X0,X1) & perp(X6,X4,X0,X1) & coll(X5,X0,X2) & perp(X5,X4,X0,X2) & circle(X3,X0,X4,X7) & circle(X3,X0,X1,X2)))),
  inference(ennf_transformation,[],[f167])).
fof(f167,plain,(
  ~! [X0,X1,X2,X3,X4,X5,X6,X7] : ((coll(X6,X0,X1) & perp(X6,X4,X0,X1) & coll(X5,X0,X2) & perp(X5,X4,X0,X2) & circle(X3,X0,X4,X7) & circle(X3,X0,X1,X2)) => (eqangle(X4,X5,X5,X6,X4,X2,X2,X1) & (eqangle(X5,X4,X4,X6,X4,X1,X1,X2) | eqangle(X4,X5,X5,X6,X2,X4,X4,X1)) & (eqangle(X5,X4,X4,X6,X4,X1,X1,X2) | eqangle(X4,X5,X5,X6,X4,X1,X1,X2)) & (eqangle(X5,X4,X4,X6,X4,X2,X2,X1) | eqangle(X4,X5,X5,X6,X2,X4,X4,X1)) & (eqangle(X5,X4,X4,X6,X4,X2,X2,X1) | eqangle(X4,X5,X5,X6,X4,X1,X1,X2)) & (eqangle(X5,X4,X4,X6,X2,X4,X4,X1) | eqangle(X4,X5,X5,X6,X4,X2,X2,X1)) & eqangle(X5,X4,X4,X6,X2,X4,X4,X1)))),
  inference(rectify,[],[f96])).
fof(f96,negated_conjecture,(
  ~! [X0,X1,X2,X3,X4,X5,X12,X18] : ((coll(X12,X0,X1) & perp(X12,X4,X0,X1) & coll(X5,X0,X2) & perp(X5,X4,X0,X2) & circle(X3,X0,X4,X18) & circle(X3,X0,X1,X2)) => (eqangle(X4,X5,X5,X12,X4,X2,X2,X1) & (eqangle(X5,X4,X4,X12,X4,X1,X1,X2) | eqangle(X4,X5,X5,X12,X2,X4,X4,X1)) & (eqangle(X5,X4,X4,X12,X4,X1,X1,X2) | eqangle(X4,X5,X5,X12,X4,X1,X1,X2)) & (eqangle(X5,X4,X4,X12,X4,X2,X2,X1) | eqangle(X4,X5,X5,X12,X2,X4,X4,X1)) & (eqangle(X5,X4,X4,X12,X4,X2,X2,X1) | eqangle(X4,X5,X5,X12,X4,X1,X1,X2)) & (eqangle(X5,X4,X4,X12,X2,X4,X4,X1) | eqangle(X4,X5,X5,X12,X4,X2,X2,X1)) & eqangle(X5,X4,X4,X12,X2,X4,X4,X1)))),
  inference(negated_conjecture,[],[f95])).
fof(f95,conjecture,(
  ! [X0,X1,X2,X3,X4,X5,X12,X18] : ((coll(X12,X0,X1) & perp(X12,X4,X0,X1) & coll(X5,X0,X2) & perp(X5,X4,X0,X2) & circle(X3,X0,X4,X18) & circle(X3,X0,X1,X2)) => (eqangle(X4,X5,X5,X12,X4,X2,X2,X1) & (eqangle(X5,X4,X4,X12,X4,X1,X1,X2) | eqangle(X4,X5,X5,X12,X2,X4,X4,X1)) & (eqangle(X5,X4,X4,X12,X4,X1,X1,X2) | eqangle(X4,X5,X5,X12,X4,X1,X1,X2)) & (eqangle(X5,X4,X4,X12,X4,X2,X2,X1) | eqangle(X4,X5,X5,X12,X2,X4,X4,X1)) & (eqangle(X5,X4,X4,X12,X4,X2,X2,X1) | eqangle(X4,X5,X5,X12,X4,X1,X1,X2)) & (eqangle(X5,X4,X4,X12,X2,X4,X4,X1) | eqangle(X4,X5,X5,X12,X4,X2,X2,X1)) & eqangle(X5,X4,X4,X12,X2,X4,X4,X1)))),
  file('./6_GDD_FULL/61-80/72.p',exemplo6GDDFULL618072)).
fof(f372,plain,(
  ( ! [X2,X0,X3,X1] : (~perp(X0,X1,X2,X3) | perp(X0,X1,X3,X2)) )),
  inference(cnf_transformation,[],[f176])).
fof(f176,plain,(
  ! [X0,X1,X2,X3] : (perp(X0,X1,X3,X2) | ~perp(X0,X1,X2,X3))),
  inference(ennf_transformation,[],[f7])).
fof(f7,axiom,(
  ! [X0,X1,X2,X3] : (perp(X0,X1,X2,X3) => perp(X0,X1,X3,X2))),
  file('./6_GDD_FULL/61-80/72.p',ruleD7)).
fof(f921,plain,(
  ( ! [X23,X22] : (~perp(X22,X23,sK27,sK26) | para(X22,X23,sK22,sK24)) )),
  inference(resolution,[],[f374,f485])).
fof(f374,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (~perp(X2,X3,X4,X5) | para(X0,X1,X4,X5) | ~perp(X0,X1,X2,X3)) )),
  inference(cnf_transformation,[],[f179])).
fof(f179,plain,(
  ! [X0,X1,X2,X3,X4,X5] : (para(X0,X1,X4,X5) | ~perp(X2,X3,X4,X5) | ~perp(X0,X1,X2,X3))),
  inference(flattening,[],[f178])).
fof(f178,plain,(
  ! [X0,X1,X2,X3,X4,X5] : (para(X0,X1,X4,X5) | (~perp(X2,X3,X4,X5) | ~perp(X0,X1,X2,X3)))),
  inference(ennf_transformation,[],[f9])).
fof(f9,axiom,(
  ! [X0,X1,X2,X3,X4,X5] : ((perp(X2,X3,X4,X5) & perp(X0,X1,X2,X3)) => para(X0,X1,X4,X5))),
  file('./6_GDD_FULL/61-80/72.p',ruleD9)).
fof(f19206,plain,(
  ( ! [X2] : (midp(sK24,X2,sK22)) ) | (~spl30_154 | ~spl30_178 | ~spl30_224)),
  inference(subsumption_resolution,[],[f19205,f14663])).
fof(f14663,plain,(
  ( ! [X3] : (midp(sK22,X3,sK23)) ) | (~spl30_154 | ~spl30_178)),
  inference(subsumption_resolution,[],[f14662,f9166])).
fof(f9166,plain,(
  ( ! [X14,X15] : (midp(X14,X15,X14)) ) | ~spl30_178),
  inference(resolution,[],[f9125,f376])).
fof(f9125,plain,(
  ( ! [X2,X3] : (midp(X2,X2,X3)) ) | ~spl30_178),
  inference(resolution,[],[f8950,f3264])).
fof(f3264,plain,(
  ( ! [X10,X8,X9] : (~midp(X10,X9,X8) | midp(X8,X9,X10)) ) | ~spl30_178),
  inference(subsumption_resolution,[],[f3258,f3113])).
fof(f3258,plain,(
  ( ! [X10,X8,X9] : (~coll(X8,X9,X10) | midp(X8,X9,X10) | ~midp(X10,X9,X8)) ) | ~spl30_178),
  inference(resolution,[],[f3052,f411])).
fof(f3052,plain,(
  ( ! [X21,X20] : (para(X20,X21,X21,X20)) ) | ~spl30_178),
  inference(resolution,[],[f2168,f369])).
fof(f369,plain,(
  ( ! [X2,X0,X3,X1] : (~para(X0,X1,X2,X3) | para(X0,X1,X3,X2)) )),
  inference(cnf_transformation,[],[f172])).
fof(f172,plain,(
  ! [X0,X1,X2,X3] : (para(X0,X1,X3,X2) | ~para(X0,X1,X2,X3))),
  inference(ennf_transformation,[],[f4])).
fof(f4,axiom,(
  ! [X0,X1,X2,X3] : (para(X0,X1,X2,X3) => para(X0,X1,X3,X2))),
  file('./6_GDD_FULL/61-80/72.p',ruleD4)).
fof(f8950,plain,(
  ( ! [X12,X11] : (midp(X11,X12,X12)) ) | ~spl30_178),
  inference(resolution,[],[f8418,f3055])).
fof(f3055,plain,(
  ( ! [X10,X8,X9] : (~midp(X8,X10,X10) | midp(X8,X9,X9)) ) | ~spl30_178),
  inference(subsumption_resolution,[],[f3048,f2168])).
fof(f3048,plain,(
  ( ! [X10,X8,X9] : (midp(X8,X9,X9) | ~para(X10,X9,X10,X9) | ~midp(X8,X10,X10)) ) | ~spl30_178),
  inference(resolution,[],[f2168,f430])).
fof(f430,plain,(
  ( ! [X4,X2,X0,X3,X1] : (~para(X0,X3,X1,X2) | midp(X4,X2,X3) | ~para(X0,X2,X1,X3) | ~midp(X4,X0,X1)) )),
  inference(cnf_transformation,[],[f264])).
fof(f264,plain,(
  ! [X0,X1,X2,X3,X4] : (midp(X4,X2,X3) | ~para(X0,X3,X1,X2) | ~para(X0,X2,X1,X3) | ~midp(X4,X0,X1))),
  inference(flattening,[],[f263])).
fof(f263,plain,(
  ! [X0,X1,X2,X3,X4] : (midp(X4,X2,X3) | (~para(X0,X3,X1,X2) | ~para(X0,X2,X1,X3) | ~midp(X4,X0,X1)))),
  inference(ennf_transformation,[],[f143])).
fof(f143,plain,(
  ! [X0,X1,X2,X3,X4] : ((para(X0,X3,X1,X2) & para(X0,X2,X1,X3) & midp(X4,X0,X1)) => midp(X4,X2,X3))),
  inference(rectify,[],[f65])).
fof(f65,axiom,(
  ! [X0,X1,X2,X3,X6] : ((para(X0,X3,X1,X2) & para(X0,X2,X1,X3) & midp(X6,X0,X1)) => midp(X6,X2,X3))),
  file('./6_GDD_FULL/61-80/72.p',ruleD64)).
fof(f8418,plain,(
  ( ! [X14] : (midp(X14,X14,X14)) ) | ~spl30_178),
  inference(subsumption_resolution,[],[f8389,f3086])).
fof(f8389,plain,(
  ( ! [X14] : (~coll(X14,X14,X14) | midp(X14,X14,X14)) ) | ~spl30_178),
  inference(resolution,[],[f8312,f433])).
fof(f433,plain,(
  ( ! [X2,X0,X1] : (~cong(X0,X1,X0,X2) | ~coll(X0,X1,X2) | midp(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f269])).
fof(f269,plain,(
  ! [X0,X1,X2] : (midp(X0,X1,X2) | ~coll(X0,X1,X2) | ~cong(X0,X1,X0,X2))),
  inference(flattening,[],[f268])).
fof(f268,plain,(
  ! [X0,X1,X2] : (midp(X0,X1,X2) | (~coll(X0,X1,X2) | ~cong(X0,X1,X0,X2)))),
  inference(ennf_transformation,[],[f68])).
fof(f68,axiom,(
  ! [X0,X1,X2] : ((coll(X0,X1,X2) & cong(X0,X1,X0,X2)) => midp(X0,X1,X2))),
  file('./6_GDD_FULL/61-80/72.p',ruleD67)).
fof(f8312,plain,(
  ( ! [X35] : (cong(X35,X35,X35,X35)) ) | ~spl30_178),
  inference(subsumption_resolution,[],[f8308,f8280])).
fof(f8280,plain,(
  ( ! [X6,X4,X5] : (cyclic(X4,X4,X5,X6)) ) | ~spl30_178),
  inference(resolution,[],[f2518,f2168])).
fof(f2518,plain,(
  ( ! [X94,X92,X93] : (~para(X92,X93,X92,X93) | cyclic(X93,X93,X92,X94)) )),
  inference(subsumption_resolution,[],[f2366,f2367])).
fof(f2367,plain,(
  ( ! [X97,X95,X96] : (~para(X95,X96,X95,X96) | coll(X95,X97,X96) | cyclic(X96,X96,X95,X97)) )),
  inference(resolution,[],[f993,f407])).
fof(f407,plain,(
  ( ! [X2,X0,X3,X1] : (~eqangle(X2,X0,X2,X1,X3,X0,X3,X1) | coll(X2,X3,X0) | cyclic(X0,X1,X2,X3)) )),
  inference(cnf_transformation,[],[f222])).
fof(f222,plain,(
  ! [X0,X1,X2,X3] : (cyclic(X0,X1,X2,X3) | coll(X2,X3,X0) | ~eqangle(X2,X0,X2,X1,X3,X0,X3,X1))),
  inference(flattening,[],[f221])).
fof(f221,plain,(
  ! [X0,X1,X2,X3] : (cyclic(X0,X1,X2,X3) | (coll(X2,X3,X0) | ~eqangle(X2,X0,X2,X1,X3,X0,X3,X1)))),
  inference(ennf_transformation,[],[f121])).
fof(f121,plain,(
  ! [X0,X1,X2,X3] : ((~coll(X2,X3,X0) & eqangle(X2,X0,X2,X1,X3,X0,X3,X1)) => cyclic(X0,X1,X2,X3))),
  inference(rectify,[],[f42])).
fof(f42,axiom,(
  ! [X0,X1,X8,X9] : ((~coll(X8,X9,X0) & eqangle(X8,X0,X8,X1,X9,X0,X9,X1)) => cyclic(X0,X1,X8,X9))),
  file('./6_GDD_FULL/61-80/72.p',ruleD42a)).
fof(f993,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (eqangle(X0,X1,X2,X3,X4,X5,X4,X5) | ~para(X0,X1,X2,X3)) )),
  inference(resolution,[],[f386,f405])).
fof(f405,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (eqangle(X0,X1,X4,X5,X2,X3,X4,X5) | ~para(X0,X1,X2,X3)) )),
  inference(cnf_transformation,[],[f219])).
fof(f219,plain,(
  ! [X0,X1,X2,X3,X4,X5] : (eqangle(X0,X1,X4,X5,X2,X3,X4,X5) | ~para(X0,X1,X2,X3))),
  inference(ennf_transformation,[],[f119])).
fof(f119,plain,(
  ! [X0,X1,X2,X3,X4,X5] : (para(X0,X1,X2,X3) => eqangle(X0,X1,X4,X5,X2,X3,X4,X5))),
  inference(rectify,[],[f40])).
fof(f40,axiom,(
  ! [X0,X1,X2,X3,X8,X9] : (para(X0,X1,X2,X3) => eqangle(X0,X1,X8,X9,X2,X3,X8,X9))),
  file('./6_GDD_FULL/61-80/72.p',ruleD40)).
fof(f386,plain,(
  ( ! [X6,X4,X2,X0,X7,X5,X3,X1] : (~eqangle(X0,X1,X2,X3,X4,X5,X6,X7) | eqangle(X0,X1,X4,X5,X2,X3,X6,X7)) )),
  inference(cnf_transformation,[],[f195])).
fof(f195,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6,X7] : (eqangle(X0,X1,X4,X5,X2,X3,X6,X7) | ~eqangle(X0,X1,X2,X3,X4,X5,X6,X7))),
  inference(ennf_transformation,[],[f103])).
fof(f103,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6,X7] : (eqangle(X0,X1,X2,X3,X4,X5,X6,X7) => eqangle(X0,X1,X4,X5,X2,X3,X6,X7))),
  inference(rectify,[],[f21])).
fof(f21,axiom,(
  ! [X0,X1,X2,X3,X8,X9,X10,X11] : (eqangle(X0,X1,X2,X3,X8,X9,X10,X11) => eqangle(X0,X1,X8,X9,X2,X3,X10,X11))),
  file('./6_GDD_FULL/61-80/72.p',ruleD21)).
fof(f2366,plain,(
  ( ! [X94,X92,X93] : (~para(X92,X93,X92,X93) | ~coll(X92,X94,X93) | cyclic(X93,X93,X92,X94)) )),
  inference(resolution,[],[f993,f408])).
fof(f408,plain,(
  ( ! [X2,X0,X3,X1] : (~eqangle(X2,X0,X2,X1,X3,X0,X3,X1) | ~coll(X2,X3,X1) | cyclic(X0,X1,X2,X3)) )),
  inference(cnf_transformation,[],[f224])).
fof(f224,plain,(
  ! [X0,X1,X2,X3] : (cyclic(X0,X1,X2,X3) | ~coll(X2,X3,X1) | ~eqangle(X2,X0,X2,X1,X3,X0,X3,X1))),
  inference(flattening,[],[f223])).
fof(f223,plain,(
  ! [X0,X1,X2,X3] : (cyclic(X0,X1,X2,X3) | (~coll(X2,X3,X1) | ~eqangle(X2,X0,X2,X1,X3,X0,X3,X1)))),
  inference(ennf_transformation,[],[f122])).
fof(f122,plain,(
  ! [X0,X1,X2,X3] : ((coll(X2,X3,X1) & eqangle(X2,X0,X2,X1,X3,X0,X3,X1)) => cyclic(X0,X1,X2,X3))),
  inference(rectify,[],[f43])).
fof(f43,axiom,(
  ! [X0,X1,X8,X9] : ((coll(X8,X9,X1) & eqangle(X8,X0,X8,X1,X9,X0,X9,X1)) => cyclic(X0,X1,X8,X9))),
  file('./6_GDD_FULL/61-80/72.p',ruleD42b)).
fof(f8308,plain,(
  ( ! [X35,X36] : (cong(X35,X35,X35,X35) | ~cyclic(X35,X35,X36,X35)) ) | ~spl30_178),
  inference(duplicate_literal_removal,[],[f8295])).
fof(f8295,plain,(
  ( ! [X35,X36] : (cong(X35,X35,X35,X35) | ~cyclic(X35,X35,X36,X35) | ~cyclic(X35,X35,X36,X35)) ) | ~spl30_178),
  inference(resolution,[],[f8280,f1451])).
fof(f1451,plain,(
  ( ! [X10,X8,X7,X9] : (~cyclic(X7,X8,X9,X10) | cong(X7,X8,X7,X8) | ~cyclic(X7,X8,X9,X8) | ~cyclic(X7,X8,X9,X7)) )),
  inference(duplicate_literal_removal,[],[f1445])).
fof(f1445,plain,(
  ( ! [X10,X8,X7,X9] : (cong(X7,X8,X7,X8) | ~cyclic(X7,X8,X9,X10) | ~cyclic(X7,X8,X9,X8) | ~cyclic(X7,X8,X9,X7) | ~cyclic(X7,X8,X9,X10)) )),
  inference(resolution,[],[f409,f406])).
fof(f406,plain,(
  ( ! [X2,X0,X3,X1] : (eqangle(X2,X0,X2,X1,X3,X0,X3,X1) | ~cyclic(X0,X1,X2,X3)) )),
  inference(cnf_transformation,[],[f220])).
fof(f220,plain,(
  ! [X0,X1,X2,X3] : (eqangle(X2,X0,X2,X1,X3,X0,X3,X1) | ~cyclic(X0,X1,X2,X3))),
  inference(ennf_transformation,[],[f120])).
fof(f120,plain,(
  ! [X0,X1,X2,X3] : (cyclic(X0,X1,X2,X3) => eqangle(X2,X0,X2,X1,X3,X0,X3,X1))),
  inference(rectify,[],[f41])).
fof(f41,axiom,(
  ! [X0,X1,X8,X9] : (cyclic(X0,X1,X8,X9) => eqangle(X8,X0,X8,X1,X9,X0,X9,X1))),
  file('./6_GDD_FULL/61-80/72.p',ruleD41)).
fof(f409,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (~eqangle(X2,X0,X2,X1,X5,X3,X5,X4) | cong(X0,X1,X3,X4) | ~cyclic(X0,X1,X2,X5) | ~cyclic(X0,X1,X2,X4) | ~cyclic(X0,X1,X2,X3)) )),
  inference(cnf_transformation,[],[f226])).
fof(f226,plain,(
  ! [X0,X1,X2,X3,X4,X5] : (cong(X0,X1,X3,X4) | ~eqangle(X2,X0,X2,X1,X5,X3,X5,X4) | ~cyclic(X0,X1,X2,X5) | ~cyclic(X0,X1,X2,X4) | ~cyclic(X0,X1,X2,X3))),
  inference(flattening,[],[f225])).
fof(f225,plain,(
  ! [X0,X1,X2,X3,X4,X5] : (cong(X0,X1,X3,X4) | (~eqangle(X2,X0,X2,X1,X5,X3,X5,X4) | ~cyclic(X0,X1,X2,X5) | ~cyclic(X0,X1,X2,X4) | ~cyclic(X0,X1,X2,X3)))),
  inference(ennf_transformation,[],[f123])).
fof(f123,plain,(
  ! [X0,X1,X2,X3,X4,X5] : ((eqangle(X2,X0,X2,X1,X5,X3,X5,X4) & cyclic(X0,X1,X2,X5) & cyclic(X0,X1,X2,X4) & cyclic(X0,X1,X2,X3)) => cong(X0,X1,X3,X4))),
  inference(rectify,[],[f44])).
fof(f44,axiom,(
  ! [X0,X1,X2,X8,X9,X14] : ((eqangle(X2,X0,X2,X1,X14,X8,X14,X9) & cyclic(X0,X1,X2,X14) & cyclic(X0,X1,X2,X9) & cyclic(X0,X1,X2,X8)) => cong(X0,X1,X8,X9))),
  file('./6_GDD_FULL/61-80/72.p',ruleD43)).
fof(f14662,plain,(
  ( ! [X3] : (midp(sK22,X3,sK23) | ~midp(sK24,X3,sK24)) ) | (~spl30_154 | ~spl30_178)),
  inference(subsumption_resolution,[],[f14653,f3113])).
fof(f14653,plain,(
  ( ! [X3] : (~coll(sK22,X3,sK23) | midp(sK22,X3,sK23) | ~midp(sK24,X3,sK24)) ) | ~spl30_154),
  inference(resolution,[],[f1909,f411])).
fof(f1909,plain,(
  para(sK24,sK22,sK24,sK23) | ~spl30_154),
  inference(avatar_component_clause,[],[f1908])).
fof(f1908,plain,(
  spl30_154 <=> para(sK24,sK22,sK24,sK23)),
  introduced(avatar_definition,[new_symbols(naming,[spl30_154])])).
fof(f19205,plain,(
  ( ! [X2] : (midp(sK24,X2,sK22) | ~midp(sK22,X2,sK23)) ) | (~spl30_178 | ~spl30_224)),
  inference(subsumption_resolution,[],[f19194,f3113])).
fof(f19194,plain,(
  ( ! [X2] : (~coll(sK24,X2,sK22) | midp(sK24,X2,sK22) | ~midp(sK22,X2,sK23)) ) | ~spl30_224),
  inference(resolution,[],[f12405,f411])).
fof(f12405,plain,(
  para(sK22,sK24,sK23,sK22) | ~spl30_224),
  inference(resolution,[],[f9738,f370])).
fof(f370,plain,(
  ( ! [X2,X0,X3,X1] : (~para(X0,X1,X2,X3) | para(X2,X3,X0,X1)) )),
  inference(cnf_transformation,[],[f173])).
fof(f173,plain,(
  ! [X0,X1,X2,X3] : (para(X2,X3,X0,X1) | ~para(X0,X1,X2,X3))),
  inference(ennf_transformation,[],[f5])).
fof(f5,axiom,(
  ! [X0,X1,X2,X3] : (para(X0,X1,X2,X3) => para(X2,X3,X0,X1))),
  file('./6_GDD_FULL/61-80/72.p',ruleD5)).
fof(f9738,plain,(
  para(sK23,sK22,sK22,sK24) | ~spl30_224),
  inference(resolution,[],[f3497,f2043])).
fof(f2043,plain,(
  ( ! [X2,X0,X3,X1] : (~para(X0,X1,X2,X3) | para(X1,X0,X2,X3)) )),
  inference(resolution,[],[f975,f404])).
fof(f404,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (~eqangle(X0,X1,X4,X5,X2,X3,X4,X5) | para(X0,X1,X2,X3)) )),
  inference(cnf_transformation,[],[f218])).
fof(f218,plain,(
  ! [X0,X1,X2,X3,X4,X5] : (para(X0,X1,X2,X3) | ~eqangle(X0,X1,X4,X5,X2,X3,X4,X5))),
  inference(ennf_transformation,[],[f118])).
fof(f118,plain,(
  ! [X0,X1,X2,X3,X4,X5] : (eqangle(X0,X1,X4,X5,X2,X3,X4,X5) => para(X0,X1,X2,X3))),
  inference(rectify,[],[f39])).
fof(f39,axiom,(
  ! [X0,X1,X2,X3,X8,X9] : (eqangle(X0,X1,X8,X9,X2,X3,X8,X9) => para(X0,X1,X2,X3))),
  file('./6_GDD_FULL/61-80/72.p',ruleD39)).
fof(f975,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (eqangle(X0,X1,X2,X3,X4,X5,X2,X3) | ~para(X1,X0,X4,X5)) )),
  inference(resolution,[],[f383,f405])).
fof(f383,plain,(
  ( ! [X6,X4,X2,X0,X7,X5,X3,X1] : (~eqangle(X0,X1,X2,X3,X4,X5,X6,X7) | eqangle(X1,X0,X2,X3,X4,X5,X6,X7)) )),
  inference(cnf_transformation,[],[f192])).
fof(f192,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6,X7] : (eqangle(X1,X0,X2,X3,X4,X5,X6,X7) | ~eqangle(X0,X1,X2,X3,X4,X5,X6,X7))),
  inference(ennf_transformation,[],[f100])).
fof(f100,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6,X7] : (eqangle(X0,X1,X2,X3,X4,X5,X6,X7) => eqangle(X1,X0,X2,X3,X4,X5,X6,X7))),
  inference(rectify,[],[f18])).
fof(f18,axiom,(
  ! [X0,X1,X2,X3,X8,X9,X10,X11] : (eqangle(X0,X1,X2,X3,X8,X9,X10,X11) => eqangle(X1,X0,X2,X3,X8,X9,X10,X11))),
  file('./6_GDD_FULL/61-80/72.p',ruleD18)).
fof(f3497,plain,(
  para(sK22,sK23,sK22,sK24) | ~spl30_224),
  inference(avatar_component_clause,[],[f3496])).
fof(f3496,plain,(
  spl30_224 <=> para(sK22,sK23,sK22,sK24)),
  introduced(avatar_definition,[new_symbols(naming,[spl30_224])])).
fof(f24412,plain,(
  ( ! [X8,X9] : (midp(X8,X9,sK22) | ~midp(sK22,X9,sK24)) ) | (~spl30_178 | ~spl30_456)),
  inference(subsumption_resolution,[],[f24401,f3113])).
fof(f24401,plain,(
  ( ! [X8,X9] : (~coll(X8,X9,sK22) | midp(X8,X9,sK22) | ~midp(sK22,X9,sK24)) ) | ~spl30_456),
  inference(resolution,[],[f11084,f411])).
fof(f11084,plain,(
  ( ! [X9] : (para(sK22,X9,sK24,sK22)) ) | ~spl30_456),
  inference(avatar_component_clause,[],[f11083])).
fof(f11083,plain,(
  spl30_456 <=> ! [X9] : para(sK22,X9,sK24,sK22)),
  introduced(avatar_definition,[new_symbols(naming,[spl30_456])])).
fof(f24662,plain,(
  ( ! [X10,X8,X11,X9] : (para(X8,X9,X10,X11) | ~midp(X8,sK22,X10)) ) | (~spl30_154 | ~spl30_178 | ~spl30_224 | ~spl30_456)),
  inference(resolution,[],[f24482,f410])).
fof(f410,plain,(
  ( ! [X4,X2,X0,X3,X1] : (~midp(X4,X0,X2) | para(X3,X4,X1,X2) | ~midp(X3,X0,X1)) )),
  inference(cnf_transformation,[],[f228])).
fof(f228,plain,(
  ! [X0,X1,X2,X3,X4] : (para(X3,X4,X1,X2) | ~midp(X4,X0,X2) | ~midp(X3,X0,X1))),
  inference(flattening,[],[f227])).
fof(f227,plain,(
  ! [X0,X1,X2,X3,X4] : (para(X3,X4,X1,X2) | (~midp(X4,X0,X2) | ~midp(X3,X0,X1)))),
  inference(ennf_transformation,[],[f124])).
fof(f124,plain,(
  ! [X0,X1,X2,X3,X4] : ((midp(X4,X0,X2) & midp(X3,X0,X1)) => para(X3,X4,X1,X2))),
  inference(rectify,[],[f45])).
fof(f45,axiom,(
  ! [X0,X1,X2,X4,X5] : ((midp(X5,X0,X2) & midp(X4,X0,X1)) => para(X4,X5,X1,X2))),
  file('./6_GDD_FULL/61-80/72.p',ruleD44)).
fof(f25133,plain,(
  ( ! [X61,X59,X57,X62,X60,X58,X56,X63] : (~para(X60,X61,X62,X63) | eqangle(X58,X59,X62,X63,X56,X57,X60,X61)) ) | (~spl30_154 | ~spl30_178 | ~spl30_224 | ~spl30_456)),
  inference(subsumption_resolution,[],[f25088,f24685])).
fof(f25088,plain,(
  ( ! [X61,X59,X57,X62,X60,X58,X56,X63] : (~para(X56,X57,X58,X59) | ~para(X60,X61,X62,X63) | eqangle(X58,X59,X62,X63,X56,X57,X60,X61)) )),
  inference(resolution,[],[f5123,f385])).
fof(f385,plain,(
  ( ! [X6,X4,X2,X0,X7,X5,X3,X1] : (~eqangle(X0,X1,X2,X3,X4,X5,X6,X7) | eqangle(X4,X5,X6,X7,X0,X1,X2,X3)) )),
  inference(cnf_transformation,[],[f194])).
fof(f194,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6,X7] : (eqangle(X4,X5,X6,X7,X0,X1,X2,X3) | ~eqangle(X0,X1,X2,X3,X4,X5,X6,X7))),
  inference(ennf_transformation,[],[f102])).
fof(f102,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6,X7] : (eqangle(X0,X1,X2,X3,X4,X5,X6,X7) => eqangle(X4,X5,X6,X7,X0,X1,X2,X3))),
  inference(rectify,[],[f20])).
fof(f20,axiom,(
  ! [X0,X1,X2,X3,X8,X9,X10,X11] : (eqangle(X0,X1,X2,X3,X8,X9,X10,X11) => eqangle(X8,X9,X10,X11,X0,X1,X2,X3))),
  file('./6_GDD_FULL/61-80/72.p',ruleD20)).
fof(f5123,plain,(
  ( ! [X80,X87,X85,X83,X81,X86,X84,X82] : (eqangle(X80,X81,X82,X83,X84,X85,X86,X87) | ~para(X80,X81,X84,X85) | ~para(X82,X83,X86,X87)) )),
  inference(resolution,[],[f1362,f979])).
fof(f979,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (eqangle(X0,X1,X2,X3,X0,X1,X4,X5) | ~para(X2,X3,X4,X5)) )),
  inference(resolution,[],[f384,f405])).
fof(f384,plain,(
  ( ! [X6,X4,X2,X0,X7,X5,X3,X1] : (~eqangle(X0,X1,X2,X3,X4,X5,X6,X7) | eqangle(X2,X3,X0,X1,X6,X7,X4,X5)) )),
  inference(cnf_transformation,[],[f193])).
fof(f193,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6,X7] : (eqangle(X2,X3,X0,X1,X6,X7,X4,X5) | ~eqangle(X0,X1,X2,X3,X4,X5,X6,X7))),
  inference(ennf_transformation,[],[f101])).
fof(f101,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6,X7] : (eqangle(X0,X1,X2,X3,X4,X5,X6,X7) => eqangle(X2,X3,X0,X1,X6,X7,X4,X5))),
  inference(rectify,[],[f19])).
fof(f19,axiom,(
  ! [X0,X1,X2,X3,X8,X9,X10,X11] : (eqangle(X0,X1,X2,X3,X8,X9,X10,X11) => eqangle(X2,X3,X0,X1,X10,X11,X8,X9))),
  file('./6_GDD_FULL/61-80/72.p',ruleD19)).
fof(f1362,plain,(
  ( ! [X6,X4,X2,X0,X8,X7,X5,X3,X1,X9] : (~eqangle(X0,X1,X2,X3,X8,X9,X6,X7) | eqangle(X0,X1,X2,X3,X4,X5,X6,X7) | ~para(X8,X9,X4,X5)) )),
  inference(resolution,[],[f387,f405])).
fof(f387,plain,(
  ( ! [X6,X4,X2,X0,X10,X8,X7,X5,X3,X1,X11,X9] : (~eqangle(X4,X5,X6,X7,X8,X9,X10,X11) | eqangle(X0,X1,X2,X3,X8,X9,X10,X11) | ~eqangle(X0,X1,X2,X3,X4,X5,X6,X7)) )),
  inference(cnf_transformation,[],[f197])).
fof(f197,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11] : (eqangle(X0,X1,X2,X3,X8,X9,X10,X11) | ~eqangle(X4,X5,X6,X7,X8,X9,X10,X11) | ~eqangle(X0,X1,X2,X3,X4,X5,X6,X7))),
  inference(flattening,[],[f196])).
fof(f196,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11] : (eqangle(X0,X1,X2,X3,X8,X9,X10,X11) | (~eqangle(X4,X5,X6,X7,X8,X9,X10,X11) | ~eqangle(X0,X1,X2,X3,X4,X5,X6,X7)))),
  inference(ennf_transformation,[],[f104])).
fof(f104,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11] : ((eqangle(X4,X5,X6,X7,X8,X9,X10,X11) & eqangle(X0,X1,X2,X3,X4,X5,X6,X7)) => eqangle(X0,X1,X2,X3,X8,X9,X10,X11))),
  inference(rectify,[],[f22])).
fof(f22,axiom,(
  ! [X0,X1,X2,X3,X8,X9,X10,X11,X4,X5,X12,X13] : ((eqangle(X8,X9,X10,X11,X4,X5,X12,X13) & eqangle(X0,X1,X2,X3,X8,X9,X10,X11)) => eqangle(X0,X1,X2,X3,X4,X5,X12,X13))),
  file('./6_GDD_FULL/61-80/72.p',ruleD22)).
fof(f549,plain,(
  ~eqangle(sK26,sK27,sK27,sK28,sK26,sK23,sK23,sK24) | ~spl30_13),
  inference(avatar_component_clause,[],[f548])).
fof(f548,plain,(
  spl30_13 <=> ~eqangle(sK26,sK27,sK27,sK28,sK26,sK23,sK23,sK24)),
  introduced(avatar_definition,[new_symbols(naming,[spl30_13])])).
fof(f25155,plain,(
  spl30_11 | ~spl30_154 | ~spl30_178 | ~spl30_224 | ~spl30_456),
  inference(avatar_contradiction_clause,[],[f25154])).
fof(f25154,plain,(
  $false | (~spl30_11 | ~spl30_154 | ~spl30_178 | ~spl30_224 | ~spl30_456)),
  inference(subsumption_resolution,[],[f540,f25134])).
fof(f540,plain,(
  ~eqangle(sK26,sK27,sK27,sK28,sK24,sK26,sK26,sK23) | ~spl30_11),
  inference(avatar_component_clause,[],[f539])).
fof(f539,plain,(
  spl30_11 <=> ~eqangle(sK26,sK27,sK27,sK28,sK24,sK26,sK26,sK23)),
  introduced(avatar_definition,[new_symbols(naming,[spl30_11])])).
fof(f25153,plain,(
  spl30_7 | ~spl30_154 | ~spl30_178 | ~spl30_224 | ~spl30_456),
  inference(avatar_contradiction_clause,[],[f25152])).
fof(f25152,plain,(
  $false | (~spl30_7 | ~spl30_154 | ~spl30_178 | ~spl30_224 | ~spl30_456)),
  inference(subsumption_resolution,[],[f526,f25134])).
fof(f526,plain,(
  ~eqangle(sK26,sK27,sK27,sK28,sK26,sK24,sK24,sK23) | ~spl30_7),
  inference(avatar_component_clause,[],[f525])).
fof(f525,plain,(
  spl30_7 <=> ~eqangle(sK26,sK27,sK27,sK28,sK26,sK24,sK24,sK23)),
  introduced(avatar_definition,[new_symbols(naming,[spl30_7])])).
fof(f25121,plain,(
  spl30_1 | ~spl30_154 | ~spl30_178 | ~spl30_224 | ~spl30_456),
  inference(avatar_contradiction_clause,[],[f25120])).
fof(f25120,plain,(
  $false | (~spl30_1 | ~spl30_154 | ~spl30_178 | ~spl30_224 | ~spl30_456)),
  inference(subsumption_resolution,[],[f25119,f24685])).
fof(f25119,plain,(
  ~para(sK26,sK28,sK26,sK23) | (~spl30_1 | ~spl30_178)),
  inference(subsumption_resolution,[],[f25080,f22568])).
fof(f22568,plain,(
  ( ! [X35,X36,X34] : (para(X34,X35,X36,X35)) ) | ~spl30_178),
  inference(subsumption_resolution,[],[f22549,f10134])).
fof(f10134,plain,(
  ( ! [X26,X24,X23,X25] : (circle(X23,X25,X26,X24)) ) | ~spl30_178),
  inference(subsumption_resolution,[],[f10133,f10051])).
fof(f10051,plain,(
  ( ! [X4,X2,X3] : (cong(X2,X3,X2,X4)) ) | ~spl30_178),
  inference(subsumption_resolution,[],[f10028,f9125])).
fof(f10028,plain,(
  ( ! [X4,X2,X3] : (cong(X2,X3,X2,X4) | ~midp(X2,X2,X3)) ) | ~spl30_178),
  inference(resolution,[],[f2028,f9125])).
fof(f2028,plain,(
  ( ! [X2,X0,X3,X1] : (~midp(X0,X3,X2) | cong(X0,X1,X0,X2) | ~midp(X0,X3,X1)) )),
  inference(resolution,[],[f947,f602])).
fof(f602,plain,(
  ( ! [X2,X0,X1] : (cong(X0,X1,X0,X2) | ~midp(X0,X2,X1)) )),
  inference(resolution,[],[f389,f434])).
fof(f434,plain,(
  ( ! [X2,X0,X1] : (cong(X0,X1,X0,X2) | ~midp(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f270])).
fof(f270,plain,(
  ! [X0,X1,X2] : (cong(X0,X1,X0,X2) | ~midp(X0,X1,X2))),
  inference(ennf_transformation,[],[f69])).
fof(f69,axiom,(
  ! [X0,X1,X2] : (midp(X0,X1,X2) => cong(X0,X1,X0,X2))),
  file('./6_GDD_FULL/61-80/72.p',ruleD68)).
fof(f389,plain,(
  ( ! [X2,X0,X3,X1] : (~cong(X0,X1,X2,X3) | cong(X2,X3,X0,X1)) )),
  inference(cnf_transformation,[],[f199])).
fof(f199,plain,(
  ! [X0,X1,X2,X3] : (cong(X2,X3,X0,X1) | ~cong(X0,X1,X2,X3))),
  inference(ennf_transformation,[],[f24])).
fof(f24,axiom,(
  ! [X0,X1,X2,X3] : (cong(X0,X1,X2,X3) => cong(X2,X3,X0,X1))),
  file('./6_GDD_FULL/61-80/72.p',ruleD24)).
fof(f947,plain,(
  ( ! [X4,X2,X0,X3,X1] : (~cong(X0,X1,X2,X4) | cong(X0,X1,X2,X3) | ~midp(X2,X4,X3)) )),
  inference(resolution,[],[f390,f434])).
fof(f390,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (~cong(X2,X3,X4,X5) | cong(X0,X1,X4,X5) | ~cong(X0,X1,X2,X3)) )),
  inference(cnf_transformation,[],[f201])).
fof(f201,plain,(
  ! [X0,X1,X2,X3,X4,X5] : (cong(X0,X1,X4,X5) | ~cong(X2,X3,X4,X5) | ~cong(X0,X1,X2,X3))),
  inference(flattening,[],[f200])).
fof(f200,plain,(
  ! [X0,X1,X2,X3,X4,X5] : (cong(X0,X1,X4,X5) | (~cong(X2,X3,X4,X5) | ~cong(X0,X1,X2,X3)))),
  inference(ennf_transformation,[],[f25])).
fof(f25,axiom,(
  ! [X0,X1,X2,X3,X4,X5] : ((cong(X2,X3,X4,X5) & cong(X0,X1,X2,X3)) => cong(X0,X1,X4,X5))),
  file('./6_GDD_FULL/61-80/72.p',ruleD25)).
fof(f10133,plain,(
  ( ! [X26,X24,X23,X25] : (circle(X23,X25,X26,X24) | ~cong(X23,X25,X23,X26)) ) | ~spl30_178),
  inference(subsumption_resolution,[],[f10132,f9166])).
fof(f10132,plain,(
  ( ! [X26,X24,X23,X25] : (~midp(X23,X25,X23) | circle(X23,X25,X26,X24) | ~cong(X23,X25,X23,X26)) ) | ~spl30_178),
  inference(subsumption_resolution,[],[f10100,f9125])).
fof(f10100,plain,(
  ( ! [X26,X24,X23,X25] : (~midp(X23,X23,X24) | ~midp(X23,X25,X23) | circle(X23,X25,X26,X24) | ~cong(X23,X25,X23,X26)) )),
  inference(resolution,[],[f2031,f377])).
fof(f377,plain,(
  ( ! [X2,X0,X3,X1] : (~cong(X3,X0,X3,X2) | circle(X3,X0,X1,X2) | ~cong(X3,X0,X3,X1)) )),
  inference(cnf_transformation,[],[f184])).
fof(f184,plain,(
  ! [X0,X1,X2,X3] : (circle(X3,X0,X1,X2) | ~cong(X3,X0,X3,X2) | ~cong(X3,X0,X3,X1))),
  inference(flattening,[],[f183])).
fof(f183,plain,(
  ! [X0,X1,X2,X3] : (circle(X3,X0,X1,X2) | (~cong(X3,X0,X3,X2) | ~cong(X3,X0,X3,X1)))),
  inference(ennf_transformation,[],[f98])).
fof(f98,plain,(
  ! [X0,X1,X2,X3] : ((cong(X3,X0,X3,X2) & cong(X3,X0,X3,X1)) => circle(X3,X0,X1,X2))),
  inference(rectify,[],[f12])).
fof(f12,axiom,(
  ! [X0,X1,X2,X7] : ((cong(X7,X0,X7,X2) & cong(X7,X0,X7,X1)) => circle(X7,X0,X1,X2))),
  file('./6_GDD_FULL/61-80/72.p',ruleD12)).
fof(f2031,plain,(
  ( ! [X14,X12,X15,X13] : (cong(X12,X13,X14,X15) | ~midp(X14,X12,X15) | ~midp(X12,X13,X14)) )),
  inference(resolution,[],[f947,f599])).
fof(f599,plain,(
  ( ! [X2,X0,X1] : (cong(X0,X1,X2,X0) | ~midp(X0,X1,X2)) )),
  inference(resolution,[],[f388,f434])).
fof(f388,plain,(
  ( ! [X2,X0,X3,X1] : (~cong(X0,X1,X2,X3) | cong(X0,X1,X3,X2)) )),
  inference(cnf_transformation,[],[f198])).
fof(f198,plain,(
  ! [X0,X1,X2,X3] : (cong(X0,X1,X3,X2) | ~cong(X0,X1,X2,X3))),
  inference(ennf_transformation,[],[f23])).
fof(f23,axiom,(
  ! [X0,X1,X2,X3] : (cong(X0,X1,X2,X3) => cong(X0,X1,X3,X2))),
  file('./6_GDD_FULL/61-80/72.p',ruleD23)).
fof(f22549,plain,(
  ( ! [X37,X35,X38,X36,X34] : (para(X34,X35,X36,X35) | ~circle(X36,X35,X37,X38)) ) | ~spl30_178),
  inference(resolution,[],[f2008,f16318])).
fof(f16318,plain,(
  ( ! [X37,X38,X36] : (perp(X37,X38,X36,X38)) ) | ~spl30_178),
  inference(subsumption_resolution,[],[f16317,f9166])).
fof(f16317,plain,(
  ( ! [X37,X38,X36] : (perp(X37,X38,X36,X38) | ~midp(X36,X37,X36)) ) | ~spl30_178),
  inference(subsumption_resolution,[],[f16316,f9125])).
fof(f16316,plain,(
  ( ! [X37,X38,X36] : (perp(X37,X38,X36,X38) | ~midp(X36,X36,X38) | ~midp(X36,X37,X36)) ) | ~spl30_178),
  inference(subsumption_resolution,[],[f16279,f15313])).
fof(f15313,plain,(
  ( ! [X6,X8,X7] : (perp(X6,X8,X8,X7)) ) | ~spl30_178),
  inference(subsumption_resolution,[],[f15294,f3113])).
fof(f15294,plain,(
  ( ! [X6,X8,X7,X5] : (~coll(X5,X6,X7) | perp(X6,X8,X8,X7)) ) | ~spl30_178),
  inference(resolution,[],[f10134,f419])).
fof(f419,plain,(
  ( ! [X2,X0,X3,X1] : (~circle(X3,X0,X1,X2) | ~coll(X3,X0,X2) | perp(X0,X1,X1,X2)) )),
  inference(cnf_transformation,[],[f245])).
fof(f245,plain,(
  ! [X0,X1,X2,X3] : (perp(X0,X1,X1,X2) | ~coll(X3,X0,X2) | ~circle(X3,X0,X1,X2))),
  inference(flattening,[],[f244])).
fof(f244,plain,(
  ! [X0,X1,X2,X3] : (perp(X0,X1,X1,X2) | (~coll(X3,X0,X2) | ~circle(X3,X0,X1,X2)))),
  inference(ennf_transformation,[],[f133])).
fof(f133,plain,(
  ! [X0,X1,X2,X3] : ((coll(X3,X0,X2) & circle(X3,X0,X1,X2)) => perp(X0,X1,X1,X2))),
  inference(rectify,[],[f54])).
fof(f54,axiom,(
  ! [X0,X1,X2,X7] : ((coll(X7,X0,X2) & circle(X7,X0,X1,X2)) => perp(X0,X1,X1,X2))),
  file('./6_GDD_FULL/61-80/72.p',ruleD53)).
fof(f16279,plain,(
  ( ! [X37,X38,X36] : (~perp(X36,X37,X37,X38) | perp(X37,X38,X36,X38) | ~midp(X36,X36,X38) | ~midp(X36,X37,X36)) )),
  inference(resolution,[],[f2318,f2031])).
fof(f2318,plain,(
  ( ! [X26,X24,X25] : (~cong(X24,X25,X24,X26) | ~perp(X24,X25,X25,X26) | perp(X25,X26,X24,X26)) )),
  inference(resolution,[],[f988,f440])).
fof(f440,plain,(
  ( ! [X6,X4,X2,X0,X7,X5,X3,X1] : (~eqangle(X0,X1,X2,X3,X4,X5,X6,X7) | ~perp(X4,X5,X6,X7) | perp(X0,X1,X2,X3)) )),
  inference(cnf_transformation,[],[f281])).
fof(f281,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6,X7] : (perp(X0,X1,X2,X3) | ~perp(X4,X5,X6,X7) | ~eqangle(X0,X1,X2,X3,X4,X5,X6,X7))),
  inference(flattening,[],[f280])).
fof(f280,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6,X7] : (perp(X0,X1,X2,X3) | (~perp(X4,X5,X6,X7) | ~eqangle(X0,X1,X2,X3,X4,X5,X6,X7)))),
  inference(ennf_transformation,[],[f147])).
fof(f147,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6,X7] : ((perp(X4,X5,X6,X7) & eqangle(X0,X1,X2,X3,X4,X5,X6,X7)) => perp(X0,X1,X2,X3))),
  inference(rectify,[],[f75])).
fof(f75,axiom,(
  ! [X0,X1,X2,X3,X8,X9,X10,X11] : ((perp(X8,X9,X10,X11) & eqangle(X0,X1,X2,X3,X8,X9,X10,X11)) => perp(X0,X1,X2,X3))),
  file('./6_GDD_FULL/61-80/72.p',ruleD74)).
fof(f988,plain,(
  ( ! [X12,X10,X11] : (eqangle(X10,X11,X12,X11,X12,X10,X10,X11) | ~cong(X12,X10,X12,X11)) )),
  inference(resolution,[],[f385,f412])).
fof(f412,plain,(
  ( ! [X2,X0,X1] : (eqangle(X2,X0,X0,X1,X0,X1,X2,X1) | ~cong(X2,X0,X2,X1)) )),
  inference(cnf_transformation,[],[f231])).
fof(f231,plain,(
  ! [X0,X1,X2] : (eqangle(X2,X0,X0,X1,X0,X1,X2,X1) | ~cong(X2,X0,X2,X1))),
  inference(ennf_transformation,[],[f126])).
fof(f126,plain,(
  ! [X0,X1,X2] : (cong(X2,X0,X2,X1) => eqangle(X2,X0,X0,X1,X0,X1,X2,X1))),
  inference(rectify,[],[f47])).
fof(f47,axiom,(
  ! [X0,X1,X7] : (cong(X7,X0,X7,X1) => eqangle(X7,X0,X0,X1,X0,X1,X7,X1))),
  file('./6_GDD_FULL/61-80/72.p',ruleD46)).
fof(f2008,plain,(
  ( ! [X24,X23,X21,X25,X22,X20] : (~perp(X24,X25,sK13(X21,X22,X23,X20),X21) | para(X24,X25,X20,X21) | ~circle(X20,X21,X22,X23)) )),
  inference(resolution,[],[f910,f374])).
fof(f910,plain,(
  ( ! [X10,X8,X11,X9] : (perp(sK13(X9,X10,X11,X8),X9,X8,X9) | ~circle(X8,X9,X10,X11)) )),
  inference(resolution,[],[f465,f372])).
fof(f465,plain,(
  ( ! [X2,X0,X3,X1] : (perp(sK13(X0,X1,X2,X3),X0,X0,X3) | ~circle(X3,X0,X1,X2)) )),
  inference(cnf_transformation,[],[f345])).
fof(f345,plain,(
  ! [X0,X1,X2,X3] : (perp(sK13(X0,X1,X2,X3),X0,X0,X3) | ~circle(X3,X0,X1,X2))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK13])],[f304,f344])).
fof(f344,plain,(
  ! [X3,X2,X1,X0] : (? [X4] : (perp(X4,X0,X0,X3) | ~circle(X3,X0,X1,X2)) => (perp(sK13(X0,X1,X2,X3),X0,X0,X3) | ~circle(X3,X0,X1,X2)))),
  introduced(choice_axiom,[])).
fof(f304,plain,(
  ! [X0,X1,X2,X3] : ? [X4] : (perp(X4,X0,X0,X3) | ~circle(X3,X0,X1,X2))),
  inference(ennf_transformation,[],[f159])).
fof(f159,plain,(
  ! [X0,X1,X2,X3] : ? [X4] : (circle(X3,X0,X1,X2) => perp(X4,X0,X0,X3))),
  inference(rectify,[],[f87])).
fof(f87,axiom,(
  ! [X0,X1,X2,X7] : ? [X8] : (circle(X7,X0,X1,X2) => perp(X8,X0,X0,X7))),
  file('./6_GDD_FULL/61-80/72.p',ruleX11)).
fof(f25080,plain,(
  ~para(sK27,sK26,sK24,sK26) | ~para(sK26,sK28,sK26,sK23) | ~spl30_1),
  inference(resolution,[],[f5123,f508])).
fof(f508,plain,(
  ~eqangle(sK27,sK26,sK26,sK28,sK24,sK26,sK26,sK23) | ~spl30_1),
  inference(avatar_component_clause,[],[f507])).
fof(f507,plain,(
  spl30_1 <=> ~eqangle(sK27,sK26,sK26,sK28,sK24,sK26,sK26,sK23)),
  introduced(avatar_definition,[new_symbols(naming,[spl30_1])])).
fof(f24334,plain,(
  ~spl30_178 | spl30_459),
  inference(avatar_contradiction_clause,[],[f24333])).
fof(f24333,plain,(
  $false | (~spl30_178 | ~spl30_459)),
  inference(subsumption_resolution,[],[f11090,f22568])).
fof(f11090,plain,(
  ~para(sK22,sK23,sK24,sK23) | ~spl30_459),
  inference(avatar_component_clause,[],[f11089])).
fof(f11089,plain,(
  spl30_459 <=> ~para(sK22,sK23,sK24,sK23)),
  introduced(avatar_definition,[new_symbols(naming,[spl30_459])])).
fof(f15686,plain,(
  spl30_456 | ~spl30_459 | ~spl30_178),
  inference(avatar_split_clause,[],[f15685,f2167,f11089,f11083])).
fof(f15685,plain,(
  ( ! [X9] : (~para(sK22,sK23,sK24,sK23) | para(sK22,X9,sK24,sK22)) ) | ~spl30_178),
  inference(subsumption_resolution,[],[f11668,f15313])).
fof(f11668,plain,(
  ( ! [X9] : (~perp(sK25,sK22,sK22,X9) | ~para(sK22,sK23,sK24,sK23) | para(sK22,X9,sK24,sK22)) )),
  inference(resolution,[],[f1898,f439])).
fof(f439,plain,(
  ( ! [X6,X4,X2,X0,X7,X5,X3,X1] : (~eqangle(X0,X1,X2,X3,X4,X5,X6,X7) | ~para(X4,X5,X6,X7) | para(X0,X1,X2,X3)) )),
  inference(cnf_transformation,[],[f279])).
fof(f279,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6,X7] : (para(X0,X1,X2,X3) | ~para(X4,X5,X6,X7) | ~eqangle(X0,X1,X2,X3,X4,X5,X6,X7))),
  inference(flattening,[],[f278])).
fof(f278,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6,X7] : (para(X0,X1,X2,X3) | (~para(X4,X5,X6,X7) | ~eqangle(X0,X1,X2,X3,X4,X5,X6,X7)))),
  inference(ennf_transformation,[],[f146])).
fof(f146,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6,X7] : ((para(X4,X5,X6,X7) & eqangle(X0,X1,X2,X3,X4,X5,X6,X7)) => para(X0,X1,X2,X3))),
  inference(rectify,[],[f74])).
fof(f74,axiom,(
  ! [X0,X1,X2,X3,X8,X9,X10,X11] : ((para(X8,X9,X10,X11) & eqangle(X0,X1,X2,X3,X8,X9,X10,X11)) => para(X0,X1,X2,X3))),
  file('./6_GDD_FULL/61-80/72.p',ruleD73)).
fof(f1898,plain,(
  ( ! [X3] : (eqangle(sK22,X3,sK24,sK22,sK22,sK23,sK24,sK23) | ~perp(sK25,sK22,sK22,X3)) )),
  inference(resolution,[],[f1105,f386])).
fof(f1105,plain,(
  ( ! [X0] : (eqangle(sK22,X0,sK22,sK23,sK24,sK22,sK24,sK23) | ~perp(sK25,sK22,sK22,X0)) )),
  inference(resolution,[],[f414,f483])).
fof(f483,plain,(
  circle(sK25,sK22,sK23,sK24)),
  inference(cnf_transformation,[],[f365])).
fof(f414,plain,(
  ( ! [X4,X2,X0,X3,X1] : (~circle(X3,X0,X1,X2) | ~perp(X3,X0,X0,X4) | eqangle(X0,X4,X0,X1,X2,X0,X2,X1)) )),
  inference(cnf_transformation,[],[f235])).
fof(f235,plain,(
  ! [X0,X1,X2,X3,X4] : (eqangle(X0,X4,X0,X1,X2,X0,X2,X1) | ~perp(X3,X0,X0,X4) | ~circle(X3,X0,X1,X2))),
  inference(flattening,[],[f234])).
fof(f234,plain,(
  ! [X0,X1,X2,X3,X4] : (eqangle(X0,X4,X0,X1,X2,X0,X2,X1) | (~perp(X3,X0,X0,X4) | ~circle(X3,X0,X1,X2)))),
  inference(ennf_transformation,[],[f128])).
fof(f128,plain,(
  ! [X0,X1,X2,X3,X4] : ((perp(X3,X0,X0,X4) & circle(X3,X0,X1,X2)) => eqangle(X0,X4,X0,X1,X2,X0,X2,X1))),
  inference(rectify,[],[f49])).
fof(f49,axiom,(
  ! [X0,X1,X2,X7,X15] : ((perp(X7,X0,X0,X15) & circle(X7,X0,X1,X2)) => eqangle(X0,X15,X0,X1,X2,X0,X2,X1))),
  file('./6_GDD_FULL/61-80/72.p',ruleD48)).
fof(f14647,plain,(
  ~spl30_178 | ~spl30_556),
  inference(avatar_contradiction_clause,[],[f14646])).
fof(f14646,plain,(
  $false | (~spl30_178 | ~spl30_556)),
  inference(subsumption_resolution,[],[f14643,f9166])).
fof(f14643,plain,(
  ~midp(sK22,sK23,sK22) | ~spl30_556),
  inference(resolution,[],[f14276,f1573])).
fof(f1573,plain,(
  ( ! [X0,X1] : (para(X0,X0,X0,X1) | ~midp(X0,X1,X0)) )),
  inference(resolution,[],[f649,f602])).
fof(f649,plain,(
  ( ! [X0,X1] : (~cong(X0,X0,X0,X1) | para(X0,X0,X0,X1)) )),
  inference(resolution,[],[f412,f404])).
fof(f14276,plain,(
  ( ! [X28] : (~para(X28,sK22,X28,sK23)) ) | ~spl30_556),
  inference(avatar_component_clause,[],[f14275])).
fof(f14275,plain,(
  spl30_556 <=> ! [X28] : ~para(X28,sK22,X28,sK23)),
  introduced(avatar_definition,[new_symbols(naming,[spl30_556])])).
fof(f14277,plain,(
  spl30_556 | spl30_506 | spl30_155 | ~spl30_178),
  inference(avatar_split_clause,[],[f14273,f2167,f1911,f12094,f14275])).
fof(f12094,plain,(
  spl30_506 <=> ! [X27] : ~perp(sK25,sK22,sK22,X27)),
  introduced(avatar_definition,[new_symbols(naming,[spl30_506])])).
fof(f1911,plain,(
  spl30_155 <=> ~para(sK24,sK22,sK24,sK23)),
  introduced(avatar_definition,[new_symbols(naming,[spl30_155])])).
fof(f14273,plain,(
  ( ! [X28,X27] : (~perp(sK25,sK22,sK22,X27) | ~para(X28,sK22,X28,sK23)) ) | (~spl30_155 | ~spl30_178)),
  inference(subsumption_resolution,[],[f12149,f11628])).
fof(f11628,plain,(
  ( ! [X17] : (~perp(sK25,sK22,sK22,X17) | ~para(sK22,X17,sK22,sK23)) ) | ~spl30_155),
  inference(subsumption_resolution,[],[f11614,f1912])).
fof(f1912,plain,(
  ~para(sK24,sK22,sK24,sK23) | ~spl30_155),
  inference(avatar_component_clause,[],[f1911])).
fof(f11614,plain,(
  ( ! [X17] : (~perp(sK25,sK22,sK22,X17) | ~para(sK22,X17,sK22,sK23) | para(sK24,sK22,sK24,sK23)) )),
  inference(resolution,[],[f1897,f439])).
fof(f1897,plain,(
  ( ! [X2] : (eqangle(sK24,sK22,sK24,sK23,sK22,X2,sK22,sK23) | ~perp(sK25,sK22,sK22,X2)) )),
  inference(resolution,[],[f1105,f385])).
fof(f12149,plain,(
  ( ! [X28,X27] : (~perp(sK25,sK22,sK22,X27) | ~para(X28,sK22,X28,sK23) | para(sK22,X27,sK22,sK23)) ) | ~spl30_178),
  inference(resolution,[],[f11440,f439])).
fof(f11440,plain,(
  ( ! [X151,X150] : (eqangle(sK22,X150,sK22,sK23,X151,sK22,X151,sK23) | ~perp(sK25,sK22,sK22,X150)) ) | ~spl30_178),
  inference(subsumption_resolution,[],[f5187,f8309])).
fof(f8309,plain,(
  ( ! [X12,X10,X11,X9] : (cyclic(X9,X10,X11,X12)) ) | ~spl30_178),
  inference(subsumption_resolution,[],[f8287,f8280])).
fof(f8287,plain,(
  ( ! [X12,X10,X11,X9] : (cyclic(X9,X10,X11,X12) | ~cyclic(X9,X9,X10,X11)) ) | ~spl30_178),
  inference(resolution,[],[f8280,f382])).
fof(f382,plain,(
  ( ! [X4,X2,X0,X3,X1] : (~cyclic(X0,X1,X2,X4) | cyclic(X1,X2,X3,X4) | ~cyclic(X0,X1,X2,X3)) )),
  inference(cnf_transformation,[],[f191])).
fof(f191,plain,(
  ! [X0,X1,X2,X3,X4] : (cyclic(X1,X2,X3,X4) | ~cyclic(X0,X1,X2,X4) | ~cyclic(X0,X1,X2,X3))),
  inference(flattening,[],[f190])).
fof(f190,plain,(
  ! [X0,X1,X2,X3,X4] : (cyclic(X1,X2,X3,X4) | (~cyclic(X0,X1,X2,X4) | ~cyclic(X0,X1,X2,X3)))),
  inference(ennf_transformation,[],[f17])).
fof(f17,axiom,(
  ! [X0,X1,X2,X3,X4] : ((cyclic(X0,X1,X2,X4) & cyclic(X0,X1,X2,X3)) => cyclic(X1,X2,X3,X4))),
  file('./6_GDD_FULL/61-80/72.p',ruleD17)).
fof(f5187,plain,(
  ( ! [X151,X150] : (eqangle(sK22,X150,sK22,sK23,X151,sK22,X151,sK23) | ~cyclic(sK22,sK23,sK24,X151) | ~perp(sK25,sK22,sK22,X150)) )),
  inference(resolution,[],[f1364,f1105])).
fof(f1364,plain,(
  ( ! [X24,X23,X21,X19,X25,X22,X20,X18] : (~eqangle(X18,X19,X20,X21,X25,X23,X25,X24) | eqangle(X18,X19,X20,X21,X22,X23,X22,X24) | ~cyclic(X23,X24,X25,X22)) )),
  inference(resolution,[],[f387,f406])).
fof(f14250,plain,(
  ~spl30_178 | ~spl30_506),
  inference(avatar_contradiction_clause,[],[f14244])).
fof(f14244,plain,(
  $false | (~spl30_178 | ~spl30_506)),
  inference(resolution,[],[f9991,f12095])).
fof(f12095,plain,(
  ( ! [X27] : (~perp(sK25,sK22,sK22,X27)) ) | ~spl30_506),
  inference(avatar_component_clause,[],[f12094])).
fof(f9991,plain,(
  ( ! [X56,X55] : (perp(X55,X56,X56,X56)) ) | ~spl30_178),
  inference(subsumption_resolution,[],[f9990,f8312])).
fof(f9990,plain,(
  ( ! [X56,X55] : (perp(X55,X56,X56,X56) | ~cong(X56,X56,X56,X56)) ) | ~spl30_178),
  inference(subsumption_resolution,[],[f9986,f8950])).
fof(f9986,plain,(
  ( ! [X56,X55] : (~midp(X55,X56,X56) | perp(X55,X56,X56,X56) | ~cong(X56,X56,X56,X56)) )),
  inference(duplicate_literal_removal,[],[f9980])).
fof(f9980,plain,(
  ( ! [X56,X55] : (~midp(X55,X56,X56) | ~midp(X55,X56,X56) | perp(X55,X56,X56,X56) | ~cong(X56,X56,X56,X56)) )),
  inference(resolution,[],[f2018,f1113])).
fof(f1113,plain,(
  ( ! [X12,X13] : (~circle(X12,X13,X13,X13) | perp(X12,X13,X13,X13) | ~cong(X13,X13,X13,X13)) )),
  inference(resolution,[],[f415,f412])).
fof(f415,plain,(
  ( ! [X4,X2,X0,X3,X1] : (~eqangle(X0,X4,X0,X1,X2,X0,X2,X1) | perp(X3,X0,X0,X4) | ~circle(X3,X0,X1,X2)) )),
  inference(cnf_transformation,[],[f237])).
fof(f237,plain,(
  ! [X0,X1,X2,X3,X4] : (perp(X3,X0,X0,X4) | ~eqangle(X0,X4,X0,X1,X2,X0,X2,X1) | ~circle(X3,X0,X1,X2))),
  inference(flattening,[],[f236])).
fof(f236,plain,(
  ! [X0,X1,X2,X3,X4] : (perp(X3,X0,X0,X4) | (~eqangle(X0,X4,X0,X1,X2,X0,X2,X1) | ~circle(X3,X0,X1,X2)))),
  inference(ennf_transformation,[],[f129])).
fof(f129,plain,(
  ! [X0,X1,X2,X3,X4] : ((eqangle(X0,X4,X0,X1,X2,X0,X2,X1) & circle(X3,X0,X1,X2)) => perp(X3,X0,X0,X4))),
  inference(rectify,[],[f50])).
fof(f50,axiom,(
  ! [X0,X1,X2,X7,X15] : ((eqangle(X0,X15,X0,X1,X2,X0,X2,X1) & circle(X7,X0,X1,X2)) => perp(X7,X0,X0,X15))),
  file('./6_GDD_FULL/61-80/72.p',ruleD49)).
fof(f2018,plain,(
  ( ! [X2,X0,X3,X1] : (circle(X0,X1,X2,X3) | ~midp(X0,X1,X3) | ~midp(X0,X2,X1)) )),
  inference(resolution,[],[f946,f602])).
fof(f946,plain,(
  ( ! [X2,X0,X3,X1] : (~cong(X0,X1,X0,X2) | circle(X0,X1,X2,X3) | ~midp(X0,X1,X3)) )),
  inference(resolution,[],[f377,f434])).
fof(f9240,plain,(
  ~spl30_178 | ~spl30_222),
  inference(avatar_contradiction_clause,[],[f9239])).
fof(f9239,plain,(
  $false | (~spl30_178 | ~spl30_222)),
  inference(subsumption_resolution,[],[f7600,f9125])).
fof(f7600,plain,(
  ~midp(sK23,sK23,sK24) | (~spl30_178 | ~spl30_222)),
  inference(resolution,[],[f3491,f2168])).
fof(f3491,plain,(
  ( ! [X0] : (~para(sK25,sK23,sK25,X0) | ~midp(X0,sK23,sK24)) ) | ~spl30_222),
  inference(avatar_component_clause,[],[f3490])).
fof(f3490,plain,(
  spl30_222 <=> ! [X0] : (~midp(X0,sK23,sK24) | ~para(sK25,sK23,sK25,X0))),
  introduced(avatar_definition,[new_symbols(naming,[spl30_222])])).
fof(f3498,plain,(
  spl30_222 | spl30_224),
  inference(avatar_split_clause,[],[f3485,f3496,f3490])).
fof(f3485,plain,(
  ( ! [X0] : (para(sK22,sK23,sK22,sK24) | ~midp(X0,sK23,sK24) | ~para(sK25,sK23,sK25,X0)) )),
  inference(resolution,[],[f1150,f483])).
fof(f1150,plain,(
  ( ! [X14,X17,X15,X18,X16] : (~circle(X14,X17,X15,X18) | para(X17,X15,X17,X18) | ~midp(X16,X15,X18) | ~para(X14,X15,X14,X16)) )),
  inference(resolution,[],[f439,f416])).
fof(f416,plain,(
  ( ! [X4,X2,X0,X3,X1] : (eqangle(X0,X1,X0,X2,X3,X1,X3,X4) | ~midp(X4,X1,X2) | ~circle(X3,X0,X1,X2)) )),
  inference(cnf_transformation,[],[f239])).
fof(f239,plain,(
  ! [X0,X1,X2,X3,X4] : (eqangle(X0,X1,X0,X2,X3,X1,X3,X4) | ~midp(X4,X1,X2) | ~circle(X3,X0,X1,X2))),
  inference(flattening,[],[f238])).
fof(f238,plain,(
  ! [X0,X1,X2,X3,X4] : (eqangle(X0,X1,X0,X2,X3,X1,X3,X4) | (~midp(X4,X1,X2) | ~circle(X3,X0,X1,X2)))),
  inference(ennf_transformation,[],[f130])).
fof(f130,plain,(
  ! [X0,X1,X2,X3,X4] : ((midp(X4,X1,X2) & circle(X3,X0,X1,X2)) => eqangle(X0,X1,X0,X2,X3,X1,X3,X4))),
  inference(rectify,[],[f51])).
fof(f51,axiom,(
  ! [X0,X1,X2,X7,X6] : ((midp(X6,X1,X2) & circle(X7,X0,X1,X2)) => eqangle(X0,X1,X0,X2,X7,X1,X7,X6))),
  file('./6_GDD_FULL/61-80/72.p',ruleD50)).
fof(f2527,plain,(
  ~spl30_180),
  inference(avatar_contradiction_clause,[],[f2524])).
fof(f2524,plain,(
  $false | ~spl30_180),
  inference(resolution,[],[f2171,f1544])).
fof(f1544,plain,(
  para(sK22,sK24,sK22,sK24)),
  inference(resolution,[],[f921,f589])).
fof(f589,plain,(
  perp(sK22,sK24,sK27,sK26)),
  inference(resolution,[],[f373,f485])).
fof(f2171,plain,(
  ( ! [X0,X1] : (~para(X0,X1,X0,X1)) ) | ~spl30_180),
  inference(avatar_component_clause,[],[f2170])).
fof(f2170,plain,(
  spl30_180 <=> ! [X1,X0] : ~para(X0,X1,X0,X1)),
  introduced(avatar_definition,[new_symbols(naming,[spl30_180])])).
fof(f2172,plain,(
  spl30_178 | spl30_180),
  inference(avatar_split_clause,[],[f2139,f2170,f2167])).
fof(f2139,plain,(
  ( ! [X2,X0,X3,X1] : (~para(X0,X1,X0,X1) | para(X2,X3,X2,X3)) )),
  inference(resolution,[],[f979,f404])).
fof(f564,plain,(
  ~spl30_1 | ~spl30_13 | ~spl30_11 | ~spl30_7),
  inference(avatar_split_clause,[],[f563,f525,f539,f548,f507])).
fof(f563,plain,(
  ~eqangle(sK26,sK27,sK27,sK28,sK26,sK24,sK24,sK23) | ~eqangle(sK26,sK27,sK27,sK28,sK24,sK26,sK26,sK23) | ~eqangle(sK26,sK27,sK27,sK28,sK26,sK23,sK23,sK24) | ~eqangle(sK27,sK26,sK26,sK28,sK24,sK26,sK26,sK23)),
  inference(subsumption_resolution,[],[f562,f479])).
fof(f479,plain,(
  ( ! [X4,X2,X0,X3,X1] : (~eqangle(X2,X4,X4,X3,X2,X0,X0,X1) | ~sP1(X0,X1,X2,X3,X4)) )),
  inference(cnf_transformation,[],[f361])).
fof(f361,plain,(
  ! [X0,X1,X2,X3,X4] : ((~eqangle(X4,X2,X2,X3,X2,X1,X1,X0) & ~eqangle(X2,X4,X4,X3,X2,X0,X0,X1)) | ~sP1(X0,X1,X2,X3,X4))),
  inference(rectify,[],[f360])).
fof(f360,plain,(
  ! [X1,X2,X4,X6,X5] : ((~eqangle(X5,X4,X4,X6,X4,X2,X2,X1) & ~eqangle(X4,X5,X5,X6,X4,X1,X1,X2)) | ~sP1(X1,X2,X4,X6,X5))),
  inference(nnf_transformation,[],[f322])).
fof(f562,plain,(
  ~eqangle(sK26,sK27,sK27,sK28,sK26,sK24,sK24,sK23) | ~eqangle(sK26,sK27,sK27,sK28,sK24,sK26,sK26,sK23) | ~eqangle(sK26,sK27,sK27,sK28,sK26,sK23,sK23,sK24) | sP1(sK23,sK24,sK26,sK28,sK27) | ~eqangle(sK27,sK26,sK26,sK28,sK24,sK26,sK26,sK23)),
  inference(subsumption_resolution,[],[f497,f482])).
fof(f482,plain,(
  ( ! [X4,X2,X0,X3,X1] : (~eqangle(X4,X1,X1,X3,X2,X1,X1,X0) | ~sP0(X0,X1,X2,X3,X4)) )),
  inference(cnf_transformation,[],[f363])).
fof(f363,plain,(
  ! [X0,X1,X2,X3,X4] : ((~eqangle(X4,X1,X1,X3,X2,X1,X1,X0) & ~eqangle(X1,X4,X4,X3,X1,X2,X2,X0)) | ~sP0(X0,X1,X2,X3,X4))),
  inference(rectify,[],[f362])).
fof(f362,plain,(
  ! [X1,X4,X2,X6,X5] : ((~eqangle(X5,X4,X4,X6,X2,X4,X4,X1) & ~eqangle(X4,X5,X5,X6,X4,X2,X2,X1)) | ~sP0(X1,X4,X2,X6,X5))),
  inference(nnf_transformation,[],[f321])).
fof(f497,plain,(
  ~eqangle(sK26,sK27,sK27,sK28,sK26,sK24,sK24,sK23) | ~eqangle(sK26,sK27,sK27,sK28,sK24,sK26,sK26,sK23) | ~eqangle(sK26,sK27,sK27,sK28,sK26,sK23,sK23,sK24) | sP1(sK23,sK24,sK26,sK28,sK27) | sP0(sK23,sK26,sK24,sK28,sK27) | ~eqangle(sK27,sK26,sK26,sK28,sK24,sK26,sK26,sK23)),
  inference(duplicate_literal_removal,[],[f489])).
fof(f489,plain,(
  ~eqangle(sK26,sK27,sK27,sK28,sK26,sK24,sK24,sK23) | ~eqangle(sK26,sK27,sK27,sK28,sK24,sK26,sK26,sK23) | ~eqangle(sK26,sK27,sK27,sK28,sK26,sK23,sK23,sK24) | ~eqangle(sK26,sK27,sK27,sK28,sK24,sK26,sK26,sK23) | sP1(sK23,sK24,sK26,sK28,sK27) | sP0(sK23,sK26,sK24,sK28,sK27) | ~eqangle(sK27,sK26,sK26,sK28,sK24,sK26,sK26,sK23)),
  inference(cnf_transformation,[],[f365])).
% SZS output end Proof for 72
% ------------------------------
% Version: Vampire 4.4.0 (commit 7916d27 on 2019-08-23 08:50:16 +0100)
% Termination reason: Refutation

% Memory used [KB]: 17270
% Time elapsed: 0.429 s
% ------------------------------
% ------------------------------
