% Refutation found. Thanks to Tanya!
% SZS status Theorem for 75y
% SZS output start Proof for 75y
18. ! [X0,X1,X2,X3,X8,X9,X10,X11] : (eqangle(X0,X1,X2,X3,X8,X9,X10,X11) => eqangle(X1,X0,X2,X3,X8,X9,X10,X11)) [input]
19. ! [X0,X1,X2,X3,X8,X9,X10,X11] : (eqangle(X0,X1,X2,X3,X8,X9,X10,X11) => eqangle(X2,X3,X0,X1,X10,X11,X8,X9)) [input]
21. ! [X0,X1,X2,X3,X8,X9,X10,X11] : (eqangle(X0,X1,X2,X3,X8,X9,X10,X11) => eqangle(X0,X1,X8,X9,X2,X3,X10,X11)) [input]
22. ! [X0,X1,X2,X3,X8,X9,X10,X11,X4,X5,X12,X13] : ((eqangle(X8,X9,X10,X11,X4,X5,X12,X13) & eqangle(X0,X1,X2,X3,X8,X9,X10,X11)) => eqangle(X0,X1,X2,X3,X4,X5,X12,X13)) [input]
40. ! [X0,X1,X2,X3,X8,X9] : (para(X0,X1,X2,X3) => eqangle(X0,X1,X8,X9,X2,X3,X8,X9)) [input]
44. ! [X0,X1,X2,X4,X5] : ((midp(X5,X0,X2) & midp(X4,X0,X1)) => para(X4,X5,X1,X2)) [input]
94. ! [X0,X1,X2,X3,X4,X5,X12,X13,X18,X19,X20,X21,X22,X23,X24,X25] : (coll(X13,X2,X12) & coll(X13,X0,X5) & coll(X12,X1,X5) & circle(X2,X1,X12,X25) & circle(X0,X1,X5,X24) & midp(X4,X3,X1) & circle(X2,X1,X3,X23) & circle(X0,X1,X3,X22) & circle(X2,X1,X20,X21) & circle(X0,X1,X18,X19)) => eqangle(X5,X3,X3,X12,X0,X13,X13,X2) [input]
95. ~! [X2,X13,X0,X12,X3,X5] : (! [X0,X1,X2,X3,X4,X5,X12,X13,X18,X19,X20,X21,X22,X23,X24,X25] : (coll(X13,X2,X12) & coll(X13,X0,X5) & coll(X12,X1,X5) & circle(X2,X1,X12,X25) & circle(X0,X1,X5,X24) & midp(X4,X3,X1) & circle(X2,X1,X3,X23) & circle(X0,X1,X3,X22) & circle(X2,X1,X20,X21) & circle(X0,X1,X18,X19)) => eqangle(X5,X3,X3,X12,X0,X13,X13,X2)) [negated conjecture 94]
99. ! [X0,X1,X2,X3,X4,X5,X6,X7] : (eqangle(X0,X1,X2,X3,X4,X5,X6,X7) => eqangle(X1,X0,X2,X3,X4,X5,X6,X7)) [rectify 18]
100. ! [X0,X1,X2,X3,X4,X5,X6,X7] : (eqangle(X0,X1,X2,X3,X4,X5,X6,X7) => eqangle(X2,X3,X0,X1,X6,X7,X4,X5)) [rectify 19]
102. ! [X0,X1,X2,X3,X4,X5,X6,X7] : (eqangle(X0,X1,X2,X3,X4,X5,X6,X7) => eqangle(X0,X1,X4,X5,X2,X3,X6,X7)) [rectify 21]
103. ! [X0,X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11] : ((eqangle(X4,X5,X6,X7,X8,X9,X10,X11) & eqangle(X0,X1,X2,X3,X4,X5,X6,X7)) => eqangle(X0,X1,X2,X3,X8,X9,X10,X11)) [rectify 22]
118. ! [X0,X1,X2,X3,X4,X5] : (para(X0,X1,X2,X3) => eqangle(X0,X1,X4,X5,X2,X3,X4,X5)) [rectify 40]
122. ! [X0,X1,X2,X3,X4] : ((midp(X4,X0,X2) & midp(X3,X0,X1)) => para(X3,X4,X1,X2)) [rectify 44]
165. ~! [X0,X1,X2,X3,X4,X5] : (! [X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21] : (coll(X13,X8,X12) & coll(X13,X6,X11) & coll(X12,X7,X11) & circle(X8,X7,X12,X21) & circle(X6,X7,X11,X20) & midp(X10,X9,X7) & circle(X8,X7,X9,X19) & circle(X6,X7,X9,X18) & circle(X8,X7,X16,X17) & circle(X6,X7,X14,X15)) => eqangle(X5,X4,X4,X3,X2,X1,X1,X0)) [rectify 95]
190. ! [X0,X1,X2,X3,X4,X5,X6,X7] : (eqangle(X1,X0,X2,X3,X4,X5,X6,X7) | ~eqangle(X0,X1,X2,X3,X4,X5,X6,X7)) [ennf transformation 99]
191. ! [X0,X1,X2,X3,X4,X5,X6,X7] : (eqangle(X2,X3,X0,X1,X6,X7,X4,X5) | ~eqangle(X0,X1,X2,X3,X4,X5,X6,X7)) [ennf transformation 100]
193. ! [X0,X1,X2,X3,X4,X5,X6,X7] : (eqangle(X0,X1,X4,X5,X2,X3,X6,X7) | ~eqangle(X0,X1,X2,X3,X4,X5,X6,X7)) [ennf transformation 102]
194. ! [X0,X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11] : (eqangle(X0,X1,X2,X3,X8,X9,X10,X11) | (~eqangle(X4,X5,X6,X7,X8,X9,X10,X11) | ~eqangle(X0,X1,X2,X3,X4,X5,X6,X7))) [ennf transformation 103]
195. ! [X0,X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11] : (eqangle(X0,X1,X2,X3,X8,X9,X10,X11) | ~eqangle(X4,X5,X6,X7,X8,X9,X10,X11) | ~eqangle(X0,X1,X2,X3,X4,X5,X6,X7)) [flattening 194]
217. ! [X0,X1,X2,X3,X4,X5] : (eqangle(X0,X1,X4,X5,X2,X3,X4,X5) | ~para(X0,X1,X2,X3)) [ennf transformation 118]
223. ! [X0,X1,X2,X3,X4] : (para(X3,X4,X1,X2) | (~midp(X4,X0,X2) | ~midp(X3,X0,X1))) [ennf transformation 122]
224. ! [X0,X1,X2,X3,X4] : (para(X3,X4,X1,X2) | ~midp(X4,X0,X2) | ~midp(X3,X0,X1)) [flattening 223]
315. ? [X0,X1,X2,X3,X4,X5] : (~eqangle(X5,X4,X4,X3,X2,X1,X1,X0) & ! [X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21] : (coll(X13,X8,X12) & coll(X13,X6,X11) & coll(X12,X7,X11) & circle(X8,X7,X12,X21) & circle(X6,X7,X11,X20) & midp(X10,X9,X7) & circle(X8,X7,X9,X19) & circle(X6,X7,X9,X18) & circle(X8,X7,X16,X17) & circle(X6,X7,X14,X15))) [ennf transformation 165]
352. ? [X0,X1,X2,X3,X4,X5] : (~eqangle(X5,X4,X4,X3,X2,X1,X1,X0) & ! [X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21] : (coll(X13,X8,X12) & coll(X13,X6,X11) & coll(X12,X7,X11) & circle(X8,X7,X12,X21) & circle(X6,X7,X11,X20) & midp(X10,X9,X7) & circle(X8,X7,X9,X19) & circle(X6,X7,X9,X18) & circle(X8,X7,X16,X17) & circle(X6,X7,X14,X15))) => (~eqangle(sK25,sK24,sK24,sK23,sK22,sK21,sK21,sK20) & ! [X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21] : (coll(X13,X8,X12) & coll(X13,X6,X11) & coll(X12,X7,X11) & circle(X8,X7,X12,X21) & circle(X6,X7,X11,X20) & midp(X10,X9,X7) & circle(X8,X7,X9,X19) & circle(X6,X7,X9,X18) & circle(X8,X7,X16,X17) & circle(X6,X7,X14,X15))) [choice axiom]
353. ~eqangle(sK25,sK24,sK24,sK23,sK22,sK21,sK21,sK20) & ! [X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21] : (coll(X13,X8,X12) & coll(X13,X6,X11) & coll(X12,X7,X11) & circle(X8,X7,X12,X21) & circle(X6,X7,X11,X20) & midp(X10,X9,X7) & circle(X8,X7,X9,X19) & circle(X6,X7,X9,X18) & circle(X8,X7,X16,X17) & circle(X6,X7,X14,X15)) [skolemisation 315,352]
371. ~eqangle(X0,X1,X2,X3,X4,X5,X6,X7) | eqangle(X1,X0,X2,X3,X4,X5,X6,X7) [cnf transformation 190]
372. ~eqangle(X0,X1,X2,X3,X4,X5,X6,X7) | eqangle(X2,X3,X0,X1,X6,X7,X4,X5) [cnf transformation 191]
374. ~eqangle(X0,X1,X2,X3,X4,X5,X6,X7) | eqangle(X0,X1,X4,X5,X2,X3,X6,X7) [cnf transformation 193]
375. ~eqangle(X4,X5,X6,X7,X8,X9,X10,X11) | eqangle(X0,X1,X2,X3,X8,X9,X10,X11) | ~eqangle(X0,X1,X2,X3,X4,X5,X6,X7) [cnf transformation 195]
393. eqangle(X0,X1,X4,X5,X2,X3,X4,X5) | ~para(X0,X1,X2,X3) [cnf transformation 217]
398. ~midp(X4,X0,X2) | para(X3,X4,X1,X2) | ~midp(X3,X0,X1) [cnf transformation 224]
471. midp(X10,X9,X7) [cnf transformation 353]
477. ~eqangle(sK25,sK24,sK24,sK23,sK22,sK21,sK21,sK20) [cnf transformation 353]
494. para(X0,X1,X2,X3) | ~midp(X0,X4,X2) [resolution 398,471]
495. para(X0,X1,X2,X3) [subsumption resolution 494,471]
599. eqangle(X0,X1,X2,X3,X4,X5,X2,X3) | ~para(X1,X0,X4,X5) [resolution 371,393]
603. eqangle(X0,X1,X2,X3,X4,X5,X2,X3) [subsumption resolution 599,495]
605. eqangle(X0,X1,X2,X3,X0,X1,X4,X5) | ~para(X2,X3,X4,X5) [resolution 372,393]
609. eqangle(X0,X1,X2,X3,X0,X1,X4,X5) [subsumption resolution 605,495]
624. eqangle(X0,X1,X2,X3,X4,X5,X4,X5) [resolution 374,603]
631. eqangle(X2,X3,X2,X3,X4,X5,X6,X7) [resolution 609,374]
1636. eqangle(X250,X251,X252,X253,X254,X255,X256,X257) | ~eqangle(X250,X251,X252,X253,X258,X259,X258,X259) [resolution 375,631]
1648. eqangle(X250,X251,X252,X253,X254,X255,X256,X257) [subsumption resolution 1636,624]
1814. $false [resolution 1648,477]
% SZS output end Proof for 75y
% ------------------------------
% Version: Vampire 4.4.0 (commit 7916d27 on 2019-08-23 08:50:16 +0100)
% Termination reason: Refutation

% Memory used [KB]: 5373
% Time elapsed: 0.039 s
% ------------------------------
% ------------------------------
