% Refutation found. Thanks to Tanya!
% SZS status Theorem for 80
% SZS output start Proof for 80
fof(f1204,plain,(
  $false),
  inference(resolution,[],[f1177,f493])).
fof(f493,plain,(
  ~midp(sK21,sK20,sK22)),
  inference(resolution,[],[f489,f476])).
fof(f476,plain,(
  ~cong(sK21,sK22,sK21,sK20)),
  inference(cnf_transformation,[],[f357])).
fof(f357,plain,(
  ~cong(sK21,sK22,sK21,sK20) & ! [X3,X4,X5,X6,X7,X8] : (coll(X8,X4,X5) & perp(X3,X7,X3,X8) & circle(X7,X3,X4,X5) & coll(X6,X4,X5) & eqangle(X6,X3,X3,X5,X6,X3,X3,X4))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK20,sK21,sK22])],[f319,f356])).
fof(f356,plain,(
  ? [X0,X1,X2] : (~cong(X1,X2,X1,X0) & ! [X3,X4,X5,X6,X7,X8] : (coll(X8,X4,X5) & perp(X3,X7,X3,X8) & circle(X7,X3,X4,X5) & coll(X6,X4,X5) & eqangle(X6,X3,X3,X5,X6,X3,X3,X4))) => (~cong(sK21,sK22,sK21,sK20) & ! [X3,X4,X5,X6,X7,X8] : (coll(X8,X4,X5) & perp(X3,X7,X3,X8) & circle(X7,X3,X4,X5) & coll(X6,X4,X5) & eqangle(X6,X3,X3,X5,X6,X3,X3,X4)))),
  introduced(choice_axiom,[])).
fof(f319,plain,(
  ? [X0,X1,X2] : (~cong(X1,X2,X1,X0) & ! [X3,X4,X5,X6,X7,X8] : (coll(X8,X4,X5) & perp(X3,X7,X3,X8) & circle(X7,X3,X4,X5) & coll(X6,X4,X5) & eqangle(X6,X3,X3,X5,X6,X3,X3,X4)))),
  inference(ennf_transformation,[],[f167])).
fof(f167,plain,(
  ~! [X0,X1,X2] : (! [X3,X4,X5,X6,X7,X8] : (coll(X8,X4,X5) & perp(X3,X7,X3,X8) & circle(X7,X3,X4,X5) & coll(X6,X4,X5) & eqangle(X6,X3,X3,X5,X6,X3,X3,X4)) => cong(X1,X2,X1,X0))),
  inference(rectify,[],[f96])).
fof(f96,negated_conjecture,(
  ~! [X10,X18,X0] : (! [X0,X1,X2,X10,X7,X18] : (coll(X18,X1,X2) & perp(X0,X7,X0,X18) & circle(X7,X0,X1,X2) & coll(X10,X1,X2) & eqangle(X10,X0,X0,X2,X10,X0,X0,X1)) => cong(X18,X0,X18,X10))),
  inference(negated_conjecture,[],[f95])).
fof(f95,conjecture,(
  ( ! [X0,X10,X18] : (! [X0,X1,X2,X10,X7,X18] : (coll(X18,X1,X2) & perp(X0,X7,X0,X18) & circle(X7,X0,X1,X2) & coll(X10,X1,X2) & eqangle(X10,X0,X0,X2,X10,X0,X0,X1)) => cong(X18,X0,X18,X10)) )),
  file('./6_GDD_FULL/61-80/80.p',exemplo6GDDFULL618080)).
fof(f489,plain,(
  ( ! [X2,X0,X1] : (cong(X0,X1,X0,X2) | ~midp(X0,X2,X1)) )),
  inference(resolution,[],[f381,f426])).
fof(f426,plain,(
  ( ! [X2,X0,X1] : (cong(X0,X1,X0,X2) | ~midp(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f270])).
fof(f270,plain,(
  ! [X0,X1,X2] : (cong(X0,X1,X0,X2) | ~midp(X0,X1,X2))),
  inference(ennf_transformation,[],[f69])).
fof(f69,axiom,(
  ! [X0,X1,X2] : (midp(X0,X1,X2) => cong(X0,X1,X0,X2))),
  file('./6_GDD_FULL/61-80/80.p',ruleD68)).
fof(f381,plain,(
  ( ! [X2,X0,X3,X1] : (~cong(X0,X1,X2,X3) | cong(X2,X3,X0,X1)) )),
  inference(cnf_transformation,[],[f199])).
fof(f199,plain,(
  ! [X0,X1,X2,X3] : (cong(X2,X3,X0,X1) | ~cong(X0,X1,X2,X3))),
  inference(ennf_transformation,[],[f24])).
fof(f24,axiom,(
  ! [X0,X1,X2,X3] : (cong(X0,X1,X2,X3) => cong(X2,X3,X0,X1))),
  file('./6_GDD_FULL/61-80/80.p',ruleD24)).
fof(f1177,plain,(
  ( ! [X21,X19,X20] : (midp(X19,X20,X21)) )),
  inference(subsumption_resolution,[],[f1176,f473])).
fof(f473,plain,(
  ( ! [X4,X7,X5,X3] : (circle(X7,X3,X4,X5)) )),
  inference(cnf_transformation,[],[f357])).
fof(f1176,plain,(
  ( ! [X21,X19,X20] : (midp(X19,X20,X21) | ~circle(X20,X20,X20,X21)) )),
  inference(subsumption_resolution,[],[f1136,f475])).
fof(f475,plain,(
  ( ! [X4,X8,X5] : (coll(X8,X4,X5)) )),
  inference(cnf_transformation,[],[f357])).
fof(f1136,plain,(
  ( ! [X21,X19,X20] : (midp(X19,X20,X21) | ~coll(X19,X20,X21) | ~circle(X20,X20,X20,X21)) )),
  inference(resolution,[],[f409,f471])).
fof(f471,plain,(
  ( ! [X6,X4,X5,X3] : (eqangle(X6,X3,X3,X5,X6,X3,X3,X4)) )),
  inference(cnf_transformation,[],[f357])).
fof(f409,plain,(
  ( ! [X4,X2,X0,X3,X1] : (~eqangle(X0,X1,X0,X2,X3,X1,X3,X4) | midp(X4,X1,X2) | ~coll(X4,X1,X2) | ~circle(X3,X0,X1,X2)) )),
  inference(cnf_transformation,[],[f241])).
fof(f241,plain,(
  ! [X0,X1,X2,X3,X4] : (midp(X4,X1,X2) | ~eqangle(X0,X1,X0,X2,X3,X1,X3,X4) | ~coll(X4,X1,X2) | ~circle(X3,X0,X1,X2))),
  inference(flattening,[],[f240])).
fof(f240,plain,(
  ! [X0,X1,X2,X3,X4] : (midp(X4,X1,X2) | (~eqangle(X0,X1,X0,X2,X3,X1,X3,X4) | ~coll(X4,X1,X2) | ~circle(X3,X0,X1,X2)))),
  inference(ennf_transformation,[],[f131])).
fof(f131,plain,(
  ! [X0,X1,X2,X3,X4] : ((eqangle(X0,X1,X0,X2,X3,X1,X3,X4) & coll(X4,X1,X2) & circle(X3,X0,X1,X2)) => midp(X4,X1,X2))),
  inference(rectify,[],[f52])).
fof(f52,axiom,(
  ! [X0,X1,X2,X7,X6] : ((eqangle(X0,X1,X0,X2,X7,X1,X7,X6) & coll(X6,X1,X2) & circle(X7,X0,X1,X2)) => midp(X6,X1,X2))),
  file('./6_GDD_FULL/61-80/80.p',ruleD51)).
% SZS output end Proof for 80
% ------------------------------
% Version: Vampire 4.4.0 (commit 7916d27 on 2019-08-23 08:50:16 +0100)
% Termination reason: Refutation

% Memory used [KB]: 5373
% Time elapsed: 0.036 s
% ------------------------------
% ------------------------------
