% Refutation found. Thanks to Tanya!
% SZS status Theorem for 45
% SZS output start Proof for 45
10. ! [X0,X1,X2,X3,X4,X5] : ((perp(X2,X3,X4,X5) & para(X0,X1,X2,X3)) => perp(X0,X1,X4,X5)) [input]
44. ! [X0,X1,X2,X4,X5] : ((midp(X5,X0,X2) & midp(X4,X0,X1)) => para(X4,X5,X1,X2)) [input]
53. ! [X0,X1,X2,X7] : ((coll(X7,X0,X2) & circle(X7,X0,X1,X2)) => perp(X0,X1,X1,X2)) [input]
94. ! [X0,X1,X2,X7,X3,X9,X18,X19,X6,X20,X21] : (coll(X20,X1,X2) & coll(X20,X0,X3) & cong(X7,X19,X7,X6) & coll(X7,X19,X6) & midp(X19,X18,X9) & midp(X18,X0,X3) & midp(X9,X2,X1) & circle(X7,X0,X3,X21) & circle(X7,X0,X1,X2)) => perp(X18,X6,X1,X2) [input]
95. ~! [X2,X1,X6,X18] : (! [X0,X1,X2,X7,X3,X9,X18,X19,X6,X20,X21] : (coll(X20,X1,X2) & coll(X20,X0,X3) & cong(X7,X19,X7,X6) & coll(X7,X19,X6) & midp(X19,X18,X9) & midp(X18,X0,X3) & midp(X9,X2,X1) & circle(X7,X0,X3,X21) & circle(X7,X0,X1,X2)) => perp(X18,X6,X1,X2)) [negated conjecture 94]
122. ! [X0,X1,X2,X3,X4] : ((midp(X4,X0,X2) & midp(X3,X0,X1)) => para(X3,X4,X1,X2)) [rectify 44]
131. ! [X0,X1,X2,X3] : ((coll(X3,X0,X2) & circle(X3,X0,X1,X2)) => perp(X0,X1,X1,X2)) [rectify 53]
165. ~! [X0,X1,X2,X3] : (! [X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14] : (coll(X13,X5,X6) & coll(X13,X4,X8) & cong(X7,X11,X7,X12) & coll(X7,X11,X12) & midp(X11,X10,X9) & midp(X10,X4,X8) & midp(X9,X6,X5) & circle(X7,X4,X8,X14) & circle(X7,X4,X5,X6)) => perp(X3,X2,X1,X0)) [rectify 95]
178. ! [X0,X1,X2,X3,X4,X5] : (perp(X0,X1,X4,X5) | (~perp(X2,X3,X4,X5) | ~para(X0,X1,X2,X3))) [ennf transformation 10]
179. ! [X0,X1,X2,X3,X4,X5] : (perp(X0,X1,X4,X5) | ~perp(X2,X3,X4,X5) | ~para(X0,X1,X2,X3)) [flattening 178]
223. ! [X0,X1,X2,X3,X4] : (para(X3,X4,X1,X2) | (~midp(X4,X0,X2) | ~midp(X3,X0,X1))) [ennf transformation 122]
224. ! [X0,X1,X2,X3,X4] : (para(X3,X4,X1,X2) | ~midp(X4,X0,X2) | ~midp(X3,X0,X1)) [flattening 223]
240. ! [X0,X1,X2,X3] : (perp(X0,X1,X1,X2) | (~coll(X3,X0,X2) | ~circle(X3,X0,X1,X2))) [ennf transformation 131]
241. ! [X0,X1,X2,X3] : (perp(X0,X1,X1,X2) | ~coll(X3,X0,X2) | ~circle(X3,X0,X1,X2)) [flattening 240]
315. ? [X0,X1,X2,X3] : (~perp(X3,X2,X1,X0) & ! [X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14] : (coll(X13,X5,X6) & coll(X13,X4,X8) & cong(X7,X11,X7,X12) & coll(X7,X11,X12) & midp(X11,X10,X9) & midp(X10,X4,X8) & midp(X9,X6,X5) & circle(X7,X4,X8,X14) & circle(X7,X4,X5,X6))) [ennf transformation 165]
352. ? [X0,X1,X2,X3] : (~perp(X3,X2,X1,X0) & ! [X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14] : (coll(X13,X5,X6) & coll(X13,X4,X8) & cong(X7,X11,X7,X12) & coll(X7,X11,X12) & midp(X11,X10,X9) & midp(X10,X4,X8) & midp(X9,X6,X5) & circle(X7,X4,X8,X14) & circle(X7,X4,X5,X6))) => (~perp(sK23,sK22,sK21,sK20) & ! [X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14] : (coll(X13,X5,X6) & coll(X13,X4,X8) & cong(X7,X11,X7,X12) & coll(X7,X11,X12) & midp(X11,X10,X9) & midp(X10,X4,X8) & midp(X9,X6,X5) & circle(X7,X4,X8,X14) & circle(X7,X4,X5,X6))) [choice axiom]
353. ~perp(sK23,sK22,sK21,sK20) & ! [X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14] : (coll(X13,X5,X6) & coll(X13,X4,X8) & cong(X7,X11,X7,X12) & coll(X7,X11,X12) & midp(X11,X10,X9) & midp(X10,X4,X8) & midp(X9,X6,X5) & circle(X7,X4,X8,X14) & circle(X7,X4,X5,X6)) [skolemisation 315,352]
363. ~perp(X2,X3,X4,X5) | perp(X0,X1,X4,X5) | ~para(X0,X1,X2,X3) [cnf transformation 179]
398. ~midp(X4,X0,X2) | para(X3,X4,X1,X2) | ~midp(X3,X0,X1) [cnf transformation 224]
407. ~circle(X3,X0,X1,X2) | ~coll(X3,X0,X2) | perp(X0,X1,X1,X2) [cnf transformation 241]
468. circle(X7,X4,X8,X14) [cnf transformation 353]
471. midp(X11,X10,X9) [cnf transformation 353]
475. coll(X13,X5,X6) [cnf transformation 353]
476. ~perp(sK23,sK22,sK21,sK20) [cnf transformation 353]
493. para(X0,X1,X2,X3) | ~midp(X0,X4,X2) [resolution 398,471]
494. para(X0,X1,X2,X3) [subsumption resolution 493,471]
507. ~coll(X0,X1,X2) | perp(X1,X3,X3,X2) [resolution 407,468]
508. perp(X1,X3,X3,X2) [subsumption resolution 507,475]
535. perp(X0,X1,X2,X3) | ~para(X0,X1,X4,X2) [resolution 363,508]
540. perp(X0,X1,X2,X3) [subsumption resolution 535,494]
541. $false [resolution 540,476]
% SZS output end Proof for 45
% ------------------------------
% Version: Vampire 4.4.0 (commit 7916d27 on 2019-08-23 08:50:16 +0100)
% Termination reason: Refutation

% Memory used [KB]: 5117
% Time elapsed: 0.027 s
% ------------------------------
% ------------------------------
