% Refutation found. Thanks to Tanya!
% SZS status Theorem for 56
% SZS output start Proof for 56
fof(f31501,plain,(
  $false),
  inference(avatar_sat_refutation,[],[f491,f620,f1060,f1879,f2074,f2253,f2929,f2942,f3171,f3195,f3227,f3228,f3229,f3244,f3262,f3317,f3336,f3419,f4482,f4538,f5090,f5659,f5678,f5859,f5875,f5936,f5939,f6425,f6432,f6434,f6489,f6686,f7839,f7878,f7879,f7941,f7996,f8470,f8936,f9070,f9071,f9198,f9199,f11155,f11172,f11881,f11942,f12789,f12808,f12816,f12840,f12846,f15651,f16404,f16502,f16698,f16752,f18946,f18991,f19007,f19359,f19360,f30503,f30674,f31280,f31285,f31500])).
fof(f31500,plain,(
  ~spl26_182 | spl26_313),
  inference(avatar_contradiction_clause,[],[f31499])).
fof(f31499,plain,(
  $false | (~spl26_182 | ~spl26_313)),
  inference(subsumption_resolution,[],[f5360,f27630])).
fof(f27630,plain,(
  ( ! [X76,X75] : (cong(X75,X75,X75,X76)) ) | ~spl26_182),
  inference(resolution,[],[f13969,f381])).
fof(f381,plain,(
  ( ! [X2,X0,X3,X1] : (~cong(X0,X1,X2,X3) | cong(X0,X1,X3,X2)) )),
  inference(cnf_transformation,[],[f198])).
fof(f198,plain,(
  ! [X0,X1,X2,X3] : (cong(X0,X1,X3,X2) | ~cong(X0,X1,X2,X3))),
  inference(ennf_transformation,[],[f23])).
fof(f23,axiom,(
  ! [X0,X1,X2,X3] : (cong(X0,X1,X2,X3) => cong(X0,X1,X3,X2))),
  file('./6_GDD_FULL/41-60/56.p',ruleD23)).
fof(f13969,plain,(
  ( ! [X45,X44] : (cong(X44,X44,X45,X44)) ) | ~spl26_182),
  inference(resolution,[],[f12909,f382])).
fof(f382,plain,(
  ( ! [X2,X0,X3,X1] : (~cong(X0,X1,X2,X3) | cong(X2,X3,X0,X1)) )),
  inference(cnf_transformation,[],[f199])).
fof(f199,plain,(
  ! [X0,X1,X2,X3] : (cong(X2,X3,X0,X1) | ~cong(X0,X1,X2,X3))),
  inference(ennf_transformation,[],[f24])).
fof(f24,axiom,(
  ! [X0,X1,X2,X3] : (cong(X0,X1,X2,X3) => cong(X2,X3,X0,X1))),
  file('./6_GDD_FULL/41-60/56.p',ruleD24)).
fof(f12909,plain,(
  ( ! [X19,X18] : (cong(X18,X19,X19,X19)) ) | ~spl26_182),
  inference(subsumption_resolution,[],[f12852,f12081])).
fof(f12081,plain,(
  ( ! [X5] : (midp(X5,X5,X5)) ) | ~spl26_182),
  inference(subsumption_resolution,[],[f12049,f4068])).
fof(f4068,plain,(
  ( ! [X4,X3] : (coll(X3,X4,X4)) ) | ~spl26_182),
  inference(resolution,[],[f2070,f425])).
fof(f425,plain,(
  ( ! [X2,X0,X1] : (~para(X0,X1,X0,X2) | coll(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f267])).
fof(f267,plain,(
  ! [X0,X1,X2] : (coll(X0,X1,X2) | ~para(X0,X1,X0,X2))),
  inference(ennf_transformation,[],[f67])).
fof(f67,axiom,(
  ! [X0,X1,X2] : (para(X0,X1,X0,X2) => coll(X0,X1,X2))),
  file('./6_GDD_FULL/41-60/56.p',ruleD66)).
fof(f2070,plain,(
  ( ! [X2,X3] : (para(X2,X3,X2,X3)) ) | ~spl26_182),
  inference(avatar_component_clause,[],[f2069])).
fof(f2069,plain,(
  spl26_182 <=> ! [X3,X2] : para(X2,X3,X2,X3)),
  introduced(avatar_definition,[new_symbols(naming,[spl26_182])])).
fof(f12049,plain,(
  ( ! [X5] : (~coll(X5,X5,X5) | midp(X5,X5,X5)) ) | ~spl26_182),
  inference(resolution,[],[f12046,f426])).
fof(f426,plain,(
  ( ! [X2,X0,X1] : (~cong(X0,X1,X0,X2) | ~coll(X0,X1,X2) | midp(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f269])).
fof(f269,plain,(
  ! [X0,X1,X2] : (midp(X0,X1,X2) | ~coll(X0,X1,X2) | ~cong(X0,X1,X0,X2))),
  inference(flattening,[],[f268])).
fof(f268,plain,(
  ! [X0,X1,X2] : (midp(X0,X1,X2) | (~coll(X0,X1,X2) | ~cong(X0,X1,X0,X2)))),
  inference(ennf_transformation,[],[f68])).
fof(f68,axiom,(
  ! [X0,X1,X2] : ((coll(X0,X1,X2) & cong(X0,X1,X0,X2)) => midp(X0,X1,X2))),
  file('./6_GDD_FULL/41-60/56.p',ruleD67)).
fof(f12046,plain,(
  ( ! [X87,X86] : (cong(X87,X87,X86,X86)) ) | ~spl26_182),
  inference(subsumption_resolution,[],[f12045,f9544])).
fof(f9544,plain,(
  ( ! [X6,X4,X5] : (cyclic(X4,X4,X5,X6)) ) | ~spl26_182),
  inference(resolution,[],[f3173,f2070])).
fof(f3173,plain,(
  ( ! [X68,X66,X67] : (~para(X66,X67,X66,X67) | cyclic(X67,X67,X66,X68)) )),
  inference(subsumption_resolution,[],[f2236,f2237])).
fof(f2237,plain,(
  ( ! [X70,X71,X69] : (~para(X69,X70,X69,X70) | coll(X69,X71,X70) | cyclic(X70,X70,X69,X71)) )),
  inference(resolution,[],[f824,f400])).
fof(f400,plain,(
  ( ! [X2,X0,X3,X1] : (~eqangle(X2,X0,X2,X1,X3,X0,X3,X1) | coll(X2,X3,X0) | cyclic(X0,X1,X2,X3)) )),
  inference(cnf_transformation,[],[f222])).
fof(f222,plain,(
  ! [X0,X1,X2,X3] : (cyclic(X0,X1,X2,X3) | coll(X2,X3,X0) | ~eqangle(X2,X0,X2,X1,X3,X0,X3,X1))),
  inference(flattening,[],[f221])).
fof(f221,plain,(
  ! [X0,X1,X2,X3] : (cyclic(X0,X1,X2,X3) | (coll(X2,X3,X0) | ~eqangle(X2,X0,X2,X1,X3,X0,X3,X1)))),
  inference(ennf_transformation,[],[f121])).
fof(f121,plain,(
  ! [X0,X1,X2,X3] : ((~coll(X2,X3,X0) & eqangle(X2,X0,X2,X1,X3,X0,X3,X1)) => cyclic(X0,X1,X2,X3))),
  inference(rectify,[],[f42])).
fof(f42,axiom,(
  ! [X0,X1,X8,X9] : ((~coll(X8,X9,X0) & eqangle(X8,X0,X8,X1,X9,X0,X9,X1)) => cyclic(X0,X1,X8,X9))),
  file('./6_GDD_FULL/41-60/56.p',ruleD42a)).
fof(f824,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (eqangle(X0,X1,X2,X3,X4,X5,X4,X5) | ~para(X0,X1,X2,X3)) )),
  inference(resolution,[],[f379,f398])).
fof(f398,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (eqangle(X0,X1,X4,X5,X2,X3,X4,X5) | ~para(X0,X1,X2,X3)) )),
  inference(cnf_transformation,[],[f219])).
fof(f219,plain,(
  ! [X0,X1,X2,X3,X4,X5] : (eqangle(X0,X1,X4,X5,X2,X3,X4,X5) | ~para(X0,X1,X2,X3))),
  inference(ennf_transformation,[],[f119])).
fof(f119,plain,(
  ! [X0,X1,X2,X3,X4,X5] : (para(X0,X1,X2,X3) => eqangle(X0,X1,X4,X5,X2,X3,X4,X5))),
  inference(rectify,[],[f40])).
fof(f40,axiom,(
  ! [X0,X1,X2,X3,X8,X9] : (para(X0,X1,X2,X3) => eqangle(X0,X1,X8,X9,X2,X3,X8,X9))),
  file('./6_GDD_FULL/41-60/56.p',ruleD40)).
fof(f379,plain,(
  ( ! [X6,X4,X2,X0,X7,X5,X3,X1] : (~eqangle(X0,X1,X2,X3,X4,X5,X6,X7) | eqangle(X0,X1,X4,X5,X2,X3,X6,X7)) )),
  inference(cnf_transformation,[],[f195])).
fof(f195,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6,X7] : (eqangle(X0,X1,X4,X5,X2,X3,X6,X7) | ~eqangle(X0,X1,X2,X3,X4,X5,X6,X7))),
  inference(ennf_transformation,[],[f103])).
fof(f103,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6,X7] : (eqangle(X0,X1,X2,X3,X4,X5,X6,X7) => eqangle(X0,X1,X4,X5,X2,X3,X6,X7))),
  inference(rectify,[],[f21])).
fof(f21,axiom,(
  ! [X0,X1,X2,X3,X8,X9,X10,X11] : (eqangle(X0,X1,X2,X3,X8,X9,X10,X11) => eqangle(X0,X1,X8,X9,X2,X3,X10,X11))),
  file('./6_GDD_FULL/41-60/56.p',ruleD21)).
fof(f2236,plain,(
  ( ! [X68,X66,X67] : (~para(X66,X67,X66,X67) | ~coll(X66,X68,X67) | cyclic(X67,X67,X66,X68)) )),
  inference(resolution,[],[f824,f401])).
fof(f401,plain,(
  ( ! [X2,X0,X3,X1] : (~eqangle(X2,X0,X2,X1,X3,X0,X3,X1) | ~coll(X2,X3,X1) | cyclic(X0,X1,X2,X3)) )),
  inference(cnf_transformation,[],[f224])).
fof(f224,plain,(
  ! [X0,X1,X2,X3] : (cyclic(X0,X1,X2,X3) | ~coll(X2,X3,X1) | ~eqangle(X2,X0,X2,X1,X3,X0,X3,X1))),
  inference(flattening,[],[f223])).
fof(f223,plain,(
  ! [X0,X1,X2,X3] : (cyclic(X0,X1,X2,X3) | (~coll(X2,X3,X1) | ~eqangle(X2,X0,X2,X1,X3,X0,X3,X1)))),
  inference(ennf_transformation,[],[f122])).
fof(f122,plain,(
  ! [X0,X1,X2,X3] : ((coll(X2,X3,X1) & eqangle(X2,X0,X2,X1,X3,X0,X3,X1)) => cyclic(X0,X1,X2,X3))),
  inference(rectify,[],[f43])).
fof(f43,axiom,(
  ! [X0,X1,X8,X9] : ((coll(X8,X9,X1) & eqangle(X8,X0,X8,X1,X9,X0,X9,X1)) => cyclic(X0,X1,X8,X9))),
  file('./6_GDD_FULL/41-60/56.p',ruleD42b)).
fof(f12045,plain,(
  ( ! [X88,X87,X86] : (cong(X87,X87,X86,X86) | ~cyclic(X87,X87,X88,X86)) ) | ~spl26_182),
  inference(subsumption_resolution,[],[f12044,f9544])).
fof(f12044,plain,(
  ( ! [X88,X87,X86] : (cong(X87,X87,X86,X86) | ~cyclic(X87,X87,X88,X88) | ~cyclic(X87,X87,X88,X86)) ) | ~spl26_182),
  inference(subsumption_resolution,[],[f12034,f11715])).
fof(f11715,plain,(
  ( ! [X12,X10,X11,X9] : (cyclic(X9,X10,X11,X12)) ) | ~spl26_182),
  inference(subsumption_resolution,[],[f11693,f9544])).
fof(f11693,plain,(
  ( ! [X12,X10,X11,X9] : (cyclic(X9,X10,X11,X12) | ~cyclic(X9,X9,X10,X11)) ) | ~spl26_182),
  inference(resolution,[],[f9544,f375])).
fof(f375,plain,(
  ( ! [X4,X2,X0,X3,X1] : (~cyclic(X0,X1,X2,X4) | cyclic(X1,X2,X3,X4) | ~cyclic(X0,X1,X2,X3)) )),
  inference(cnf_transformation,[],[f191])).
fof(f191,plain,(
  ! [X0,X1,X2,X3,X4] : (cyclic(X1,X2,X3,X4) | ~cyclic(X0,X1,X2,X4) | ~cyclic(X0,X1,X2,X3))),
  inference(flattening,[],[f190])).
fof(f190,plain,(
  ! [X0,X1,X2,X3,X4] : (cyclic(X1,X2,X3,X4) | (~cyclic(X0,X1,X2,X4) | ~cyclic(X0,X1,X2,X3)))),
  inference(ennf_transformation,[],[f17])).
fof(f17,axiom,(
  ! [X0,X1,X2,X3,X4] : ((cyclic(X0,X1,X2,X4) & cyclic(X0,X1,X2,X3)) => cyclic(X1,X2,X3,X4))),
  file('./6_GDD_FULL/41-60/56.p',ruleD17)).
fof(f12034,plain,(
  ( ! [X88,X87,X86] : (~cyclic(X86,X87,X88,X88) | cong(X87,X87,X86,X86) | ~cyclic(X87,X87,X88,X88) | ~cyclic(X87,X87,X88,X86)) )),
  inference(duplicate_literal_removal,[],[f12023])).
fof(f12023,plain,(
  ( ! [X88,X87,X86] : (~cyclic(X86,X87,X88,X88) | cong(X87,X87,X86,X86) | ~cyclic(X87,X87,X88,X88) | ~cyclic(X87,X87,X88,X86) | ~cyclic(X87,X87,X88,X86)) )),
  inference(resolution,[],[f2085,f402])).
fof(f402,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (~eqangle(X2,X0,X2,X1,X5,X3,X5,X4) | cong(X0,X1,X3,X4) | ~cyclic(X0,X1,X2,X5) | ~cyclic(X0,X1,X2,X4) | ~cyclic(X0,X1,X2,X3)) )),
  inference(cnf_transformation,[],[f226])).
fof(f226,plain,(
  ! [X0,X1,X2,X3,X4,X5] : (cong(X0,X1,X3,X4) | ~eqangle(X2,X0,X2,X1,X5,X3,X5,X4) | ~cyclic(X0,X1,X2,X5) | ~cyclic(X0,X1,X2,X4) | ~cyclic(X0,X1,X2,X3))),
  inference(flattening,[],[f225])).
fof(f225,plain,(
  ! [X0,X1,X2,X3,X4,X5] : (cong(X0,X1,X3,X4) | (~eqangle(X2,X0,X2,X1,X5,X3,X5,X4) | ~cyclic(X0,X1,X2,X5) | ~cyclic(X0,X1,X2,X4) | ~cyclic(X0,X1,X2,X3)))),
  inference(ennf_transformation,[],[f123])).
fof(f123,plain,(
  ! [X0,X1,X2,X3,X4,X5] : ((eqangle(X2,X0,X2,X1,X5,X3,X5,X4) & cyclic(X0,X1,X2,X5) & cyclic(X0,X1,X2,X4) & cyclic(X0,X1,X2,X3)) => cong(X0,X1,X3,X4))),
  inference(rectify,[],[f44])).
fof(f44,axiom,(
  ! [X0,X1,X2,X8,X9,X14] : ((eqangle(X2,X0,X2,X1,X14,X8,X14,X9) & cyclic(X0,X1,X2,X14) & cyclic(X0,X1,X2,X9) & cyclic(X0,X1,X2,X8)) => cong(X0,X1,X8,X9))),
  file('./6_GDD_FULL/41-60/56.p',ruleD43)).
fof(f2085,plain,(
  ( ! [X17,X15,X18,X16] : (eqangle(X17,X16,X18,X16,X17,X15,X18,X15) | ~cyclic(X15,X16,X17,X18)) )),
  inference(resolution,[],[f812,f379])).
fof(f812,plain,(
  ( ! [X6,X8,X7,X9] : (eqangle(X6,X7,X6,X8,X9,X7,X9,X8) | ~cyclic(X8,X7,X6,X9)) )),
  inference(resolution,[],[f377,f399])).
fof(f399,plain,(
  ( ! [X2,X0,X3,X1] : (eqangle(X2,X0,X2,X1,X3,X0,X3,X1) | ~cyclic(X0,X1,X2,X3)) )),
  inference(cnf_transformation,[],[f220])).
fof(f220,plain,(
  ! [X0,X1,X2,X3] : (eqangle(X2,X0,X2,X1,X3,X0,X3,X1) | ~cyclic(X0,X1,X2,X3))),
  inference(ennf_transformation,[],[f120])).
fof(f120,plain,(
  ! [X0,X1,X2,X3] : (cyclic(X0,X1,X2,X3) => eqangle(X2,X0,X2,X1,X3,X0,X3,X1))),
  inference(rectify,[],[f41])).
fof(f41,axiom,(
  ! [X0,X1,X8,X9] : (cyclic(X0,X1,X8,X9) => eqangle(X8,X0,X8,X1,X9,X0,X9,X1))),
  file('./6_GDD_FULL/41-60/56.p',ruleD41)).
fof(f377,plain,(
  ( ! [X6,X4,X2,X0,X7,X5,X3,X1] : (~eqangle(X0,X1,X2,X3,X4,X5,X6,X7) | eqangle(X2,X3,X0,X1,X6,X7,X4,X5)) )),
  inference(cnf_transformation,[],[f193])).
fof(f193,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6,X7] : (eqangle(X2,X3,X0,X1,X6,X7,X4,X5) | ~eqangle(X0,X1,X2,X3,X4,X5,X6,X7))),
  inference(ennf_transformation,[],[f101])).
fof(f101,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6,X7] : (eqangle(X0,X1,X2,X3,X4,X5,X6,X7) => eqangle(X2,X3,X0,X1,X6,X7,X4,X5))),
  inference(rectify,[],[f19])).
fof(f19,axiom,(
  ! [X0,X1,X2,X3,X8,X9,X10,X11] : (eqangle(X0,X1,X2,X3,X8,X9,X10,X11) => eqangle(X2,X3,X0,X1,X10,X11,X8,X9))),
  file('./6_GDD_FULL/41-60/56.p',ruleD19)).
fof(f12852,plain,(
  ( ! [X19,X18] : (cong(X18,X19,X19,X19) | ~midp(X18,X18,X18)) ) | ~spl26_182),
  inference(resolution,[],[f4399,f2070])).
fof(f4399,plain,(
  ( ! [X21,X19,X20,X18] : (~para(X19,X21,X20,X18) | cong(X18,X19,X19,X20) | ~midp(X18,X21,X21)) ) | ~spl26_182),
  inference(subsumption_resolution,[],[f4398,f4087])).
fof(f4087,plain,(
  ( ! [X4,X3] : (coll(X3,X4,X3)) ) | ~spl26_182),
  inference(resolution,[],[f4068,f360])).
fof(f360,plain,(
  ( ! [X2,X0,X1] : (~coll(X0,X1,X2) | coll(X1,X0,X2)) )),
  inference(cnf_transformation,[],[f169])).
fof(f169,plain,(
  ! [X0,X1,X2] : (coll(X1,X0,X2) | ~coll(X0,X1,X2))),
  inference(ennf_transformation,[],[f2])).
fof(f2,axiom,(
  ! [X0,X1,X2] : (coll(X0,X1,X2) => coll(X1,X0,X2))),
  file('./6_GDD_FULL/41-60/56.p',ruleD2)).
fof(f4398,plain,(
  ( ! [X21,X19,X20,X18] : (~para(X19,X21,X20,X18) | ~coll(X18,X21,X18) | cong(X18,X19,X19,X20) | ~midp(X18,X21,X21)) ) | ~spl26_182),
  inference(subsumption_resolution,[],[f4393,f4151])).
fof(f4151,plain,(
  ( ! [X2,X0,X1] : (coll(X0,X1,X2)) ) | ~spl26_182),
  inference(subsumption_resolution,[],[f4142,f4116])).
fof(f4116,plain,(
  ( ! [X6,X5] : (coll(X5,X5,X6)) ) | ~spl26_182),
  inference(resolution,[],[f4087,f359])).
fof(f359,plain,(
  ( ! [X2,X0,X1] : (~coll(X0,X1,X2) | coll(X0,X2,X1)) )),
  inference(cnf_transformation,[],[f168])).
fof(f168,plain,(
  ! [X0,X1,X2] : (coll(X0,X2,X1) | ~coll(X0,X1,X2))),
  inference(ennf_transformation,[],[f1])).
fof(f1,axiom,(
  ! [X0,X1,X2] : (coll(X0,X1,X2) => coll(X0,X2,X1))),
  file('./6_GDD_FULL/41-60/56.p',ruleD1)).
fof(f4142,plain,(
  ( ! [X2,X0,X1] : (coll(X0,X1,X2) | ~coll(X2,X2,X0)) ) | ~spl26_182),
  inference(resolution,[],[f4116,f361])).
fof(f361,plain,(
  ( ! [X2,X0,X3,X1] : (~coll(X0,X1,X3) | coll(X2,X3,X0) | ~coll(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f171])).
fof(f171,plain,(
  ! [X0,X1,X2,X3] : (coll(X2,X3,X0) | ~coll(X0,X1,X3) | ~coll(X0,X1,X2))),
  inference(flattening,[],[f170])).
fof(f170,plain,(
  ! [X0,X1,X2,X3] : (coll(X2,X3,X0) | (~coll(X0,X1,X3) | ~coll(X0,X1,X2)))),
  inference(ennf_transformation,[],[f3])).
fof(f3,axiom,(
  ! [X0,X1,X2,X3] : ((coll(X0,X1,X3) & coll(X0,X1,X2)) => coll(X2,X3,X0))),
  file('./6_GDD_FULL/41-60/56.p',ruleD3)).
fof(f4393,plain,(
  ( ! [X21,X19,X20,X18] : (~coll(X18,X19,X20) | ~para(X19,X21,X20,X18) | ~coll(X18,X21,X18) | cong(X18,X19,X19,X20) | ~midp(X18,X21,X21)) )),
  inference(resolution,[],[f1160,f526])).
fof(f526,plain,(
  ( ! [X2,X0,X1] : (cong(X0,X1,X2,X0) | ~midp(X0,X1,X2)) )),
  inference(resolution,[],[f381,f427])).
fof(f427,plain,(
  ( ! [X2,X0,X1] : (cong(X0,X1,X0,X2) | ~midp(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f270])).
fof(f270,plain,(
  ! [X0,X1,X2] : (cong(X0,X1,X0,X2) | ~midp(X0,X1,X2))),
  inference(ennf_transformation,[],[f69])).
fof(f69,axiom,(
  ! [X0,X1,X2] : (midp(X0,X1,X2) => cong(X0,X1,X0,X2))),
  file('./6_GDD_FULL/41-60/56.p',ruleD68)).
fof(f1160,plain,(
  ( ! [X4,X2,X0,X3,X1] : (~cong(X0,X1,X1,X2) | ~coll(X0,X3,X4) | ~para(X3,X1,X4,X2) | ~coll(X0,X1,X2) | cong(X0,X3,X3,X4)) )),
  inference(resolution,[],[f424,f434])).
fof(f434,plain,(
  ( ! [X6,X4,X2,X0,X7,X5,X3,X1] : (~eqratio(X0,X1,X2,X3,X4,X5,X6,X7) | ~cong(X4,X5,X6,X7) | cong(X0,X1,X2,X3)) )),
  inference(cnf_transformation,[],[f283])).
fof(f283,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6,X7] : (cong(X0,X1,X2,X3) | ~cong(X4,X5,X6,X7) | ~eqratio(X0,X1,X2,X3,X4,X5,X6,X7))),
  inference(flattening,[],[f282])).
fof(f282,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6,X7] : (cong(X0,X1,X2,X3) | (~cong(X4,X5,X6,X7) | ~eqratio(X0,X1,X2,X3,X4,X5,X6,X7)))),
  inference(ennf_transformation,[],[f148])).
fof(f148,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6,X7] : ((cong(X4,X5,X6,X7) & eqratio(X0,X1,X2,X3,X4,X5,X6,X7)) => cong(X0,X1,X2,X3))),
  inference(rectify,[],[f76])).
fof(f76,axiom,(
  ! [X0,X1,X2,X3,X8,X9,X10,X11] : ((cong(X8,X9,X10,X11) & eqratio(X0,X1,X2,X3,X8,X9,X10,X11)) => cong(X0,X1,X2,X3))),
  file('./6_GDD_FULL/41-60/56.p',ruleD75)).
fof(f424,plain,(
  ( ! [X4,X2,X0,X3,X1] : (eqratio(X4,X0,X0,X2,X4,X1,X1,X3) | ~coll(X4,X1,X3) | ~coll(X4,X0,X2) | ~para(X0,X1,X2,X3)) )),
  inference(cnf_transformation,[],[f266])).
fof(f266,plain,(
  ! [X0,X1,X2,X3,X4] : (eqratio(X4,X0,X0,X2,X4,X1,X1,X3) | ~coll(X4,X1,X3) | ~coll(X4,X0,X2) | ~para(X0,X1,X2,X3))),
  inference(flattening,[],[f265])).
fof(f265,plain,(
  ! [X0,X1,X2,X3,X4] : (eqratio(X4,X0,X0,X2,X4,X1,X1,X3) | (~coll(X4,X1,X3) | ~coll(X4,X0,X2) | ~para(X0,X1,X2,X3)))),
  inference(ennf_transformation,[],[f144])).
fof(f144,plain,(
  ! [X0,X1,X2,X3,X4] : ((coll(X4,X1,X3) & coll(X4,X0,X2) & para(X0,X1,X2,X3)) => eqratio(X4,X0,X0,X2,X4,X1,X1,X3))),
  inference(rectify,[],[f66])).
fof(f66,axiom,(
  ! [X0,X1,X2,X3,X7] : ((coll(X7,X1,X3) & coll(X7,X0,X2) & para(X0,X1,X2,X3)) => eqratio(X7,X0,X0,X2,X7,X1,X1,X3))),
  file('./6_GDD_FULL/41-60/56.p',ruleD65)).
fof(f5360,plain,(
  ~cong(sK20,sK20,sK20,sK22) | ~spl26_313),
  inference(avatar_component_clause,[],[f5359])).
fof(f5359,plain,(
  spl26_313 <=> ~cong(sK20,sK20,sK20,sK22)),
  introduced(avatar_definition,[new_symbols(naming,[spl26_313])])).
fof(f31285,plain,(
  spl26_162 | ~spl26_294),
  inference(avatar_split_clause,[],[f12385,f4536,f1842])).
fof(f1842,plain,(
  spl26_162 <=> para(sK22,sK20,sK22,sK21)),
  introduced(avatar_definition,[new_symbols(naming,[spl26_162])])).
fof(f4536,plain,(
  spl26_294 <=> para(sK20,sK22,sK22,sK21)),
  introduced(avatar_definition,[new_symbols(naming,[spl26_294])])).
fof(f12385,plain,(
  para(sK22,sK20,sK22,sK21) | ~spl26_294),
  inference(resolution,[],[f4537,f1961])).
fof(f1961,plain,(
  ( ! [X2,X0,X3,X1] : (~para(X0,X1,X2,X3) | para(X1,X0,X2,X3)) )),
  inference(resolution,[],[f804,f397])).
fof(f397,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (~eqangle(X0,X1,X4,X5,X2,X3,X4,X5) | para(X0,X1,X2,X3)) )),
  inference(cnf_transformation,[],[f218])).
fof(f218,plain,(
  ! [X0,X1,X2,X3,X4,X5] : (para(X0,X1,X2,X3) | ~eqangle(X0,X1,X4,X5,X2,X3,X4,X5))),
  inference(ennf_transformation,[],[f118])).
fof(f118,plain,(
  ! [X0,X1,X2,X3,X4,X5] : (eqangle(X0,X1,X4,X5,X2,X3,X4,X5) => para(X0,X1,X2,X3))),
  inference(rectify,[],[f39])).
fof(f39,axiom,(
  ! [X0,X1,X2,X3,X8,X9] : (eqangle(X0,X1,X8,X9,X2,X3,X8,X9) => para(X0,X1,X2,X3))),
  file('./6_GDD_FULL/41-60/56.p',ruleD39)).
fof(f804,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (eqangle(X0,X1,X2,X3,X4,X5,X2,X3) | ~para(X1,X0,X4,X5)) )),
  inference(resolution,[],[f376,f398])).
fof(f376,plain,(
  ( ! [X6,X4,X2,X0,X7,X5,X3,X1] : (~eqangle(X0,X1,X2,X3,X4,X5,X6,X7) | eqangle(X1,X0,X2,X3,X4,X5,X6,X7)) )),
  inference(cnf_transformation,[],[f192])).
fof(f192,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6,X7] : (eqangle(X1,X0,X2,X3,X4,X5,X6,X7) | ~eqangle(X0,X1,X2,X3,X4,X5,X6,X7))),
  inference(ennf_transformation,[],[f100])).
fof(f100,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6,X7] : (eqangle(X0,X1,X2,X3,X4,X5,X6,X7) => eqangle(X1,X0,X2,X3,X4,X5,X6,X7))),
  inference(rectify,[],[f18])).
fof(f18,axiom,(
  ! [X0,X1,X2,X3,X8,X9,X10,X11] : (eqangle(X0,X1,X2,X3,X8,X9,X10,X11) => eqangle(X1,X0,X2,X3,X8,X9,X10,X11))),
  file('./6_GDD_FULL/41-60/56.p',ruleD18)).
fof(f4537,plain,(
  para(sK20,sK22,sK22,sK21) | ~spl26_294),
  inference(avatar_component_clause,[],[f4536])).
fof(f31280,plain,(
  ~spl26_182 | spl26_341),
  inference(avatar_contradiction_clause,[],[f31279])).
fof(f31279,plain,(
  $false | (~spl26_182 | ~spl26_341)),
  inference(subsumption_resolution,[],[f5700,f11715])).
fof(f5700,plain,(
  ~cyclic(sK21,sK22,sK20,sK23) | ~spl26_341),
  inference(avatar_component_clause,[],[f5699])).
fof(f5699,plain,(
  spl26_341 <=> ~cyclic(sK21,sK22,sK20,sK23)),
  introduced(avatar_definition,[new_symbols(naming,[spl26_341])])).
fof(f30674,plain,(
  ~spl26_182 | spl26_305 | ~spl26_340 | ~spl26_498),
  inference(avatar_contradiction_clause,[],[f30673])).
fof(f30673,plain,(
  $false | (~spl26_182 | ~spl26_305 | ~spl26_340 | ~spl26_498)),
  inference(subsumption_resolution,[],[f5086,f30628])).
fof(f30628,plain,(
  ( ! [X84] : (cong(sK21,sK20,X84,X84)) ) | (~spl26_182 | ~spl26_340 | ~spl26_498)),
  inference(subsumption_resolution,[],[f30627,f11715])).
fof(f30627,plain,(
  ( ! [X84] : (cong(sK21,sK20,X84,X84) | ~cyclic(sK21,sK20,sK22,X84)) ) | (~spl26_340 | ~spl26_498)),
  inference(subsumption_resolution,[],[f28147,f8927])).
fof(f8927,plain,(
  para(sK22,sK21,sK22,sK20) | ~spl26_498),
  inference(avatar_component_clause,[],[f8926])).
fof(f8926,plain,(
  spl26_498 <=> para(sK22,sK21,sK22,sK20)),
  introduced(avatar_definition,[new_symbols(naming,[spl26_498])])).
fof(f28147,plain,(
  ( ! [X84] : (cong(sK21,sK20,X84,X84) | ~para(sK22,sK21,sK22,sK20) | ~cyclic(sK21,sK20,sK22,X84)) ) | ~spl26_340),
  inference(resolution,[],[f2246,f9627])).
fof(f9627,plain,(
  cyclic(sK21,sK20,sK22,sK23) | ~spl26_340),
  inference(resolution,[],[f5703,f373])).
fof(f373,plain,(
  ( ! [X2,X0,X3,X1] : (~cyclic(X0,X1,X2,X3) | cyclic(X0,X2,X1,X3)) )),
  inference(cnf_transformation,[],[f188])).
fof(f188,plain,(
  ! [X0,X1,X2,X3] : (cyclic(X0,X2,X1,X3) | ~cyclic(X0,X1,X2,X3))),
  inference(ennf_transformation,[],[f15])).
fof(f15,axiom,(
  ! [X0,X1,X2,X3] : (cyclic(X0,X1,X2,X3) => cyclic(X0,X2,X1,X3))),
  file('./6_GDD_FULL/41-60/56.p',ruleD15)).
fof(f5703,plain,(
  cyclic(sK21,sK22,sK20,sK23) | ~spl26_340),
  inference(avatar_component_clause,[],[f5702])).
fof(f5702,plain,(
  spl26_340 <=> cyclic(sK21,sK22,sK20,sK23)),
  introduced(avatar_definition,[new_symbols(naming,[spl26_340])])).
fof(f2246,plain,(
  ( ! [X80,X83,X81,X84,X82] : (~cyclic(X81,X82,X80,X84) | cong(X81,X82,X83,X83) | ~para(X80,X81,X80,X82) | ~cyclic(X81,X82,X80,X83)) )),
  inference(duplicate_literal_removal,[],[f2240])).
fof(f2240,plain,(
  ( ! [X80,X83,X81,X84,X82] : (~para(X80,X81,X80,X82) | cong(X81,X82,X83,X83) | ~cyclic(X81,X82,X80,X84) | ~cyclic(X81,X82,X80,X83) | ~cyclic(X81,X82,X80,X83)) )),
  inference(resolution,[],[f824,f402])).
fof(f5086,plain,(
  ~cong(sK21,sK20,sK21,sK21) | ~spl26_305),
  inference(avatar_component_clause,[],[f5085])).
fof(f5085,plain,(
  spl26_305 <=> ~cong(sK21,sK20,sK21,sK21)),
  introduced(avatar_definition,[new_symbols(naming,[spl26_305])])).
fof(f30503,plain,(
  ~spl26_2 | ~spl26_68 | ~spl26_182 | ~spl26_188 | ~spl26_462 | ~spl26_672 | ~spl26_804),
  inference(avatar_contradiction_clause,[],[f30502])).
fof(f30502,plain,(
  $false | (~spl26_2 | ~spl26_68 | ~spl26_182 | ~spl26_188 | ~spl26_462 | ~spl26_672 | ~spl26_804)),
  inference(subsumption_resolution,[],[f30501,f16610])).
fof(f16610,plain,(
  midp(sK22,sK25,sK22) | ~spl26_804),
  inference(avatar_component_clause,[],[f16609])).
fof(f16609,plain,(
  spl26_804 <=> midp(sK22,sK25,sK22)),
  introduced(avatar_definition,[new_symbols(naming,[spl26_804])])).
fof(f30501,plain,(
  ~midp(sK22,sK25,sK22) | (~spl26_2 | ~spl26_68 | ~spl26_182 | ~spl26_188 | ~spl26_462 | ~spl26_672)),
  inference(subsumption_resolution,[],[f10061,f30487])).
fof(f30487,plain,(
  ( ! [X1] : (midp(sK22,sK22,X1)) ) | (~spl26_2 | ~spl26_68 | ~spl26_182 | ~spl26_188 | ~spl26_462 | ~spl26_672)),
  inference(resolution,[],[f30480,f369])).
fof(f369,plain,(
  ( ! [X2,X0,X1] : (~midp(X2,X1,X0) | midp(X2,X0,X1)) )),
  inference(cnf_transformation,[],[f182])).
fof(f182,plain,(
  ! [X0,X1,X2] : (midp(X2,X0,X1) | ~midp(X2,X1,X0))),
  inference(ennf_transformation,[],[f97])).
fof(f97,plain,(
  ! [X0,X1,X2] : (midp(X2,X1,X0) => midp(X2,X0,X1))),
  inference(rectify,[],[f11])).
fof(f11,axiom,(
  ! [X0,X1,X6] : (midp(X6,X1,X0) => midp(X6,X0,X1))),
  file('./6_GDD_FULL/41-60/56.p',ruleD11)).
fof(f30480,plain,(
  ( ! [X2] : (midp(sK22,X2,sK22)) ) | (~spl26_2 | ~spl26_68 | ~spl26_182 | ~spl26_188 | ~spl26_462 | ~spl26_672)),
  inference(subsumption_resolution,[],[f30479,f28657])).
fof(f28657,plain,(
  ( ! [X2] : (midp(sK25,X2,sK20)) ) | (~spl26_2 | ~spl26_68 | ~spl26_182 | ~spl26_188 | ~spl26_462)),
  inference(subsumption_resolution,[],[f28656,f28082])).
fof(f28082,plain,(
  ( ! [X1] : (midp(sK23,X1,sK22)) ) | (~spl26_2 | ~spl26_182 | ~spl26_188 | ~spl26_462)),
  inference(resolution,[],[f28072,f7263])).
fof(f7263,plain,(
  ( ! [X0] : (~midp(sK20,X0,sK25) | midp(sK23,X0,sK22)) ) | ~spl26_182),
  inference(subsumption_resolution,[],[f785,f4151])).
fof(f785,plain,(
  ( ! [X0] : (~midp(sK20,X0,sK25) | midp(sK23,X0,sK22) | ~coll(sK23,X0,sK22)) )),
  inference(resolution,[],[f404,f477])).
fof(f477,plain,(
  para(sK20,sK23,sK25,sK22)),
  inference(cnf_transformation,[],[f358])).
fof(f358,plain,(
  ~cong(sK22,sK25,sK22,sK24) & para(sK20,sK23,sK25,sK22) & perp(sK20,sK23,sK20,sK25) & coll(sK24,sK20,sK21) & coll(sK24,sK22,sK23) & circle(sK23,sK20,sK21,sK22) & (cong(sK22,sK20,sK22,sK21) | cong(sK22,sK20,sK20,sK21))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK20,sK21,sK22,sK23,sK24,sK25])],[f320,f357])).
fof(f357,plain,(
  ? [X0,X1,X2,X3,X4,X5] : (~cong(X2,X5,X2,X4) & para(X0,X3,X5,X2) & perp(X0,X3,X0,X5) & coll(X4,X0,X1) & coll(X4,X2,X3) & circle(X3,X0,X1,X2) & (cong(X2,X0,X2,X1) | cong(X2,X0,X0,X1))) => (~cong(sK22,sK25,sK22,sK24) & para(sK20,sK23,sK25,sK22) & perp(sK20,sK23,sK20,sK25) & coll(sK24,sK20,sK21) & coll(sK24,sK22,sK23) & circle(sK23,sK20,sK21,sK22) & (cong(sK22,sK20,sK22,sK21) | cong(sK22,sK20,sK20,sK21)))),
  introduced(choice_axiom,[])).
fof(f320,plain,(
  ? [X0,X1,X2,X3,X4,X5] : (~cong(X2,X5,X2,X4) & para(X0,X3,X5,X2) & perp(X0,X3,X0,X5) & coll(X4,X0,X1) & coll(X4,X2,X3) & circle(X3,X0,X1,X2) & (cong(X2,X0,X2,X1) | cong(X2,X0,X0,X1)))),
  inference(flattening,[],[f319])).
fof(f319,plain,(
  ? [X0,X1,X2,X3,X4,X5] : (~cong(X2,X5,X2,X4) & (para(X0,X3,X5,X2) & perp(X0,X3,X0,X5) & coll(X4,X0,X1) & coll(X4,X2,X3) & circle(X3,X0,X1,X2) & (cong(X2,X0,X2,X1) | cong(X2,X0,X0,X1))))),
  inference(ennf_transformation,[],[f167])).
fof(f167,plain,(
  ~! [X0,X1,X2,X3,X4,X5] : ((para(X0,X3,X5,X2) & perp(X0,X3,X0,X5) & coll(X4,X0,X1) & coll(X4,X2,X3) & circle(X3,X0,X1,X2) & (cong(X2,X0,X2,X1) | cong(X2,X0,X0,X1))) => cong(X2,X5,X2,X4))),
  inference(rectify,[],[f96])).
fof(f96,negated_conjecture,(
  ~! [X0,X1,X6,X7,X3,X4] : ((para(X0,X7,X4,X6) & perp(X0,X7,X0,X4) & coll(X3,X0,X1) & coll(X3,X6,X7) & circle(X7,X0,X1,X6) & (cong(X6,X0,X6,X1) | cong(X6,X0,X0,X1))) => cong(X6,X4,X6,X3))),
  inference(negated_conjecture,[],[f95])).
fof(f95,conjecture,(
  ! [X0,X1,X6,X7,X3,X4] : ((para(X0,X7,X4,X6) & perp(X0,X7,X0,X4) & coll(X3,X0,X1) & coll(X3,X6,X7) & circle(X7,X0,X1,X6) & (cong(X6,X0,X6,X1) | cong(X6,X0,X0,X1))) => cong(X6,X4,X6,X3))),
  file('./6_GDD_FULL/41-60/56.p',exemplo6GDDFULL416056)).
fof(f404,plain,(
  ( ! [X4,X2,X0,X3,X1] : (~para(X3,X4,X1,X2) | ~coll(X4,X0,X2) | midp(X4,X0,X2) | ~midp(X3,X0,X1)) )),
  inference(cnf_transformation,[],[f230])).
fof(f230,plain,(
  ! [X0,X1,X2,X3,X4] : (midp(X4,X0,X2) | ~coll(X4,X0,X2) | ~para(X3,X4,X1,X2) | ~midp(X3,X0,X1))),
  inference(flattening,[],[f229])).
fof(f229,plain,(
  ! [X0,X1,X2,X3,X4] : (midp(X4,X0,X2) | (~coll(X4,X0,X2) | ~para(X3,X4,X1,X2) | ~midp(X3,X0,X1)))),
  inference(ennf_transformation,[],[f125])).
fof(f125,plain,(
  ! [X0,X1,X2,X3,X4] : ((coll(X4,X0,X2) & para(X3,X4,X1,X2) & midp(X3,X0,X1)) => midp(X4,X0,X2))),
  inference(rectify,[],[f46])).
fof(f46,axiom,(
  ! [X0,X1,X2,X4,X5] : ((coll(X5,X0,X2) & para(X4,X5,X1,X2) & midp(X4,X0,X1)) => midp(X5,X0,X2))),
  file('./6_GDD_FULL/41-60/56.p',ruleD45)).
fof(f28072,plain,(
  ( ! [X64,X63] : (midp(sK20,X63,X64)) ) | (~spl26_2 | ~spl26_182 | ~spl26_188 | ~spl26_462)),
  inference(subsumption_resolution,[],[f28042,f25915])).
fof(f25915,plain,(
  ( ! [X66,X64,X65] : (perp(X64,X65,X65,X66)) ) | (~spl26_182 | ~spl26_188 | ~spl26_462)),
  inference(resolution,[],[f25887,f365])).
fof(f365,plain,(
  ( ! [X2,X0,X3,X1] : (~perp(X0,X1,X2,X3) | perp(X0,X1,X3,X2)) )),
  inference(cnf_transformation,[],[f176])).
fof(f176,plain,(
  ! [X0,X1,X2,X3] : (perp(X0,X1,X3,X2) | ~perp(X0,X1,X2,X3))),
  inference(ennf_transformation,[],[f7])).
fof(f7,axiom,(
  ! [X0,X1,X2,X3] : (perp(X0,X1,X2,X3) => perp(X0,X1,X3,X2))),
  file('./6_GDD_FULL/41-60/56.p',ruleD7)).
fof(f25887,plain,(
  ( ! [X39,X37,X38] : (perp(X37,X39,X38,X39)) ) | (~spl26_182 | ~spl26_188 | ~spl26_462)),
  inference(subsumption_resolution,[],[f25886,f22672])).
fof(f22672,plain,(
  ( ! [X48,X49] : (perp(X48,X48,X49,X48)) ) | (~spl26_182 | ~spl26_188 | ~spl26_462)),
  inference(resolution,[],[f22646,f365])).
fof(f22646,plain,(
  ( ! [X64,X65] : (perp(X65,X65,X65,X64)) ) | (~spl26_182 | ~spl26_188 | ~spl26_462)),
  inference(subsumption_resolution,[],[f22645,f11341])).
fof(f11341,plain,(
  ( ! [X74,X72,X75,X73] : (~para(X72,X73,X74,X75) | perp(X73,X72,X74,X75)) ) | (~spl26_182 | ~spl26_188)),
  inference(subsumption_resolution,[],[f11318,f9421])).
fof(f9421,plain,(
  ( ! [X17,X16] : (perp(X16,X17,X16,X17)) ) | (~spl26_182 | ~spl26_188)),
  inference(resolution,[],[f2252,f2070])).
fof(f2252,plain,(
  ( ! [X47,X45,X46,X44] : (~para(X44,X45,X46,X47) | perp(X44,X45,X46,X47)) ) | ~spl26_188),
  inference(avatar_component_clause,[],[f2251])).
fof(f2251,plain,(
  spl26_188 <=> ! [X44,X46,X45,X47] : (~para(X44,X45,X46,X47) | perp(X44,X45,X46,X47))),
  introduced(avatar_definition,[new_symbols(naming,[spl26_188])])).
fof(f11318,plain,(
  ( ! [X76,X74,X72,X77,X75,X73] : (~para(X72,X73,X74,X75) | ~perp(X76,X77,X76,X77) | perp(X73,X72,X74,X75)) )),
  inference(resolution,[],[f1965,f433])).
fof(f433,plain,(
  ( ! [X6,X4,X2,X0,X7,X5,X3,X1] : (~eqangle(X0,X1,X2,X3,X4,X5,X6,X7) | ~perp(X4,X5,X6,X7) | perp(X0,X1,X2,X3)) )),
  inference(cnf_transformation,[],[f281])).
fof(f281,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6,X7] : (perp(X0,X1,X2,X3) | ~perp(X4,X5,X6,X7) | ~eqangle(X0,X1,X2,X3,X4,X5,X6,X7))),
  inference(flattening,[],[f280])).
fof(f280,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6,X7] : (perp(X0,X1,X2,X3) | (~perp(X4,X5,X6,X7) | ~eqangle(X0,X1,X2,X3,X4,X5,X6,X7)))),
  inference(ennf_transformation,[],[f147])).
fof(f147,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6,X7] : ((perp(X4,X5,X6,X7) & eqangle(X0,X1,X2,X3,X4,X5,X6,X7)) => perp(X0,X1,X2,X3))),
  inference(rectify,[],[f75])).
fof(f75,axiom,(
  ! [X0,X1,X2,X3,X8,X9,X10,X11] : ((perp(X8,X9,X10,X11) & eqangle(X0,X1,X2,X3,X8,X9,X10,X11)) => perp(X0,X1,X2,X3))),
  file('./6_GDD_FULL/41-60/56.p',ruleD74)).
fof(f1965,plain,(
  ( ! [X26,X24,X23,X27,X25,X22] : (eqangle(X23,X22,X24,X25,X26,X27,X26,X27) | ~para(X22,X23,X24,X25)) )),
  inference(resolution,[],[f804,f379])).
fof(f22645,plain,(
  ( ! [X64,X65] : (para(X65,X65,X65,X64) | perp(X65,X65,X65,X64)) ) | (~spl26_182 | ~spl26_462)),
  inference(subsumption_resolution,[],[f22644,f9544])).
fof(f22644,plain,(
  ( ! [X64,X65] : (~cyclic(X64,X64,X65,X65) | para(X65,X65,X65,X64) | perp(X65,X65,X65,X64)) ) | ~spl26_462),
  inference(subsumption_resolution,[],[f22611,f8469])).
fof(f8469,plain,(
  ( ! [X37,X38] : (para(X37,X37,X38,X38)) ) | ~spl26_462),
  inference(avatar_component_clause,[],[f8468])).
fof(f8468,plain,(
  spl26_462 <=> ! [X38,X37] : para(X37,X37,X38,X38)),
  introduced(avatar_definition,[new_symbols(naming,[spl26_462])])).
fof(f22611,plain,(
  ( ! [X64,X65] : (~para(X64,X64,X65,X65) | ~cyclic(X64,X64,X65,X65) | para(X65,X65,X65,X64) | perp(X65,X65,X65,X64)) )),
  inference(resolution,[],[f3598,f430])).
fof(f430,plain,(
  ( ! [X2,X0,X3,X1] : (~eqangle(X0,X1,X2,X3,X2,X3,X0,X1) | para(X0,X1,X2,X3) | perp(X0,X1,X2,X3)) )),
  inference(cnf_transformation,[],[f275])).
fof(f275,plain,(
  ! [X0,X1,X2,X3] : (perp(X0,X1,X2,X3) | para(X0,X1,X2,X3) | ~eqangle(X0,X1,X2,X3,X2,X3,X0,X1))),
  inference(flattening,[],[f274])).
fof(f274,plain,(
  ! [X0,X1,X2,X3] : (perp(X0,X1,X2,X3) | (para(X0,X1,X2,X3) | ~eqangle(X0,X1,X2,X3,X2,X3,X0,X1)))),
  inference(ennf_transformation,[],[f72])).
fof(f72,axiom,(
  ! [X0,X1,X2,X3] : ((~para(X0,X1,X2,X3) & eqangle(X0,X1,X2,X3,X2,X3,X0,X1)) => perp(X0,X1,X2,X3))),
  file('./6_GDD_FULL/41-60/56.p',ruleD71)).
fof(f3598,plain,(
  ( ! [X10,X8,X7,X9] : (eqangle(X9,X10,X10,X7,X9,X8,X9,X10) | ~para(X7,X8,X9,X10) | ~cyclic(X7,X8,X9,X10)) )),
  inference(resolution,[],[f1012,f377])).
fof(f1012,plain,(
  ( ! [X17,X15,X18,X16] : (eqangle(X18,X15,X17,X18,X17,X18,X17,X16) | ~cyclic(X15,X16,X17,X18) | ~para(X15,X16,X17,X18)) )),
  inference(resolution,[],[f413,f376])).
fof(f413,plain,(
  ( ! [X2,X0,X3,X1] : (eqangle(X0,X3,X2,X3,X2,X3,X2,X1) | ~para(X0,X1,X2,X3) | ~cyclic(X0,X1,X2,X3)) )),
  inference(cnf_transformation,[],[f247])).
fof(f247,plain,(
  ! [X0,X1,X2,X3] : (eqangle(X0,X3,X2,X3,X2,X3,X2,X1) | ~para(X0,X1,X2,X3) | ~cyclic(X0,X1,X2,X3))),
  inference(flattening,[],[f246])).
fof(f246,plain,(
  ! [X0,X1,X2,X3] : (eqangle(X0,X3,X2,X3,X2,X3,X2,X1) | (~para(X0,X1,X2,X3) | ~cyclic(X0,X1,X2,X3)))),
  inference(ennf_transformation,[],[f55])).
fof(f55,axiom,(
  ! [X0,X1,X2,X3] : ((para(X0,X1,X2,X3) & cyclic(X0,X1,X2,X3)) => eqangle(X0,X3,X2,X3,X2,X3,X2,X1))),
  file('./6_GDD_FULL/41-60/56.p',ruleD54)).
fof(f25886,plain,(
  ( ! [X39,X37,X38] : (~perp(X38,X38,X37,X38) | perp(X37,X39,X38,X39)) ) | (~spl26_182 | ~spl26_188 | ~spl26_462)),
  inference(subsumption_resolution,[],[f25885,f11715])).
fof(f25885,plain,(
  ( ! [X39,X37,X38] : (~cyclic(X39,X38,X37,X38) | ~perp(X38,X38,X37,X38) | perp(X37,X39,X38,X39)) ) | (~spl26_182 | ~spl26_188 | ~spl26_462)),
  inference(subsumption_resolution,[],[f25849,f24810])).
fof(f24810,plain,(
  ( ! [X70,X71] : (cong(X70,X71,X70,X71)) ) | (~spl26_182 | ~spl26_188 | ~spl26_462)),
  inference(subsumption_resolution,[],[f24799,f12081])).
fof(f24799,plain,(
  ( ! [X70,X71] : (cong(X70,X71,X70,X71) | ~midp(X71,X71,X71)) ) | (~spl26_182 | ~spl26_188 | ~spl26_462)),
  inference(resolution,[],[f24772,f414])).
fof(f414,plain,(
  ( ! [X2,X0,X3,X1] : (~perp(X3,X2,X0,X1) | cong(X3,X0,X3,X1) | ~midp(X2,X0,X1)) )),
  inference(cnf_transformation,[],[f249])).
fof(f249,plain,(
  ! [X0,X1,X2,X3] : (cong(X3,X0,X3,X1) | ~perp(X3,X2,X0,X1) | ~midp(X2,X0,X1))),
  inference(flattening,[],[f248])).
fof(f248,plain,(
  ! [X0,X1,X2,X3] : (cong(X3,X0,X3,X1) | (~perp(X3,X2,X0,X1) | ~midp(X2,X0,X1)))),
  inference(ennf_transformation,[],[f134])).
fof(f134,plain,(
  ! [X0,X1,X2,X3] : ((perp(X3,X2,X0,X1) & midp(X2,X0,X1)) => cong(X3,X0,X3,X1))),
  inference(rectify,[],[f56])).
fof(f56,axiom,(
  ! [X0,X1,X6,X7] : ((perp(X7,X6,X0,X1) & midp(X6,X0,X1)) => cong(X7,X0,X7,X1))),
  file('./6_GDD_FULL/41-60/56.p',ruleD55)).
fof(f24772,plain,(
  ( ! [X92,X91] : (perp(X91,X92,X92,X92)) ) | (~spl26_182 | ~spl26_188 | ~spl26_462)),
  inference(subsumption_resolution,[],[f24771,f2252])).
fof(f24771,plain,(
  ( ! [X92,X91] : (para(X91,X92,X92,X92) | perp(X91,X92,X92,X92)) ) | (~spl26_182 | ~spl26_462)),
  inference(subsumption_resolution,[],[f24770,f11715])).
fof(f24770,plain,(
  ( ! [X92,X91] : (~cyclic(X91,X92,X92,X91) | para(X91,X92,X92,X92) | perp(X91,X92,X92,X92)) ) | ~spl26_462),
  inference(subsumption_resolution,[],[f24736,f8469])).
fof(f24736,plain,(
  ( ! [X92,X91] : (~para(X91,X91,X92,X92) | ~cyclic(X91,X92,X92,X91) | para(X91,X92,X92,X92) | perp(X91,X92,X92,X92)) )),
  inference(resolution,[],[f4575,f430])).
fof(f4575,plain,(
  ( ! [X187,X185,X188,X186,X184,X183] : (eqangle(X183,X184,X184,X185,X186,X187,X188,X185) | ~para(X188,X183,X186,X187) | ~cyclic(X183,X185,X184,X188)) )),
  inference(resolution,[],[f1253,f805])).
fof(f805,plain,(
  ( ! [X6,X8,X7,X9] : (eqangle(X6,X7,X7,X8,X9,X6,X9,X8) | ~cyclic(X6,X8,X7,X9)) )),
  inference(resolution,[],[f376,f399])).
fof(f1253,plain,(
  ( ! [X6,X4,X2,X0,X8,X7,X5,X3,X1,X9] : (~eqangle(X0,X1,X2,X3,X8,X9,X6,X7) | eqangle(X0,X1,X2,X3,X4,X5,X6,X7) | ~para(X8,X9,X4,X5)) )),
  inference(resolution,[],[f380,f398])).
fof(f380,plain,(
  ( ! [X6,X4,X2,X0,X10,X8,X7,X5,X3,X1,X11,X9] : (~eqangle(X4,X5,X6,X7,X8,X9,X10,X11) | eqangle(X0,X1,X2,X3,X8,X9,X10,X11) | ~eqangle(X0,X1,X2,X3,X4,X5,X6,X7)) )),
  inference(cnf_transformation,[],[f197])).
fof(f197,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11] : (eqangle(X0,X1,X2,X3,X8,X9,X10,X11) | ~eqangle(X4,X5,X6,X7,X8,X9,X10,X11) | ~eqangle(X0,X1,X2,X3,X4,X5,X6,X7))),
  inference(flattening,[],[f196])).
fof(f196,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11] : (eqangle(X0,X1,X2,X3,X8,X9,X10,X11) | (~eqangle(X4,X5,X6,X7,X8,X9,X10,X11) | ~eqangle(X0,X1,X2,X3,X4,X5,X6,X7)))),
  inference(ennf_transformation,[],[f104])).
fof(f104,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11] : ((eqangle(X4,X5,X6,X7,X8,X9,X10,X11) & eqangle(X0,X1,X2,X3,X4,X5,X6,X7)) => eqangle(X0,X1,X2,X3,X8,X9,X10,X11))),
  inference(rectify,[],[f22])).
fof(f22,axiom,(
  ! [X0,X1,X2,X3,X8,X9,X10,X11,X4,X5,X12,X13] : ((eqangle(X8,X9,X10,X11,X4,X5,X12,X13) & eqangle(X0,X1,X2,X3,X8,X9,X10,X11)) => eqangle(X0,X1,X2,X3,X4,X5,X12,X13))),
  file('./6_GDD_FULL/41-60/56.p',ruleD22)).
fof(f25849,plain,(
  ( ! [X39,X37,X38] : (~cong(X37,X38,X37,X38) | ~cyclic(X39,X38,X37,X38) | ~perp(X38,X38,X37,X38) | perp(X37,X39,X38,X39)) )),
  inference(resolution,[],[f4631,f433])).
fof(f4631,plain,(
  ( ! [X28,X26,X27] : (eqangle(X26,X27,X28,X27,X28,X28,X26,X28) | ~cong(X26,X28,X26,X28) | ~cyclic(X27,X28,X26,X28)) )),
  inference(resolution,[],[f1257,f825])).
fof(f825,plain,(
  ( ! [X6,X8,X7,X9] : (eqangle(X6,X7,X8,X7,X6,X9,X8,X9) | ~cyclic(X7,X9,X6,X8)) )),
  inference(resolution,[],[f379,f399])).
fof(f1257,plain,(
  ( ! [X39,X37,X35,X41,X38,X36,X40] : (~eqangle(X35,X36,X37,X38,X41,X39,X39,X40) | eqangle(X35,X36,X37,X38,X39,X40,X41,X40) | ~cong(X41,X39,X41,X40)) )),
  inference(resolution,[],[f380,f405])).
fof(f405,plain,(
  ( ! [X2,X0,X1] : (eqangle(X2,X0,X0,X1,X0,X1,X2,X1) | ~cong(X2,X0,X2,X1)) )),
  inference(cnf_transformation,[],[f231])).
fof(f231,plain,(
  ! [X0,X1,X2] : (eqangle(X2,X0,X0,X1,X0,X1,X2,X1) | ~cong(X2,X0,X2,X1))),
  inference(ennf_transformation,[],[f126])).
fof(f126,plain,(
  ! [X0,X1,X2] : (cong(X2,X0,X2,X1) => eqangle(X2,X0,X0,X1,X0,X1,X2,X1))),
  inference(rectify,[],[f47])).
fof(f47,axiom,(
  ! [X0,X1,X7] : (cong(X7,X0,X7,X1) => eqangle(X7,X0,X0,X1,X0,X1,X7,X1))),
  file('./6_GDD_FULL/41-60/56.p',ruleD46)).
fof(f28042,plain,(
  ( ! [X64,X63] : (midp(sK20,X63,X64) | ~perp(X63,X64,X64,sK22)) ) | (~spl26_2 | ~spl26_182 | ~spl26_188 | ~spl26_462)),
  inference(resolution,[],[f28008,f6054])).
fof(f6054,plain,(
  ( ! [X4,X5,X3] : (~midp(sK8(X3,X4,sK20,sK21,sK22),X5,X3) | midp(sK20,X5,X4) | ~perp(X3,X4,X4,sK22)) ) | (~spl26_2 | ~spl26_182)),
  inference(subsumption_resolution,[],[f6043,f4151])).
fof(f6043,plain,(
  ( ! [X4,X5,X3] : (~perp(X3,X4,X4,sK22) | ~coll(sK20,X5,X4) | midp(sK20,X5,X4) | ~midp(sK8(X3,X4,sK20,sK21,sK22),X5,X3)) ) | ~spl26_2),
  inference(resolution,[],[f5949,f404])).
fof(f5949,plain,(
  ( ! [X6,X7] : (para(sK8(X6,X7,sK20,sK21,sK22),sK20,X6,X7) | ~perp(X6,X7,X7,sK22)) ) | ~spl26_2),
  inference(resolution,[],[f490,f452])).
fof(f452,plain,(
  ( ! [X4,X2,X0,X3,X1] : (~cong(X4,X2,X4,X3) | ~perp(X0,X1,X1,X4) | para(sK8(X0,X1,X2,X3,X4),X2,X0,X1)) )),
  inference(cnf_transformation,[],[f338])).
fof(f338,plain,(
  ! [X0,X1,X2,X3,X4] : ((cong(X1,X2,X1,sK8(X0,X1,X2,X3,X4)) & para(sK8(X0,X1,X2,X3,X4),X2,X0,X1) & cong(X4,X2,X4,sK8(X0,X1,X2,X3,X4))) | ~perp(X0,X1,X1,X4) | ~cong(X4,X2,X4,X3))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK8])],[f301,f337])).
fof(f337,plain,(
  ! [X4,X3,X2,X1,X0] : (? [X5] : ((cong(X1,X2,X1,X5) & para(X5,X2,X0,X1) & cong(X4,X2,X4,X5)) | ~perp(X0,X1,X1,X4) | ~cong(X4,X2,X4,X3)) => ((cong(X1,X2,X1,sK8(X0,X1,X2,X3,X4)) & para(sK8(X0,X1,X2,X3,X4),X2,X0,X1) & cong(X4,X2,X4,sK8(X0,X1,X2,X3,X4))) | ~perp(X0,X1,X1,X4) | ~cong(X4,X2,X4,X3)))),
  introduced(choice_axiom,[])).
fof(f301,plain,(
  ! [X0,X1,X2,X3,X4] : ? [X5] : ((cong(X1,X2,X1,X5) & para(X5,X2,X0,X1) & cong(X4,X2,X4,X5)) | ~perp(X0,X1,X1,X4) | ~cong(X4,X2,X4,X3))),
  inference(flattening,[],[f300])).
fof(f300,plain,(
  ! [X0,X1,X2,X3,X4] : ? [X5] : ((cong(X1,X2,X1,X5) & para(X5,X2,X0,X1) & cong(X4,X2,X4,X5)) | (~perp(X0,X1,X1,X4) | ~cong(X4,X2,X4,X3)))),
  inference(ennf_transformation,[],[f157])).
fof(f157,plain,(
  ! [X0,X1,X2,X3,X4] : ? [X5] : ((perp(X0,X1,X1,X4) & cong(X4,X2,X4,X3)) => (cong(X1,X2,X1,X5) & para(X5,X2,X0,X1) & cong(X4,X2,X4,X5)))),
  inference(rectify,[],[f85])).
fof(f85,axiom,(
  ! [X0,X1,X2,X3,X7] : ? [X8] : ((perp(X0,X1,X1,X7) & cong(X7,X2,X7,X3)) => (cong(X1,X2,X1,X8) & para(X8,X2,X0,X1) & cong(X7,X2,X7,X8)))),
  file('./6_GDD_FULL/41-60/56.p',ruleX9)).
fof(f490,plain,(
  cong(sK22,sK20,sK22,sK21) | ~spl26_2),
  inference(avatar_component_clause,[],[f489])).
fof(f489,plain,(
  spl26_2 <=> cong(sK22,sK20,sK22,sK21)),
  introduced(avatar_definition,[new_symbols(naming,[spl26_2])])).
fof(f28008,plain,(
  ( ! [X8,X7] : (midp(X7,X8,X8)) ) | (~spl26_182 | ~spl26_188 | ~spl26_462)),
  inference(subsumption_resolution,[],[f27934,f4068])).
fof(f27934,plain,(
  ( ! [X8,X7] : (~coll(X7,X8,X8) | midp(X7,X8,X8)) ) | (~spl26_182 | ~spl26_188 | ~spl26_462)),
  inference(resolution,[],[f24810,f426])).
fof(f28656,plain,(
  ( ! [X2] : (midp(sK25,X2,sK20) | ~midp(sK23,X2,sK22)) ) | (~spl26_68 | ~spl26_182)),
  inference(subsumption_resolution,[],[f28638,f4151])).
fof(f28638,plain,(
  ( ! [X2] : (~coll(sK25,X2,sK20) | midp(sK25,X2,sK20) | ~midp(sK23,X2,sK22)) ) | ~spl26_68),
  inference(resolution,[],[f942,f404])).
fof(f942,plain,(
  para(sK23,sK25,sK22,sK20) | ~spl26_68),
  inference(avatar_component_clause,[],[f941])).
fof(f941,plain,(
  spl26_68 <=> para(sK23,sK25,sK22,sK20)),
  introduced(avatar_definition,[new_symbols(naming,[spl26_68])])).
fof(f30479,plain,(
  ( ! [X2] : (midp(sK22,X2,sK22) | ~midp(sK25,X2,sK20)) ) | (~spl26_182 | ~spl26_672)),
  inference(subsumption_resolution,[],[f30461,f4087])).
fof(f30461,plain,(
  ( ! [X2] : (~coll(sK22,X2,sK22) | midp(sK22,X2,sK22) | ~midp(sK25,X2,sK20)) ) | ~spl26_672),
  inference(resolution,[],[f12585,f404])).
fof(f12585,plain,(
  para(sK25,sK22,sK20,sK22) | ~spl26_672),
  inference(avatar_component_clause,[],[f12584])).
fof(f12584,plain,(
  spl26_672 <=> para(sK25,sK22,sK20,sK22)),
  introduced(avatar_definition,[new_symbols(naming,[spl26_672])])).
fof(f10061,plain,(
  ~midp(sK22,sK22,sK24) | ~midp(sK22,sK25,sK22)),
  inference(resolution,[],[f1939,f478])).
fof(f478,plain,(
  ~cong(sK22,sK25,sK22,sK24)),
  inference(cnf_transformation,[],[f358])).
fof(f1939,plain,(
  ( ! [X14,X12,X15,X13] : (cong(X12,X13,X14,X15) | ~midp(X14,X12,X15) | ~midp(X12,X13,X14)) )),
  inference(resolution,[],[f749,f526])).
fof(f749,plain,(
  ( ! [X4,X2,X0,X3,X1] : (~cong(X0,X1,X2,X4) | cong(X0,X1,X2,X3) | ~midp(X2,X4,X3)) )),
  inference(resolution,[],[f383,f427])).
fof(f383,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (~cong(X2,X3,X4,X5) | cong(X0,X1,X4,X5) | ~cong(X0,X1,X2,X3)) )),
  inference(cnf_transformation,[],[f201])).
fof(f201,plain,(
  ! [X0,X1,X2,X3,X4,X5] : (cong(X0,X1,X4,X5) | ~cong(X2,X3,X4,X5) | ~cong(X0,X1,X2,X3))),
  inference(flattening,[],[f200])).
fof(f200,plain,(
  ! [X0,X1,X2,X3,X4,X5] : (cong(X0,X1,X4,X5) | (~cong(X2,X3,X4,X5) | ~cong(X0,X1,X2,X3)))),
  inference(ennf_transformation,[],[f25])).
fof(f25,axiom,(
  ! [X0,X1,X2,X3,X4,X5] : ((cong(X2,X3,X4,X5) & cong(X0,X1,X2,X3)) => cong(X0,X1,X4,X5))),
  file('./6_GDD_FULL/41-60/56.p',ruleD25)).
fof(f19360,plain,(
  spl26_674 | ~spl26_188),
  inference(avatar_split_clause,[],[f19333,f2251,f12593])).
fof(f12593,plain,(
  spl26_674 <=> para(sK20,sK25,sK25,sK22)),
  introduced(avatar_definition,[new_symbols(naming,[spl26_674])])).
fof(f19333,plain,(
  para(sK20,sK25,sK25,sK22) | ~spl26_188),
  inference(resolution,[],[f10328,f363])).
fof(f363,plain,(
  ( ! [X2,X0,X3,X1] : (~para(X0,X1,X2,X3) | para(X2,X3,X0,X1)) )),
  inference(cnf_transformation,[],[f173])).
fof(f173,plain,(
  ! [X0,X1,X2,X3] : (para(X2,X3,X0,X1) | ~para(X0,X1,X2,X3))),
  inference(ennf_transformation,[],[f5])).
fof(f5,axiom,(
  ! [X0,X1,X2,X3] : (para(X0,X1,X2,X3) => para(X2,X3,X0,X1))),
  file('./6_GDD_FULL/41-60/56.p',ruleD5)).
fof(f10328,plain,(
  para(sK25,sK22,sK20,sK25) | ~spl26_188),
  inference(resolution,[],[f9454,f731])).
fof(f731,plain,(
  ( ! [X6,X7] : (~perp(X6,X7,sK20,sK23) | para(X6,X7,sK20,sK25)) )),
  inference(resolution,[],[f367,f476])).
fof(f476,plain,(
  perp(sK20,sK23,sK20,sK25)),
  inference(cnf_transformation,[],[f358])).
fof(f367,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (~perp(X2,X3,X4,X5) | para(X0,X1,X4,X5) | ~perp(X0,X1,X2,X3)) )),
  inference(cnf_transformation,[],[f179])).
fof(f179,plain,(
  ! [X0,X1,X2,X3,X4,X5] : (para(X0,X1,X4,X5) | ~perp(X2,X3,X4,X5) | ~perp(X0,X1,X2,X3))),
  inference(flattening,[],[f178])).
fof(f178,plain,(
  ! [X0,X1,X2,X3,X4,X5] : (para(X0,X1,X4,X5) | (~perp(X2,X3,X4,X5) | ~perp(X0,X1,X2,X3)))),
  inference(ennf_transformation,[],[f9])).
fof(f9,axiom,(
  ! [X0,X1,X2,X3,X4,X5] : ((perp(X2,X3,X4,X5) & perp(X0,X1,X2,X3)) => para(X0,X1,X4,X5))),
  file('./6_GDD_FULL/41-60/56.p',ruleD9)).
fof(f9454,plain,(
  perp(sK25,sK22,sK20,sK23) | ~spl26_188),
  inference(resolution,[],[f2252,f510])).
fof(f510,plain,(
  para(sK25,sK22,sK20,sK23)),
  inference(resolution,[],[f363,f477])).
fof(f19359,plain,(
  spl26_672 | ~spl26_675 | ~spl26_334),
  inference(avatar_split_clause,[],[f15171,f5676,f12590,f12584])).
fof(f12590,plain,(
  spl26_675 <=> ~para(sK20,sK25,sK25,sK22)),
  introduced(avatar_definition,[new_symbols(naming,[spl26_675])])).
fof(f5676,plain,(
  spl26_334 <=> cong(sK20,sK25,sK20,sK22)),
  introduced(avatar_definition,[new_symbols(naming,[spl26_334])])).
fof(f15171,plain,(
  ~para(sK20,sK25,sK25,sK22) | para(sK25,sK22,sK20,sK22) | ~spl26_334),
  inference(resolution,[],[f2200,f5677])).
fof(f5677,plain,(
  cong(sK20,sK25,sK20,sK22) | ~spl26_334),
  inference(avatar_component_clause,[],[f5676])).
fof(f2200,plain,(
  ( ! [X23,X21,X22] : (~cong(X21,X22,X21,X23) | ~para(X21,X22,X22,X23) | para(X22,X23,X21,X23)) )),
  inference(resolution,[],[f819,f432])).
fof(f432,plain,(
  ( ! [X6,X4,X2,X0,X7,X5,X3,X1] : (~eqangle(X0,X1,X2,X3,X4,X5,X6,X7) | ~para(X4,X5,X6,X7) | para(X0,X1,X2,X3)) )),
  inference(cnf_transformation,[],[f279])).
fof(f279,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6,X7] : (para(X0,X1,X2,X3) | ~para(X4,X5,X6,X7) | ~eqangle(X0,X1,X2,X3,X4,X5,X6,X7))),
  inference(flattening,[],[f278])).
fof(f278,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6,X7] : (para(X0,X1,X2,X3) | (~para(X4,X5,X6,X7) | ~eqangle(X0,X1,X2,X3,X4,X5,X6,X7)))),
  inference(ennf_transformation,[],[f146])).
fof(f146,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6,X7] : ((para(X4,X5,X6,X7) & eqangle(X0,X1,X2,X3,X4,X5,X6,X7)) => para(X0,X1,X2,X3))),
  inference(rectify,[],[f74])).
fof(f74,axiom,(
  ! [X0,X1,X2,X3,X8,X9,X10,X11] : ((para(X8,X9,X10,X11) & eqangle(X0,X1,X2,X3,X8,X9,X10,X11)) => para(X0,X1,X2,X3))),
  file('./6_GDD_FULL/41-60/56.p',ruleD73)).
fof(f819,plain,(
  ( ! [X12,X10,X11] : (eqangle(X10,X11,X12,X11,X12,X10,X10,X11) | ~cong(X12,X10,X12,X11)) )),
  inference(resolution,[],[f378,f405])).
fof(f378,plain,(
  ( ! [X6,X4,X2,X0,X7,X5,X3,X1] : (~eqangle(X0,X1,X2,X3,X4,X5,X6,X7) | eqangle(X4,X5,X6,X7,X0,X1,X2,X3)) )),
  inference(cnf_transformation,[],[f194])).
fof(f194,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6,X7] : (eqangle(X4,X5,X6,X7,X0,X1,X2,X3) | ~eqangle(X0,X1,X2,X3,X4,X5,X6,X7))),
  inference(ennf_transformation,[],[f102])).
fof(f102,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6,X7] : (eqangle(X0,X1,X2,X3,X4,X5,X6,X7) => eqangle(X4,X5,X6,X7,X0,X1,X2,X3))),
  inference(rectify,[],[f20])).
fof(f20,axiom,(
  ! [X0,X1,X2,X3,X8,X9,X10,X11] : (eqangle(X0,X1,X2,X3,X8,X9,X10,X11) => eqangle(X8,X9,X10,X11,X0,X1,X2,X3))),
  file('./6_GDD_FULL/41-60/56.p',ruleD20)).
fof(f19007,plain,(
  spl26_804 | ~spl26_786),
  inference(avatar_split_clause,[],[f18995,f16361,f16609])).
fof(f16361,plain,(
  spl26_786 <=> midp(sK22,sK22,sK25)),
  introduced(avatar_definition,[new_symbols(naming,[spl26_786])])).
fof(f18995,plain,(
  midp(sK22,sK25,sK22) | ~spl26_786),
  inference(resolution,[],[f16362,f369])).
fof(f16362,plain,(
  midp(sK22,sK22,sK25) | ~spl26_786),
  inference(avatar_component_clause,[],[f16361])).
fof(f18991,plain,(
  spl26_786 | ~spl26_18 | ~spl26_312 | ~spl26_780),
  inference(avatar_split_clause,[],[f18990,f16347,f5362,f638,f16361])).
fof(f638,plain,(
  spl26_18 <=> cong(sK20,sK20,sK20,sK25)),
  introduced(avatar_definition,[new_symbols(naming,[spl26_18])])).
fof(f5362,plain,(
  spl26_312 <=> cong(sK20,sK20,sK20,sK22)),
  introduced(avatar_definition,[new_symbols(naming,[spl26_312])])).
fof(f16347,plain,(
  spl26_780 <=> para(sK20,sK22,sK20,sK25)),
  introduced(avatar_definition,[new_symbols(naming,[spl26_780])])).
fof(f18990,plain,(
  midp(sK22,sK22,sK25) | (~spl26_18 | ~spl26_312 | ~spl26_780)),
  inference(subsumption_resolution,[],[f18989,f16348])).
fof(f16348,plain,(
  para(sK20,sK22,sK20,sK25) | ~spl26_780),
  inference(avatar_component_clause,[],[f16347])).
fof(f18989,plain,(
  midp(sK22,sK22,sK25) | ~para(sK20,sK22,sK20,sK25) | (~spl26_18 | ~spl26_312)),
  inference(subsumption_resolution,[],[f18954,f1472])).
fof(f1472,plain,(
  coll(sK22,sK22,sK25)),
  inference(resolution,[],[f1463,f359])).
fof(f1463,plain,(
  coll(sK22,sK25,sK22)),
  inference(resolution,[],[f1451,f360])).
fof(f1451,plain,(
  coll(sK25,sK22,sK22)),
  inference(resolution,[],[f1441,f425])).
fof(f1441,plain,(
  para(sK25,sK22,sK25,sK22)),
  inference(resolution,[],[f723,f510])).
fof(f723,plain,(
  ( ! [X0,X1] : (~para(X0,X1,sK20,sK23) | para(X0,X1,sK25,sK22)) )),
  inference(resolution,[],[f364,f477])).
fof(f364,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (~para(X2,X3,X4,X5) | para(X0,X1,X4,X5) | ~para(X0,X1,X2,X3)) )),
  inference(cnf_transformation,[],[f175])).
fof(f175,plain,(
  ! [X0,X1,X2,X3,X4,X5] : (para(X0,X1,X4,X5) | ~para(X2,X3,X4,X5) | ~para(X0,X1,X2,X3))),
  inference(flattening,[],[f174])).
fof(f174,plain,(
  ! [X0,X1,X2,X3,X4,X5] : (para(X0,X1,X4,X5) | (~para(X2,X3,X4,X5) | ~para(X0,X1,X2,X3)))),
  inference(ennf_transformation,[],[f6])).
fof(f6,axiom,(
  ! [X0,X1,X2,X3,X4,X5] : ((para(X2,X3,X4,X5) & para(X0,X1,X2,X3)) => para(X0,X1,X4,X5))),
  file('./6_GDD_FULL/41-60/56.p',ruleD6)).
fof(f18954,plain,(
  midp(sK22,sK22,sK25) | ~coll(sK22,sK22,sK25) | ~para(sK20,sK22,sK20,sK25) | (~spl26_18 | ~spl26_312)),
  inference(resolution,[],[f2239,f9252])).
fof(f9252,plain,(
  circle(sK20,sK20,sK22,sK25) | (~spl26_18 | ~spl26_312)),
  inference(resolution,[],[f7893,f5363])).
fof(f5363,plain,(
  cong(sK20,sK20,sK20,sK22) | ~spl26_312),
  inference(avatar_component_clause,[],[f5362])).
fof(f7893,plain,(
  ( ! [X0] : (~cong(sK20,sK20,sK20,X0) | circle(sK20,sK20,X0,sK25)) ) | ~spl26_18),
  inference(resolution,[],[f639,f370])).
fof(f370,plain,(
  ( ! [X2,X0,X3,X1] : (~cong(X3,X0,X3,X2) | circle(X3,X0,X1,X2) | ~cong(X3,X0,X3,X1)) )),
  inference(cnf_transformation,[],[f184])).
fof(f184,plain,(
  ! [X0,X1,X2,X3] : (circle(X3,X0,X1,X2) | ~cong(X3,X0,X3,X2) | ~cong(X3,X0,X3,X1))),
  inference(flattening,[],[f183])).
fof(f183,plain,(
  ! [X0,X1,X2,X3] : (circle(X3,X0,X1,X2) | (~cong(X3,X0,X3,X2) | ~cong(X3,X0,X3,X1)))),
  inference(ennf_transformation,[],[f98])).
fof(f98,plain,(
  ! [X0,X1,X2,X3] : ((cong(X3,X0,X3,X2) & cong(X3,X0,X3,X1)) => circle(X3,X0,X1,X2))),
  inference(rectify,[],[f12])).
fof(f12,axiom,(
  ! [X0,X1,X2,X7] : ((cong(X7,X0,X7,X2) & cong(X7,X0,X7,X1)) => circle(X7,X0,X1,X2))),
  file('./6_GDD_FULL/41-60/56.p',ruleD12)).
fof(f639,plain,(
  cong(sK20,sK20,sK20,sK25) | ~spl26_18),
  inference(avatar_component_clause,[],[f638])).
fof(f2239,plain,(
  ( ! [X78,X76,X79,X77] : (~circle(X79,X76,X77,X78) | midp(X77,X77,X78) | ~coll(X77,X77,X78) | ~para(X76,X77,X76,X78)) )),
  inference(resolution,[],[f824,f410])).
fof(f410,plain,(
  ( ! [X4,X2,X0,X3,X1] : (~eqangle(X0,X1,X0,X2,X3,X1,X3,X4) | midp(X4,X1,X2) | ~coll(X4,X1,X2) | ~circle(X3,X0,X1,X2)) )),
  inference(cnf_transformation,[],[f241])).
fof(f241,plain,(
  ! [X0,X1,X2,X3,X4] : (midp(X4,X1,X2) | ~eqangle(X0,X1,X0,X2,X3,X1,X3,X4) | ~coll(X4,X1,X2) | ~circle(X3,X0,X1,X2))),
  inference(flattening,[],[f240])).
fof(f240,plain,(
  ! [X0,X1,X2,X3,X4] : (midp(X4,X1,X2) | (~eqangle(X0,X1,X0,X2,X3,X1,X3,X4) | ~coll(X4,X1,X2) | ~circle(X3,X0,X1,X2)))),
  inference(ennf_transformation,[],[f131])).
fof(f131,plain,(
  ! [X0,X1,X2,X3,X4] : ((eqangle(X0,X1,X0,X2,X3,X1,X3,X4) & coll(X4,X1,X2) & circle(X3,X0,X1,X2)) => midp(X4,X1,X2))),
  inference(rectify,[],[f52])).
fof(f52,axiom,(
  ! [X0,X1,X2,X7,X6] : ((eqangle(X0,X1,X0,X2,X7,X1,X7,X6) & coll(X6,X1,X2) & circle(X7,X0,X1,X2)) => midp(X6,X1,X2))),
  file('./6_GDD_FULL/41-60/56.p',ruleD51)).
fof(f18946,plain,(
  spl26_808 | ~spl26_188),
  inference(avatar_split_clause,[],[f18928,f2251,f16738])).
fof(f16738,plain,(
  spl26_808 <=> para(sK20,sK25,sK22,sK25)),
  introduced(avatar_definition,[new_symbols(naming,[spl26_808])])).
fof(f18928,plain,(
  para(sK20,sK25,sK22,sK25) | ~spl26_188),
  inference(resolution,[],[f10311,f363])).
fof(f10311,plain,(
  para(sK22,sK25,sK20,sK25) | ~spl26_188),
  inference(resolution,[],[f9451,f731])).
fof(f9451,plain,(
  perp(sK22,sK25,sK20,sK23) | ~spl26_188),
  inference(resolution,[],[f2252,f521])).
fof(f521,plain,(
  para(sK22,sK25,sK20,sK23)),
  inference(resolution,[],[f507,f363])).
fof(f507,plain,(
  para(sK20,sK23,sK22,sK25)),
  inference(resolution,[],[f362,f477])).
fof(f362,plain,(
  ( ! [X2,X0,X3,X1] : (~para(X0,X1,X2,X3) | para(X0,X1,X3,X2)) )),
  inference(cnf_transformation,[],[f172])).
fof(f172,plain,(
  ! [X0,X1,X2,X3] : (para(X0,X1,X3,X2) | ~para(X0,X1,X2,X3))),
  inference(ennf_transformation,[],[f4])).
fof(f4,axiom,(
  ! [X0,X1,X2,X3] : (para(X0,X1,X2,X3) => para(X0,X1,X3,X2))),
  file('./6_GDD_FULL/41-60/56.p',ruleD4)).
fof(f16752,plain,(
  spl26_780 | ~spl26_809 | ~spl26_338),
  inference(avatar_split_clause,[],[f16721,f5692,f16741,f16347])).
fof(f16741,plain,(
  spl26_809 <=> ~para(sK20,sK25,sK22,sK25)),
  introduced(avatar_definition,[new_symbols(naming,[spl26_809])])).
fof(f5692,plain,(
  spl26_338 <=> cong(sK22,sK20,sK22,sK25)),
  introduced(avatar_definition,[new_symbols(naming,[spl26_338])])).
fof(f16721,plain,(
  ~para(sK20,sK25,sK22,sK25) | para(sK20,sK22,sK20,sK25) | ~spl26_338),
  inference(resolution,[],[f5693,f2022])).
fof(f2022,plain,(
  ( ! [X23,X21,X22] : (~cong(X21,X22,X21,X23) | ~para(X22,X23,X21,X23) | para(X22,X21,X22,X23)) )),
  inference(resolution,[],[f806,f432])).
fof(f806,plain,(
  ( ! [X12,X10,X11] : (eqangle(X10,X11,X10,X12,X10,X12,X11,X12) | ~cong(X11,X10,X11,X12)) )),
  inference(resolution,[],[f376,f405])).
fof(f5693,plain,(
  cong(sK22,sK20,sK22,sK25) | ~spl26_338),
  inference(avatar_component_clause,[],[f5692])).
fof(f16698,plain,(
  spl26_338 | ~spl26_16 | ~spl26_56 | ~spl26_188),
  inference(avatar_split_clause,[],[f16697,f2251,f911,f629,f5692])).
fof(f629,plain,(
  spl26_16 <=> midp(sK23,sK20,sK25)),
  introduced(avatar_definition,[new_symbols(naming,[spl26_16])])).
fof(f911,plain,(
  spl26_56 <=> para(sK22,sK23,sK20,sK25)),
  introduced(avatar_definition,[new_symbols(naming,[spl26_56])])).
fof(f16697,plain,(
  cong(sK22,sK20,sK22,sK25) | (~spl26_16 | ~spl26_56 | ~spl26_188)),
  inference(subsumption_resolution,[],[f16692,f630])).
fof(f630,plain,(
  midp(sK23,sK20,sK25) | ~spl26_16),
  inference(avatar_component_clause,[],[f629])).
fof(f16692,plain,(
  cong(sK22,sK20,sK22,sK25) | ~midp(sK23,sK20,sK25) | (~spl26_56 | ~spl26_188)),
  inference(resolution,[],[f9449,f414])).
fof(f9449,plain,(
  perp(sK22,sK23,sK20,sK25) | (~spl26_56 | ~spl26_188)),
  inference(resolution,[],[f2252,f912])).
fof(f912,plain,(
  para(sK22,sK23,sK20,sK25) | ~spl26_56),
  inference(avatar_component_clause,[],[f911])).
fof(f16502,plain,(
  spl26_328 | ~spl26_18 | ~spl26_182 | ~spl26_312),
  inference(avatar_split_clause,[],[f16501,f5362,f2069,f638,f5645])).
fof(f5645,plain,(
  spl26_328 <=> midp(sK25,sK22,sK25)),
  introduced(avatar_definition,[new_symbols(naming,[spl26_328])])).
fof(f16501,plain,(
  midp(sK25,sK22,sK25) | (~spl26_18 | ~spl26_182 | ~spl26_312)),
  inference(subsumption_resolution,[],[f16500,f2070])).
fof(f16500,plain,(
  midp(sK25,sK22,sK25) | ~para(sK20,sK22,sK20,sK22) | (~spl26_18 | ~spl26_182 | ~spl26_312)),
  inference(subsumption_resolution,[],[f16338,f4087])).
fof(f16338,plain,(
  ~coll(sK25,sK22,sK25) | midp(sK25,sK22,sK25) | ~para(sK20,sK22,sK20,sK22) | (~spl26_18 | ~spl26_312)),
  inference(resolution,[],[f9252,f1150])).
fof(f1150,plain,(
  ( ! [X2,X0,X1] : (~circle(X2,X2,X1,X0) | ~coll(X0,X1,X0) | midp(X0,X1,X0) | ~para(X2,X1,X2,X1)) )),
  inference(resolution,[],[f410,f398])).
fof(f16404,plain,(
  spl26_332 | ~spl26_328),
  inference(avatar_split_clause,[],[f16392,f5645,f5667])).
fof(f5667,plain,(
  spl26_332 <=> midp(sK25,sK25,sK22)),
  introduced(avatar_definition,[new_symbols(naming,[spl26_332])])).
fof(f16392,plain,(
  midp(sK25,sK25,sK22) | ~spl26_328),
  inference(resolution,[],[f5646,f369])).
fof(f5646,plain,(
  midp(sK25,sK22,sK25) | ~spl26_328),
  inference(avatar_component_clause,[],[f5645])).
fof(f15651,plain,(
  spl26_174 | ~spl26_40),
  inference(avatar_split_clause,[],[f15642,f707,f1874])).
fof(f1874,plain,(
  spl26_174 <=> perp(sK23,sK20,sK20,sK20)),
  introduced(avatar_definition,[new_symbols(naming,[spl26_174])])).
fof(f707,plain,(
  spl26_40 <=> midp(sK20,sK20,sK25)),
  introduced(avatar_definition,[new_symbols(naming,[spl26_40])])).
fof(f15642,plain,(
  perp(sK23,sK20,sK20,sK20) | ~spl26_40),
  inference(resolution,[],[f8811,f366])).
fof(f366,plain,(
  ( ! [X2,X0,X3,X1] : (~perp(X0,X1,X2,X3) | perp(X2,X3,X0,X1)) )),
  inference(cnf_transformation,[],[f177])).
fof(f177,plain,(
  ! [X0,X1,X2,X3] : (perp(X2,X3,X0,X1) | ~perp(X0,X1,X2,X3))),
  inference(ennf_transformation,[],[f8])).
fof(f8,axiom,(
  ! [X0,X1,X2,X3] : (perp(X0,X1,X2,X3) => perp(X2,X3,X0,X1))),
  file('./6_GDD_FULL/41-60/56.p',ruleD8)).
fof(f8811,plain,(
  perp(sK20,sK20,sK23,sK20) | ~spl26_40),
  inference(resolution,[],[f8680,f365])).
fof(f8680,plain,(
  perp(sK20,sK20,sK20,sK23) | ~spl26_40),
  inference(subsumption_resolution,[],[f8662,f708])).
fof(f708,plain,(
  midp(sK20,sK20,sK25) | ~spl26_40),
  inference(avatar_component_clause,[],[f707])).
fof(f8662,plain,(
  perp(sK20,sK20,sK20,sK23) | ~midp(sK20,sK20,sK25)),
  inference(resolution,[],[f741,f1487])).
fof(f1487,plain,(
  ( ! [X2,X3] : (para(X2,X2,X2,X3) | ~midp(X2,X2,X3)) )),
  inference(resolution,[],[f600,f427])).
fof(f600,plain,(
  ( ! [X0,X1] : (~cong(X0,X0,X0,X1) | para(X0,X0,X0,X1)) )),
  inference(resolution,[],[f405,f397])).
fof(f741,plain,(
  ( ! [X10,X11] : (~para(X10,X11,sK20,sK25) | perp(X10,X11,sK20,sK23)) )),
  inference(resolution,[],[f368,f516])).
fof(f516,plain,(
  perp(sK20,sK25,sK20,sK23)),
  inference(resolution,[],[f366,f476])).
fof(f368,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (~perp(X2,X3,X4,X5) | perp(X0,X1,X4,X5) | ~para(X0,X1,X2,X3)) )),
  inference(cnf_transformation,[],[f181])).
fof(f181,plain,(
  ! [X0,X1,X2,X3,X4,X5] : (perp(X0,X1,X4,X5) | ~perp(X2,X3,X4,X5) | ~para(X0,X1,X2,X3))),
  inference(flattening,[],[f180])).
fof(f180,plain,(
  ! [X0,X1,X2,X3,X4,X5] : (perp(X0,X1,X4,X5) | (~perp(X2,X3,X4,X5) | ~para(X0,X1,X2,X3)))),
  inference(ennf_transformation,[],[f10])).
fof(f10,axiom,(
  ! [X0,X1,X2,X3,X4,X5] : ((perp(X2,X3,X4,X5) & para(X0,X1,X2,X3)) => perp(X0,X1,X4,X5))),
  file('./6_GDD_FULL/41-60/56.p',ruleD10)).
fof(f12846,plain,(
  spl26_244 | ~spl26_170 | ~spl26_276 | ~spl26_312),
  inference(avatar_split_clause,[],[f12845,f5362,f3971,f1865,f2931])).
fof(f2931,plain,(
  spl26_244 <=> para(sK21,sK22,sK20,sK22)),
  introduced(avatar_definition,[new_symbols(naming,[spl26_244])])).
fof(f1865,plain,(
  spl26_170 <=> cyclic(sK20,sK21,sK20,sK22)),
  introduced(avatar_definition,[new_symbols(naming,[spl26_170])])).
fof(f3971,plain,(
  spl26_276 <=> para(sK21,sK20,sK20,sK22)),
  introduced(avatar_definition,[new_symbols(naming,[spl26_276])])).
fof(f12845,plain,(
  para(sK21,sK22,sK20,sK22) | (~spl26_170 | ~spl26_276 | ~spl26_312)),
  inference(subsumption_resolution,[],[f12844,f9405])).
fof(f9405,plain,(
  para(sK20,sK22,sK20,sK20) | ~spl26_312),
  inference(resolution,[],[f2194,f5363])).
fof(f2194,plain,(
  ( ! [X0,X1] : (~cong(X0,X0,X0,X1) | para(X0,X1,X0,X0)) )),
  inference(resolution,[],[f819,f397])).
fof(f12844,plain,(
  para(sK21,sK22,sK20,sK22) | ~para(sK20,sK22,sK20,sK20) | (~spl26_170 | ~spl26_276)),
  inference(subsumption_resolution,[],[f9554,f3972])).
fof(f3972,plain,(
  para(sK21,sK20,sK20,sK22) | ~spl26_276),
  inference(avatar_component_clause,[],[f3971])).
fof(f9554,plain,(
  para(sK21,sK22,sK20,sK22) | ~para(sK21,sK20,sK20,sK22) | ~para(sK20,sK22,sK20,sK20) | ~spl26_170),
  inference(resolution,[],[f9102,f1034])).
fof(f1034,plain,(
  ( ! [X6,X8,X7,X9] : (~cyclic(X9,X8,X6,X7) | para(X9,X7,X6,X7) | ~para(X9,X8,X6,X7) | ~para(X6,X7,X6,X8)) )),
  inference(resolution,[],[f432,f413])).
fof(f9102,plain,(
  cyclic(sK21,sK20,sK20,sK22) | ~spl26_170),
  inference(resolution,[],[f1866,f374])).
fof(f374,plain,(
  ( ! [X2,X0,X3,X1] : (~cyclic(X0,X1,X2,X3) | cyclic(X1,X0,X2,X3)) )),
  inference(cnf_transformation,[],[f189])).
fof(f189,plain,(
  ! [X0,X1,X2,X3] : (cyclic(X1,X0,X2,X3) | ~cyclic(X0,X1,X2,X3))),
  inference(ennf_transformation,[],[f16])).
fof(f16,axiom,(
  ! [X0,X1,X2,X3] : (cyclic(X0,X1,X2,X3) => cyclic(X1,X0,X2,X3))),
  file('./6_GDD_FULL/41-60/56.p',ruleD16)).
fof(f1866,plain,(
  cyclic(sK20,sK21,sK20,sK22) | ~spl26_170),
  inference(avatar_component_clause,[],[f1865])).
fof(f12840,plain,(
  spl26_276 | ~spl26_264),
  inference(avatar_split_clause,[],[f12834,f3633,f3971])).
fof(f3633,plain,(
  spl26_264 <=> para(sK20,sK21,sK20,sK22)),
  introduced(avatar_definition,[new_symbols(naming,[spl26_264])])).
fof(f12834,plain,(
  para(sK21,sK20,sK20,sK22) | ~spl26_264),
  inference(resolution,[],[f3634,f1961])).
fof(f3634,plain,(
  para(sK20,sK21,sK20,sK22) | ~spl26_264),
  inference(avatar_component_clause,[],[f3633])).
fof(f12816,plain,(
  spl26_264 | ~spl26_350),
  inference(avatar_split_clause,[],[f12815,f5891,f3633])).
fof(f5891,plain,(
  spl26_350 <=> midp(sK21,sK21,sK22)),
  introduced(avatar_definition,[new_symbols(naming,[spl26_350])])).
fof(f12815,plain,(
  para(sK20,sK21,sK20,sK22) | ~spl26_350),
  inference(subsumption_resolution,[],[f12760,f5892])).
fof(f5892,plain,(
  midp(sK21,sK21,sK22) | ~spl26_350),
  inference(avatar_component_clause,[],[f5891])).
fof(f12760,plain,(
  ~midp(sK21,sK21,sK22) | para(sK20,sK21,sK20,sK22)),
  inference(resolution,[],[f2853,f473])).
fof(f473,plain,(
  circle(sK23,sK20,sK21,sK22)),
  inference(cnf_transformation,[],[f358])).
fof(f2853,plain,(
  ( ! [X2,X0,X3,X1] : (~circle(X0,X1,X2,X3) | ~midp(X2,X2,X3) | para(X1,X2,X1,X3)) )),
  inference(resolution,[],[f872,f397])).
fof(f872,plain,(
  ( ! [X6,X4,X7,X5,X3] : (eqangle(X7,X4,X6,X4,X7,X5,X6,X3) | ~circle(X6,X7,X4,X5) | ~midp(X3,X4,X5)) )),
  inference(resolution,[],[f409,f379])).
fof(f409,plain,(
  ( ! [X4,X2,X0,X3,X1] : (eqangle(X0,X1,X0,X2,X3,X1,X3,X4) | ~midp(X4,X1,X2) | ~circle(X3,X0,X1,X2)) )),
  inference(cnf_transformation,[],[f239])).
fof(f239,plain,(
  ! [X0,X1,X2,X3,X4] : (eqangle(X0,X1,X0,X2,X3,X1,X3,X4) | ~midp(X4,X1,X2) | ~circle(X3,X0,X1,X2))),
  inference(flattening,[],[f238])).
fof(f238,plain,(
  ! [X0,X1,X2,X3,X4] : (eqangle(X0,X1,X0,X2,X3,X1,X3,X4) | (~midp(X4,X1,X2) | ~circle(X3,X0,X1,X2)))),
  inference(ennf_transformation,[],[f130])).
fof(f130,plain,(
  ! [X0,X1,X2,X3,X4] : ((midp(X4,X1,X2) & circle(X3,X0,X1,X2)) => eqangle(X0,X1,X0,X2,X3,X1,X3,X4))),
  inference(rectify,[],[f51])).
fof(f51,axiom,(
  ! [X0,X1,X2,X7,X6] : ((midp(X6,X1,X2) & circle(X7,X0,X1,X2)) => eqangle(X0,X1,X0,X2,X7,X1,X7,X6))),
  file('./6_GDD_FULL/41-60/56.p',ruleD50)).
fof(f12808,plain,(
  spl26_350 | ~spl26_346),
  inference(avatar_split_clause,[],[f12796,f5878,f5891])).
fof(f5878,plain,(
  spl26_346 <=> midp(sK21,sK22,sK21)),
  introduced(avatar_definition,[new_symbols(naming,[spl26_346])])).
fof(f12796,plain,(
  midp(sK21,sK21,sK22) | ~spl26_346),
  inference(resolution,[],[f5879,f369])).
fof(f5879,plain,(
  midp(sK21,sK22,sK21) | ~spl26_346),
  inference(avatar_component_clause,[],[f5878])).
fof(f12789,plain,(
  spl26_346 | ~spl26_226 | ~spl26_362),
  inference(avatar_split_clause,[],[f12785,f5934,f2730,f5878])).
fof(f2730,plain,(
  spl26_226 <=> coll(sK22,sK21,sK21)),
  introduced(avatar_definition,[new_symbols(naming,[spl26_226])])).
fof(f5934,plain,(
  spl26_362 <=> cong(sK21,sK22,sK21,sK21)),
  introduced(avatar_definition,[new_symbols(naming,[spl26_362])])).
fof(f12785,plain,(
  midp(sK21,sK22,sK21) | (~spl26_226 | ~spl26_362)),
  inference(subsumption_resolution,[],[f12765,f3209])).
fof(f3209,plain,(
  coll(sK21,sK22,sK21) | ~spl26_226),
  inference(resolution,[],[f2731,f360])).
fof(f2731,plain,(
  coll(sK22,sK21,sK21) | ~spl26_226),
  inference(avatar_component_clause,[],[f2730])).
fof(f12765,plain,(
  ~coll(sK21,sK22,sK21) | midp(sK21,sK22,sK21) | ~spl26_362),
  inference(resolution,[],[f5935,f426])).
fof(f5935,plain,(
  cong(sK21,sK22,sK21,sK21) | ~spl26_362),
  inference(avatar_component_clause,[],[f5934])).
fof(f11942,plain,(
  spl26_68 | ~spl26_60),
  inference(avatar_split_clause,[],[f11930,f921,f941])).
fof(f921,plain,(
  spl26_60 <=> para(sK22,sK20,sK23,sK25)),
  introduced(avatar_definition,[new_symbols(naming,[spl26_60])])).
fof(f11930,plain,(
  para(sK23,sK25,sK22,sK20) | ~spl26_60),
  inference(resolution,[],[f922,f363])).
fof(f922,plain,(
  para(sK22,sK20,sK23,sK25) | ~spl26_60),
  inference(avatar_component_clause,[],[f921])).
fof(f11881,plain,(
  ~spl26_16 | ~spl26_54 | spl26_61 | ~spl26_182 | ~spl26_376),
  inference(avatar_contradiction_clause,[],[f11880])).
fof(f11880,plain,(
  $false | (~spl26_16 | ~spl26_54 | ~spl26_61 | ~spl26_182 | ~spl26_376)),
  inference(subsumption_resolution,[],[f11879,f9389])).
fof(f9389,plain,(
  midp(sK23,sK22,sK23) | (~spl26_54 | ~spl26_182 | ~spl26_376)),
  inference(resolution,[],[f9319,f369])).
fof(f9319,plain,(
  midp(sK23,sK23,sK22) | (~spl26_54 | ~spl26_182 | ~spl26_376)),
  inference(resolution,[],[f9251,f7263])).
fof(f9251,plain,(
  midp(sK20,sK23,sK25) | (~spl26_54 | ~spl26_376)),
  inference(resolution,[],[f9220,f369])).
fof(f9220,plain,(
  midp(sK20,sK25,sK23) | (~spl26_54 | ~spl26_376)),
  inference(resolution,[],[f908,f6439])).
fof(f6439,plain,(
  midp(sK20,sK20,sK22) | ~spl26_376),
  inference(avatar_component_clause,[],[f6438])).
fof(f6438,plain,(
  spl26_376 <=> midp(sK20,sK20,sK22)),
  introduced(avatar_definition,[new_symbols(naming,[spl26_376])])).
fof(f908,plain,(
  ( ! [X1] : (~midp(X1,sK20,sK22) | midp(X1,sK25,sK23)) ) | ~spl26_54),
  inference(avatar_component_clause,[],[f907])).
fof(f907,plain,(
  spl26_54 <=> ! [X1] : (midp(X1,sK25,sK23) | ~midp(X1,sK20,sK22))),
  introduced(avatar_definition,[new_symbols(naming,[spl26_54])])).
fof(f11879,plain,(
  ~midp(sK23,sK22,sK23) | (~spl26_16 | ~spl26_61)),
  inference(resolution,[],[f925,f7881])).
fof(f7881,plain,(
  ( ! [X0,X1] : (para(X0,sK20,X1,sK25) | ~midp(sK23,X0,X1)) ) | ~spl26_16),
  inference(resolution,[],[f630,f422])).
fof(f422,plain,(
  ( ! [X4,X2,X0,X3,X1] : (~midp(X4,X2,X3) | para(X0,X2,X1,X3) | ~midp(X4,X0,X1)) )),
  inference(cnf_transformation,[],[f262])).
fof(f262,plain,(
  ! [X0,X1,X2,X3,X4] : (para(X0,X2,X1,X3) | ~midp(X4,X2,X3) | ~midp(X4,X0,X1))),
  inference(flattening,[],[f261])).
fof(f261,plain,(
  ! [X0,X1,X2,X3,X4] : (para(X0,X2,X1,X3) | (~midp(X4,X2,X3) | ~midp(X4,X0,X1)))),
  inference(ennf_transformation,[],[f142])).
fof(f142,plain,(
  ! [X0,X1,X2,X3,X4] : ((midp(X4,X2,X3) & midp(X4,X0,X1)) => para(X0,X2,X1,X3))),
  inference(rectify,[],[f64])).
fof(f64,axiom,(
  ! [X0,X1,X2,X3,X6] : ((midp(X6,X2,X3) & midp(X6,X0,X1)) => para(X0,X2,X1,X3))),
  file('./6_GDD_FULL/41-60/56.p',ruleD63)).
fof(f925,plain,(
  ~para(sK22,sK20,sK23,sK25) | ~spl26_61),
  inference(avatar_component_clause,[],[f924])).
fof(f924,plain,(
  spl26_61 <=> ~para(sK22,sK20,sK23,sK25)),
  introduced(avatar_definition,[new_symbols(naming,[spl26_61])])).
fof(f11172,plain,(
  spl26_292 | ~spl26_208),
  inference(avatar_split_clause,[],[f11162,f2443,f4527])).
fof(f4527,plain,(
  spl26_292 <=> para(sK22,sK21,sK20,sK21)),
  introduced(avatar_definition,[new_symbols(naming,[spl26_292])])).
fof(f2443,plain,(
  spl26_208 <=> para(sK20,sK21,sK22,sK21)),
  introduced(avatar_definition,[new_symbols(naming,[spl26_208])])).
fof(f11162,plain,(
  para(sK22,sK21,sK20,sK21) | ~spl26_208),
  inference(resolution,[],[f2444,f363])).
fof(f2444,plain,(
  para(sK20,sK21,sK22,sK21) | ~spl26_208),
  inference(avatar_component_clause,[],[f2443])).
fof(f11155,plain,(
  spl26_208 | ~spl26_246),
  inference(avatar_split_clause,[],[f11135,f2940,f2443])).
fof(f2940,plain,(
  spl26_246 <=> para(sK20,sK21,sK21,sK22)),
  introduced(avatar_definition,[new_symbols(naming,[spl26_246])])).
fof(f11135,plain,(
  para(sK20,sK21,sK22,sK21) | ~spl26_246),
  inference(resolution,[],[f2941,f362])).
fof(f2941,plain,(
  para(sK20,sK21,sK21,sK22) | ~spl26_246),
  inference(avatar_component_clause,[],[f2940])).
fof(f9199,plain,(
  ~spl26_53 | spl26_54),
  inference(avatar_split_clause,[],[f7285,f907,f904])).
fof(f904,plain,(
  spl26_53 <=> ~para(sK20,sK25,sK22,sK23)),
  introduced(avatar_definition,[new_symbols(naming,[spl26_53])])).
fof(f7285,plain,(
  ( ! [X0] : (midp(X0,sK25,sK23) | ~para(sK20,sK25,sK22,sK23) | ~midp(X0,sK20,sK22)) )),
  inference(resolution,[],[f507,f423])).
fof(f423,plain,(
  ( ! [X4,X2,X0,X3,X1] : (~para(X0,X3,X1,X2) | midp(X4,X2,X3) | ~para(X0,X2,X1,X3) | ~midp(X4,X0,X1)) )),
  inference(cnf_transformation,[],[f264])).
fof(f264,plain,(
  ! [X0,X1,X2,X3,X4] : (midp(X4,X2,X3) | ~para(X0,X3,X1,X2) | ~para(X0,X2,X1,X3) | ~midp(X4,X0,X1))),
  inference(flattening,[],[f263])).
fof(f263,plain,(
  ! [X0,X1,X2,X3,X4] : (midp(X4,X2,X3) | (~para(X0,X3,X1,X2) | ~para(X0,X2,X1,X3) | ~midp(X4,X0,X1)))),
  inference(ennf_transformation,[],[f143])).
fof(f143,plain,(
  ! [X0,X1,X2,X3,X4] : ((para(X0,X3,X1,X2) & para(X0,X2,X1,X3) & midp(X4,X0,X1)) => midp(X4,X2,X3))),
  inference(rectify,[],[f65])).
fof(f65,axiom,(
  ! [X0,X1,X2,X3,X6] : ((para(X0,X3,X1,X2) & para(X0,X2,X1,X3) & midp(X6,X0,X1)) => midp(X6,X2,X3))),
  file('./6_GDD_FULL/41-60/56.p',ruleD64)).
fof(f9198,plain,(
  spl26_52 | ~spl26_56),
  inference(avatar_split_clause,[],[f9192,f911,f901])).
fof(f901,plain,(
  spl26_52 <=> para(sK20,sK25,sK22,sK23)),
  introduced(avatar_definition,[new_symbols(naming,[spl26_52])])).
fof(f9192,plain,(
  para(sK20,sK25,sK22,sK23) | ~spl26_56),
  inference(resolution,[],[f912,f363])).
fof(f9071,plain,(
  ~spl26_407 | ~spl26_16 | spl26_57),
  inference(avatar_split_clause,[],[f8103,f914,f629,f7467])).
fof(f7467,plain,(
  spl26_407 <=> ~midp(sK22,sK20,sK20)),
  introduced(avatar_definition,[new_symbols(naming,[spl26_407])])).
fof(f914,plain,(
  spl26_57 <=> ~para(sK22,sK23,sK20,sK25)),
  introduced(avatar_definition,[new_symbols(naming,[spl26_57])])).
fof(f8103,plain,(
  ~midp(sK22,sK20,sK20) | (~spl26_16 | ~spl26_57)),
  inference(resolution,[],[f7882,f915])).
fof(f915,plain,(
  ~para(sK22,sK23,sK20,sK25) | ~spl26_57),
  inference(avatar_component_clause,[],[f914])).
fof(f7882,plain,(
  ( ! [X2,X3] : (para(X2,sK23,X3,sK25) | ~midp(X2,sK20,X3)) ) | ~spl26_16),
  inference(resolution,[],[f630,f403])).
fof(f403,plain,(
  ( ! [X4,X2,X0,X3,X1] : (~midp(X4,X0,X2) | para(X3,X4,X1,X2) | ~midp(X3,X0,X1)) )),
  inference(cnf_transformation,[],[f228])).
fof(f228,plain,(
  ! [X0,X1,X2,X3,X4] : (para(X3,X4,X1,X2) | ~midp(X4,X0,X2) | ~midp(X3,X0,X1))),
  inference(flattening,[],[f227])).
fof(f227,plain,(
  ! [X0,X1,X2,X3,X4] : (para(X3,X4,X1,X2) | (~midp(X4,X0,X2) | ~midp(X3,X0,X1)))),
  inference(ennf_transformation,[],[f124])).
fof(f124,plain,(
  ! [X0,X1,X2,X3,X4] : ((midp(X4,X0,X2) & midp(X3,X0,X1)) => para(X3,X4,X1,X2))),
  inference(rectify,[],[f45])).
fof(f45,axiom,(
  ! [X0,X1,X2,X4,X5] : ((midp(X5,X0,X2) & midp(X4,X0,X1)) => para(X4,X5,X1,X2))),
  file('./6_GDD_FULL/41-60/56.p',ruleD44)).
fof(f9070,plain,(
  spl26_406 | ~spl26_0 | ~spl26_14),
  inference(avatar_split_clause,[],[f9053,f615,f483,f7464])).
fof(f7464,plain,(
  spl26_406 <=> midp(sK22,sK20,sK20)),
  introduced(avatar_definition,[new_symbols(naming,[spl26_406])])).
fof(f483,plain,(
  spl26_0 <=> cong(sK22,sK20,sK20,sK21)),
  introduced(avatar_definition,[new_symbols(naming,[spl26_0])])).
fof(f615,plain,(
  spl26_14 <=> coll(sK23,sK20,sK22)),
  introduced(avatar_definition,[new_symbols(naming,[spl26_14])])).
fof(f9053,plain,(
  midp(sK22,sK20,sK20) | (~spl26_0 | ~spl26_14)),
  inference(subsumption_resolution,[],[f9032,f3356])).
fof(f3356,plain,(
  coll(sK22,sK20,sK20) | ~spl26_14),
  inference(resolution,[],[f3337,f1021])).
fof(f1021,plain,(
  ( ! [X0] : (~coll(sK20,sK24,X0) | coll(X0,sK20,sK20)) )),
  inference(resolution,[],[f1019,f361])).
fof(f1019,plain,(
  coll(sK20,sK24,sK20)),
  inference(resolution,[],[f1005,f359])).
fof(f1005,plain,(
  coll(sK20,sK20,sK24)),
  inference(resolution,[],[f542,f493])).
fof(f493,plain,(
  coll(sK24,sK21,sK20)),
  inference(resolution,[],[f359,f475])).
fof(f475,plain,(
  coll(sK24,sK20,sK21)),
  inference(cnf_transformation,[],[f358])).
fof(f542,plain,(
  ( ! [X9] : (~coll(sK24,sK21,X9) | coll(X9,sK20,sK24)) )),
  inference(resolution,[],[f361,f493])).
fof(f3337,plain,(
  coll(sK20,sK24,sK22) | ~spl26_14),
  inference(resolution,[],[f3324,f538])).
fof(f538,plain,(
  ( ! [X5] : (~coll(sK22,sK23,X5) | coll(X5,sK24,sK22)) )),
  inference(resolution,[],[f361,f501])).
fof(f501,plain,(
  coll(sK22,sK23,sK24)),
  inference(resolution,[],[f495,f359])).
fof(f495,plain,(
  coll(sK22,sK24,sK23)),
  inference(resolution,[],[f360,f474])).
fof(f474,plain,(
  coll(sK24,sK22,sK23)),
  inference(cnf_transformation,[],[f358])).
fof(f3324,plain,(
  coll(sK22,sK23,sK20) | ~spl26_14),
  inference(resolution,[],[f3320,f360])).
fof(f3320,plain,(
  coll(sK23,sK22,sK20) | ~spl26_14),
  inference(resolution,[],[f616,f359])).
fof(f616,plain,(
  coll(sK23,sK20,sK22) | ~spl26_14),
  inference(avatar_component_clause,[],[f615])).
fof(f9032,plain,(
  ~coll(sK22,sK20,sK20) | midp(sK22,sK20,sK20) | ~spl26_0),
  inference(resolution,[],[f9019,f426])).
fof(f9019,plain,(
  cong(sK22,sK20,sK22,sK20) | ~spl26_0),
  inference(resolution,[],[f6479,f484])).
fof(f484,plain,(
  cong(sK22,sK20,sK20,sK21) | ~spl26_0),
  inference(avatar_component_clause,[],[f483])).
fof(f6479,plain,(
  ( ! [X2,X1] : (~cong(X1,X2,sK20,sK21) | cong(X1,X2,sK22,sK20)) ) | ~spl26_0),
  inference(resolution,[],[f6466,f383])).
fof(f6466,plain,(
  cong(sK20,sK21,sK22,sK20) | ~spl26_0),
  inference(resolution,[],[f484,f382])).
fof(f8936,plain,(
  spl26_498 | ~spl26_162),
  inference(avatar_split_clause,[],[f8923,f1842,f8926])).
fof(f8923,plain,(
  para(sK22,sK21,sK22,sK20) | ~spl26_162),
  inference(resolution,[],[f1843,f363])).
fof(f1843,plain,(
  para(sK22,sK20,sK22,sK21) | ~spl26_162),
  inference(avatar_component_clause,[],[f1842])).
fof(f8470,plain,(
  spl26_234 | spl26_234 | spl26_462 | ~spl26_182 | ~spl26_344),
  inference(avatar_split_clause,[],[f8466,f5857,f2069,f8468,f2788,f2788])).
fof(f2788,plain,(
  spl26_234 <=> ! [X11,X10,X12] : ~midp(X10,X11,X12)),
  introduced(avatar_definition,[new_symbols(naming,[spl26_234])])).
fof(f5857,plain,(
  spl26_344 <=> ! [X1,X5,X2,X6] : (~midp(X5,X1,X2) | para(X1,X1,X6,X6))),
  introduced(avatar_definition,[new_symbols(naming,[spl26_344])])).
fof(f8466,plain,(
  ( ! [X39,X37,X43,X41,X38,X44,X42,X40] : (para(X37,X37,X38,X38) | ~midp(X42,X39,X43) | ~midp(X44,X40,X41)) ) | (~spl26_182 | ~spl26_344)),
  inference(subsumption_resolution,[],[f8465,f4151])).
fof(f8465,plain,(
  ( ! [X39,X37,X43,X41,X38,X44,X42,X40] : (para(X37,X37,X38,X38) | ~midp(X42,X39,X43) | ~midp(X44,X40,X41) | ~coll(X43,X40,X41)) ) | (~spl26_182 | ~spl26_344)),
  inference(subsumption_resolution,[],[f8440,f4151])).
fof(f8440,plain,(
  ( ! [X39,X37,X43,X41,X38,X44,X42,X40] : (para(X37,X37,X38,X38) | ~coll(X39,X40,X41) | ~midp(X42,X39,X43) | ~midp(X44,X40,X41) | ~coll(X43,X40,X41)) ) | ~spl26_344),
  inference(resolution,[],[f5858,f1272])).
fof(f1272,plain,(
  ( ! [X30,X28,X26,X31,X29,X27,X25] : (midp(sK6(X26,X27,X28,X25,X30,X29,X31),X31,X26) | ~coll(X28,X26,X27) | ~midp(X29,X28,X25) | ~midp(X30,X26,X27) | ~coll(X25,X26,X27)) )),
  inference(resolution,[],[f448,f369])).
fof(f448,plain,(
  ( ! [X6,X4,X2,X0,X5,X3,X1] : (midp(sK6(X0,X1,X2,X3,X4,X5,X6),X0,X6) | ~coll(X3,X0,X1) | ~coll(X2,X0,X1) | ~midp(X5,X2,X3) | ~midp(X4,X0,X1)) )),
  inference(cnf_transformation,[],[f334])).
fof(f334,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6] : (midp(sK6(X0,X1,X2,X3,X4,X5,X6),X0,X6) | ~coll(X3,X0,X1) | ~coll(X2,X0,X1) | ~midp(X5,X2,X3) | ~midp(X4,X0,X1))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK6])],[f297,f333])).
fof(f333,plain,(
  ! [X6,X5,X4,X3,X2,X1,X0] : (? [X7] : (midp(X7,X0,X6) | ~coll(X3,X0,X1) | ~coll(X2,X0,X1) | ~midp(X5,X2,X3) | ~midp(X4,X0,X1)) => (midp(sK6(X0,X1,X2,X3,X4,X5,X6),X0,X6) | ~coll(X3,X0,X1) | ~coll(X2,X0,X1) | ~midp(X5,X2,X3) | ~midp(X4,X0,X1)))),
  introduced(choice_axiom,[])).
fof(f297,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6] : ? [X7] : (midp(X7,X0,X6) | ~coll(X3,X0,X1) | ~coll(X2,X0,X1) | ~midp(X5,X2,X3) | ~midp(X4,X0,X1))),
  inference(flattening,[],[f296])).
fof(f296,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6] : ? [X7] : (midp(X7,X0,X6) | (~coll(X3,X0,X1) | ~coll(X2,X0,X1) | ~midp(X5,X2,X3) | ~midp(X4,X0,X1)))),
  inference(ennf_transformation,[],[f155])).
fof(f155,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6] : ? [X7] : ((coll(X3,X0,X1) & coll(X2,X0,X1) & midp(X5,X2,X3) & midp(X4,X0,X1)) => midp(X7,X0,X6))),
  inference(rectify,[],[f83])).
fof(f83,axiom,(
  ! [X0,X1,X2,X3,X6,X16,X9] : ? [X8] : ((coll(X3,X0,X1) & coll(X2,X0,X1) & midp(X16,X2,X3) & midp(X6,X0,X1)) => midp(X8,X0,X9))),
  file('./6_GDD_FULL/41-60/56.p',ruleX7)).
fof(f5858,plain,(
  ( ! [X6,X2,X5,X1] : (~midp(X5,X1,X2) | para(X1,X1,X6,X6)) ) | ~spl26_344),
  inference(avatar_component_clause,[],[f5857])).
fof(f7996,plain,(
  ~spl26_182 | spl26_195),
  inference(avatar_contradiction_clause,[],[f7995])).
fof(f7995,plain,(
  $false | (~spl26_182 | ~spl26_195)),
  inference(subsumption_resolution,[],[f2306,f4151])).
fof(f2306,plain,(
  ~coll(sK20,sK20,sK25) | ~spl26_195),
  inference(avatar_component_clause,[],[f2305])).
fof(f2305,plain,(
  spl26_195 <=> ~coll(sK20,sK20,sK25)),
  introduced(avatar_definition,[new_symbols(naming,[spl26_195])])).
fof(f7941,plain,(
  spl26_40 | ~spl26_18 | ~spl26_194),
  inference(avatar_split_clause,[],[f7914,f2308,f638,f707])).
fof(f2308,plain,(
  spl26_194 <=> coll(sK20,sK20,sK25)),
  introduced(avatar_definition,[new_symbols(naming,[spl26_194])])).
fof(f7914,plain,(
  midp(sK20,sK20,sK25) | (~spl26_18 | ~spl26_194)),
  inference(subsumption_resolution,[],[f7895,f2309])).
fof(f2309,plain,(
  coll(sK20,sK20,sK25) | ~spl26_194),
  inference(avatar_component_clause,[],[f2308])).
fof(f7895,plain,(
  ~coll(sK20,sK20,sK25) | midp(sK20,sK20,sK25) | ~spl26_18),
  inference(resolution,[],[f639,f426])).
fof(f7879,plain,(
  ~spl26_17 | spl26_18),
  inference(avatar_split_clause,[],[f7277,f638,f632])).
fof(f632,plain,(
  spl26_17 <=> ~midp(sK23,sK20,sK25)),
  introduced(avatar_definition,[new_symbols(naming,[spl26_17])])).
fof(f7277,plain,(
  cong(sK20,sK20,sK20,sK25) | ~midp(sK23,sK20,sK25)),
  inference(resolution,[],[f476,f414])).
fof(f7878,plain,(
  spl26_16 | ~spl26_182 | ~spl26_376),
  inference(avatar_split_clause,[],[f7869,f6438,f2069,f629])).
fof(f7869,plain,(
  midp(sK23,sK20,sK25) | (~spl26_182 | ~spl26_376)),
  inference(resolution,[],[f6439,f7262])).
fof(f7262,plain,(
  ( ! [X1] : (~midp(sK20,X1,sK22) | midp(sK23,X1,sK25)) ) | ~spl26_182),
  inference(subsumption_resolution,[],[f786,f4151])).
fof(f786,plain,(
  ( ! [X1] : (~coll(sK23,X1,sK25) | midp(sK23,X1,sK25) | ~midp(sK20,X1,sK22)) )),
  inference(resolution,[],[f404,f507])).
fof(f7839,plain,(
  spl26_376 | ~spl26_176 | ~spl26_312),
  inference(avatar_split_clause,[],[f7812,f5362,f1884,f6438])).
fof(f1884,plain,(
  spl26_176 <=> coll(sK20,sK22,sK20)),
  introduced(avatar_definition,[new_symbols(naming,[spl26_176])])).
fof(f7812,plain,(
  midp(sK20,sK20,sK22) | (~spl26_176 | ~spl26_312)),
  inference(subsumption_resolution,[],[f7793,f3425])).
fof(f3425,plain,(
  coll(sK20,sK20,sK22) | ~spl26_176),
  inference(resolution,[],[f1885,f359])).
fof(f1885,plain,(
  coll(sK20,sK22,sK20) | ~spl26_176),
  inference(avatar_component_clause,[],[f1884])).
fof(f7793,plain,(
  ~coll(sK20,sK20,sK22) | midp(sK20,sK20,sK22) | ~spl26_312),
  inference(resolution,[],[f5363,f426])).
fof(f6686,plain,(
  ~spl26_186 | ~spl26_282),
  inference(avatar_contradiction_clause,[],[f6685])).
fof(f6685,plain,(
  $false | (~spl26_186 | ~spl26_282)),
  inference(subsumption_resolution,[],[f6683,f2249])).
fof(f2249,plain,(
  ( ! [X48,X49] : (~perp(X48,X49,X48,X49)) ) | ~spl26_186),
  inference(avatar_component_clause,[],[f2248])).
fof(f2248,plain,(
  spl26_186 <=> ! [X49,X48] : ~perp(X48,X49,X48,X49)),
  introduced(avatar_definition,[new_symbols(naming,[spl26_186])])).
fof(f6683,plain,(
  perp(sK21,sK20,sK21,sK20) | ~spl26_282),
  inference(resolution,[],[f4022,f365])).
fof(f4022,plain,(
  perp(sK21,sK20,sK20,sK21) | ~spl26_282),
  inference(avatar_component_clause,[],[f4021])).
fof(f4021,plain,(
  spl26_282 <=> perp(sK21,sK20,sK20,sK21)),
  introduced(avatar_definition,[new_symbols(naming,[spl26_282])])).
fof(f6489,plain,(
  spl26_360 | ~spl26_240),
  inference(avatar_split_clause,[],[f6486,f2921,f5925])).
fof(f5925,plain,(
  spl26_360 <=> midp(sK20,sK22,sK21)),
  introduced(avatar_definition,[new_symbols(naming,[spl26_360])])).
fof(f2921,plain,(
  spl26_240 <=> midp(sK20,sK21,sK22)),
  introduced(avatar_definition,[new_symbols(naming,[spl26_240])])).
fof(f6486,plain,(
  midp(sK20,sK22,sK21) | ~spl26_240),
  inference(resolution,[],[f2922,f369])).
fof(f2922,plain,(
  midp(sK20,sK21,sK22) | ~spl26_240),
  inference(avatar_component_clause,[],[f2921])).
fof(f6434,plain,(
  spl26_282 | ~spl26_252 | ~spl26_304),
  inference(avatar_split_clause,[],[f6433,f5088,f3034,f4021])).
fof(f3034,plain,(
  spl26_252 <=> perp(sK20,sK21,sK21,sK21)),
  introduced(avatar_definition,[new_symbols(naming,[spl26_252])])).
fof(f5088,plain,(
  spl26_304 <=> cong(sK21,sK20,sK21,sK21)),
  introduced(avatar_definition,[new_symbols(naming,[spl26_304])])).
fof(f6433,plain,(
  perp(sK21,sK20,sK20,sK21) | (~spl26_252 | ~spl26_304)),
  inference(subsumption_resolution,[],[f6414,f3035])).
fof(f3035,plain,(
  perp(sK20,sK21,sK21,sK21) | ~spl26_252),
  inference(avatar_component_clause,[],[f3034])).
fof(f6414,plain,(
  perp(sK21,sK20,sK20,sK21) | ~perp(sK20,sK21,sK21,sK21) | ~spl26_304),
  inference(resolution,[],[f5089,f1047])).
fof(f1047,plain,(
  ( ! [X21,X19,X20] : (~cong(X21,X19,X21,X20) | perp(X21,X19,X19,X20) | ~perp(X19,X20,X21,X20)) )),
  inference(resolution,[],[f433,f405])).
fof(f5089,plain,(
  cong(sK21,sK20,sK21,sK21) | ~spl26_304),
  inference(avatar_component_clause,[],[f5088])).
fof(f6432,plain,(
  ~spl26_82 | ~spl26_304),
  inference(avatar_contradiction_clause,[],[f6431])).
fof(f6431,plain,(
  $false | (~spl26_82 | ~spl26_304)),
  inference(subsumption_resolution,[],[f6423,f5996])).
fof(f5996,plain,(
  ( ! [X0,X1] : (~midp(X0,X1,X0)) ) | ~spl26_82),
  inference(subsumption_resolution,[],[f5986,f369])).
fof(f5986,plain,(
  ( ! [X0,X1] : (~midp(X0,X0,X1) | ~midp(X0,X1,X0)) ) | ~spl26_82),
  inference(resolution,[],[f1059,f530])).
fof(f530,plain,(
  ( ! [X2,X0,X1] : (cong(X0,X1,X0,X2) | ~midp(X0,X2,X1)) )),
  inference(resolution,[],[f382,f427])).
fof(f1059,plain,(
  ( ! [X6,X8,X7] : (~cong(X6,X7,X7,X8) | ~midp(X6,X7,X8)) ) | ~spl26_82),
  inference(avatar_component_clause,[],[f1058])).
fof(f1058,plain,(
  spl26_82 <=> ! [X7,X8,X6] : (~cong(X6,X7,X7,X8) | ~midp(X6,X7,X8))),
  introduced(avatar_definition,[new_symbols(naming,[spl26_82])])).
fof(f6423,plain,(
  midp(sK21,sK20,sK21) | ~spl26_304),
  inference(subsumption_resolution,[],[f6406,f844])).
fof(f844,plain,(
  coll(sK21,sK20,sK21)),
  inference(resolution,[],[f838,f359])).
fof(f838,plain,(
  coll(sK21,sK21,sK20)),
  inference(resolution,[],[f533,f496])).
fof(f496,plain,(
  coll(sK20,sK24,sK21)),
  inference(resolution,[],[f360,f475])).
fof(f533,plain,(
  ( ! [X0] : (~coll(sK20,sK24,X0) | coll(X0,sK21,sK20)) )),
  inference(resolution,[],[f361,f496])).
fof(f6406,plain,(
  ~coll(sK21,sK20,sK21) | midp(sK21,sK20,sK21) | ~spl26_304),
  inference(resolution,[],[f5089,f426])).
fof(f6425,plain,(
  spl26_181 | ~spl26_304),
  inference(avatar_contradiction_clause,[],[f6424])).
fof(f6424,plain,(
  $false | (~spl26_181 | ~spl26_304)),
  inference(subsumption_resolution,[],[f6423,f1896])).
fof(f1896,plain,(
  ~midp(sK21,sK20,sK21) | ~spl26_181),
  inference(avatar_component_clause,[],[f1895])).
fof(f1895,plain,(
  spl26_181 <=> ~midp(sK21,sK20,sK21)),
  introduced(avatar_definition,[new_symbols(naming,[spl26_181])])).
fof(f5939,plain,(
  spl26_1 | ~spl26_4 | ~spl26_80),
  inference(avatar_contradiction_clause,[],[f5938])).
fof(f5938,plain,(
  $false | (~spl26_1 | ~spl26_4 | ~spl26_80)),
  inference(subsumption_resolution,[],[f5937,f561])).
fof(f561,plain,(
  midp(sK22,sK20,sK21) | ~spl26_4),
  inference(avatar_component_clause,[],[f560])).
fof(f560,plain,(
  spl26_4 <=> midp(sK22,sK20,sK21)),
  introduced(avatar_definition,[new_symbols(naming,[spl26_4])])).
fof(f5937,plain,(
  ~midp(sK22,sK20,sK21) | (~spl26_1 | ~spl26_80)),
  inference(resolution,[],[f481,f1056])).
fof(f1056,plain,(
  ( ! [X10,X11,X9] : (cong(X9,X10,X10,X11) | ~midp(X9,X10,X11)) ) | ~spl26_80),
  inference(avatar_component_clause,[],[f1055])).
fof(f1055,plain,(
  spl26_80 <=> ! [X9,X11,X10] : (cong(X9,X10,X10,X11) | ~midp(X9,X10,X11))),
  introduced(avatar_definition,[new_symbols(naming,[spl26_80])])).
fof(f481,plain,(
  ~cong(sK22,sK20,sK20,sK21) | ~spl26_1),
  inference(avatar_component_clause,[],[f480])).
fof(f480,plain,(
  spl26_1 <=> ~cong(sK22,sK20,sK20,sK21)),
  introduced(avatar_definition,[new_symbols(naming,[spl26_1])])).
fof(f5936,plain,(
  ~spl26_361 | spl26_362 | ~spl26_222),
  inference(avatar_split_clause,[],[f4718,f2505,f5934,f5928])).
fof(f5928,plain,(
  spl26_361 <=> ~midp(sK20,sK22,sK21)),
  introduced(avatar_definition,[new_symbols(naming,[spl26_361])])).
fof(f2505,plain,(
  spl26_222 <=> perp(sK22,sK21,sK21,sK20)),
  introduced(avatar_definition,[new_symbols(naming,[spl26_222])])).
fof(f4718,plain,(
  cong(sK21,sK22,sK21,sK21) | ~midp(sK20,sK22,sK21) | ~spl26_222),
  inference(resolution,[],[f4493,f414])).
fof(f4493,plain,(
  perp(sK21,sK20,sK22,sK21) | ~spl26_222),
  inference(resolution,[],[f2506,f366])).
fof(f2506,plain,(
  perp(sK22,sK21,sK21,sK20) | ~spl26_222),
  inference(avatar_component_clause,[],[f2505])).
fof(f5875,plain,(
  ~spl26_0 | ~spl26_234),
  inference(avatar_contradiction_clause,[],[f5870])).
fof(f5870,plain,(
  $false | (~spl26_0 | ~spl26_234)),
  inference(resolution,[],[f2789,f5229])).
fof(f5229,plain,(
  midp(sK20,sK21,sK21) | ~spl26_0),
  inference(subsumption_resolution,[],[f5211,f848])).
fof(f848,plain,(
  coll(sK20,sK21,sK21)),
  inference(resolution,[],[f844,f360])).
fof(f5211,plain,(
  ~coll(sK20,sK21,sK21) | midp(sK20,sK21,sK21) | ~spl26_0),
  inference(resolution,[],[f5132,f426])).
fof(f5132,plain,(
  cong(sK20,sK21,sK20,sK21) | ~spl26_0),
  inference(resolution,[],[f5096,f5097])).
fof(f5097,plain,(
  cong(sK20,sK21,sK22,sK20) | ~spl26_0),
  inference(resolution,[],[f484,f382])).
fof(f5096,plain,(
  ( ! [X4,X3] : (~cong(X3,X4,sK22,sK20) | cong(X3,X4,sK20,sK21)) ) | ~spl26_0),
  inference(resolution,[],[f484,f383])).
fof(f2789,plain,(
  ( ! [X12,X10,X11] : (~midp(X10,X11,X12)) ) | ~spl26_234),
  inference(avatar_component_clause,[],[f2788])).
fof(f5859,plain,(
  spl26_344 | spl26_234 | ~spl26_182),
  inference(avatar_split_clause,[],[f5855,f2069,f2788,f5857])).
fof(f5855,plain,(
  ( ! [X6,X4,X2,X0,X5,X3,X1] : (~midp(X3,X0,X4) | ~midp(X5,X1,X2) | para(X1,X1,X6,X6)) ) | ~spl26_182),
  inference(subsumption_resolution,[],[f5854,f4151])).
fof(f5854,plain,(
  ( ! [X6,X4,X2,X0,X5,X3,X1] : (~midp(X3,X0,X4) | ~midp(X5,X1,X2) | para(X1,X1,X6,X6) | ~coll(X4,X1,X2)) ) | ~spl26_182),
  inference(subsumption_resolution,[],[f5853,f4151])).
fof(f5853,plain,(
  ( ! [X6,X4,X2,X0,X5,X3,X1] : (~coll(X0,X1,X2) | ~midp(X3,X0,X4) | ~midp(X5,X1,X2) | para(X1,X1,X6,X6) | ~coll(X4,X1,X2)) )),
  inference(duplicate_literal_removal,[],[f5849])).
fof(f5849,plain,(
  ( ! [X6,X4,X2,X0,X5,X3,X1] : (~coll(X0,X1,X2) | ~midp(X3,X0,X4) | ~midp(X5,X1,X2) | para(X1,X1,X6,X6) | ~coll(X4,X1,X2) | ~coll(X4,X1,X2) | ~coll(X0,X1,X2) | ~midp(X3,X0,X4) | ~midp(X5,X1,X2)) )),
  inference(resolution,[],[f1269,f448])).
fof(f1269,plain,(
  ( ! [X6,X4,X2,X0,X8,X7,X5,X3,X1] : (~midp(sK6(X1,X2,X3,X0,X5,X4,X8),X6,X7) | ~coll(X3,X1,X2) | ~midp(X4,X3,X0) | ~midp(X5,X1,X2) | para(X6,X1,X7,X8) | ~coll(X0,X1,X2)) )),
  inference(resolution,[],[f448,f422])).
fof(f5678,plain,(
  ~spl26_333 | spl26_334),
  inference(avatar_split_clause,[],[f4353,f5676,f5670])).
fof(f5670,plain,(
  spl26_333 <=> ~midp(sK25,sK25,sK22)),
  introduced(avatar_definition,[new_symbols(naming,[spl26_333])])).
fof(f4353,plain,(
  cong(sK20,sK25,sK20,sK22) | ~midp(sK25,sK25,sK22)),
  inference(resolution,[],[f2261,f414])).
fof(f2261,plain,(
  perp(sK20,sK25,sK25,sK22)),
  inference(resolution,[],[f1461,f366])).
fof(f1461,plain,(
  perp(sK25,sK22,sK20,sK25)),
  inference(resolution,[],[f739,f510])).
fof(f739,plain,(
  ( ! [X6,X7] : (~para(X6,X7,sK20,sK23) | perp(X6,X7,sK20,sK25)) )),
  inference(resolution,[],[f368,f476])).
fof(f5659,plain,(
  ~spl26_181 | spl26_3 | ~spl26_216),
  inference(avatar_split_clause,[],[f5658,f2483,f486,f1895])).
fof(f486,plain,(
  spl26_3 <=> ~cong(sK22,sK20,sK22,sK21)),
  introduced(avatar_definition,[new_symbols(naming,[spl26_3])])).
fof(f2483,plain,(
  spl26_216 <=> perp(sK20,sK21,sK22,sK21)),
  introduced(avatar_definition,[new_symbols(naming,[spl26_216])])).
fof(f5658,plain,(
  ~midp(sK21,sK20,sK21) | (~spl26_3 | ~spl26_216)),
  inference(subsumption_resolution,[],[f4478,f487])).
fof(f487,plain,(
  ~cong(sK22,sK20,sK22,sK21) | ~spl26_3),
  inference(avatar_component_clause,[],[f486])).
fof(f4478,plain,(
  cong(sK22,sK20,sK22,sK21) | ~midp(sK21,sK20,sK21) | ~spl26_216),
  inference(resolution,[],[f3913,f414])).
fof(f3913,plain,(
  perp(sK22,sK21,sK20,sK21) | ~spl26_216),
  inference(resolution,[],[f2484,f366])).
fof(f2484,plain,(
  perp(sK20,sK21,sK22,sK21) | ~spl26_216),
  inference(avatar_component_clause,[],[f2483])).
fof(f5090,plain,(
  ~spl26_5 | spl26_304 | ~spl26_12),
  inference(avatar_split_clause,[],[f4058,f612,f5088,f557])).
fof(f557,plain,(
  spl26_5 <=> ~midp(sK22,sK20,sK21)),
  introduced(avatar_definition,[new_symbols(naming,[spl26_5])])).
fof(f612,plain,(
  spl26_12 <=> perp(sK20,sK21,sK21,sK22)),
  introduced(avatar_definition,[new_symbols(naming,[spl26_12])])).
fof(f4058,plain,(
  cong(sK21,sK20,sK21,sK21) | ~midp(sK22,sK20,sK21) | ~spl26_12),
  inference(resolution,[],[f3329,f414])).
fof(f3329,plain,(
  perp(sK21,sK22,sK20,sK21) | ~spl26_12),
  inference(resolution,[],[f613,f366])).
fof(f613,plain,(
  perp(sK20,sK21,sK21,sK22) | ~spl26_12),
  inference(avatar_component_clause,[],[f612])).
fof(f4538,plain,(
  ~spl26_293 | spl26_294 | ~spl26_0),
  inference(avatar_split_clause,[],[f4518,f483,f4536,f4530])).
fof(f4530,plain,(
  spl26_293 <=> ~para(sK22,sK21,sK20,sK21)),
  introduced(avatar_definition,[new_symbols(naming,[spl26_293])])).
fof(f4518,plain,(
  para(sK20,sK22,sK22,sK21) | ~para(sK22,sK21,sK20,sK21) | ~spl26_0),
  inference(resolution,[],[f3958,f1037])).
fof(f1037,plain,(
  ( ! [X21,X19,X20] : (~cong(X21,X19,X21,X20) | para(X21,X19,X19,X20) | ~para(X19,X20,X21,X20)) )),
  inference(resolution,[],[f432,f405])).
fof(f3958,plain,(
  cong(sK20,sK22,sK20,sK21) | ~spl26_0),
  inference(resolution,[],[f3256,f382])).
fof(f3256,plain,(
  cong(sK20,sK21,sK20,sK22) | ~spl26_0),
  inference(resolution,[],[f3239,f381])).
fof(f3239,plain,(
  cong(sK20,sK21,sK22,sK20) | ~spl26_0),
  inference(resolution,[],[f484,f382])).
fof(f4482,plain,(
  spl26_222 | ~spl26_216),
  inference(avatar_split_clause,[],[f4480,f2483,f2505])).
fof(f4480,plain,(
  perp(sK22,sK21,sK21,sK20) | ~spl26_216),
  inference(resolution,[],[f3913,f365])).
fof(f3419,plain,(
  spl26_176 | ~spl26_14),
  inference(avatar_split_clause,[],[f3417,f615,f1884])).
fof(f3417,plain,(
  coll(sK20,sK22,sK20) | ~spl26_14),
  inference(resolution,[],[f3356,f360])).
fof(f3336,plain,(
  spl26_216 | ~spl26_12),
  inference(avatar_split_clause,[],[f3330,f612,f2483])).
fof(f3330,plain,(
  perp(sK20,sK21,sK22,sK21) | ~spl26_12),
  inference(resolution,[],[f613,f365])).
fof(f3317,plain,(
  spl26_14 | ~spl26_242),
  inference(avatar_split_clause,[],[f3313,f2924,f615])).
fof(f2924,plain,(
  spl26_242 <=> coll(sK20,sK21,sK22)),
  introduced(avatar_definition,[new_symbols(naming,[spl26_242])])).
fof(f3313,plain,(
  coll(sK23,sK20,sK22) | ~spl26_242),
  inference(resolution,[],[f3297,f360])).
fof(f3297,plain,(
  coll(sK20,sK23,sK22) | ~spl26_242),
  inference(resolution,[],[f3276,f537])).
fof(f537,plain,(
  ( ! [X4] : (~coll(sK22,sK24,X4) | coll(X4,sK23,sK22)) )),
  inference(resolution,[],[f361,f495])).
fof(f3276,plain,(
  coll(sK22,sK24,sK20) | ~spl26_242),
  inference(resolution,[],[f2925,f534])).
fof(f534,plain,(
  ( ! [X1] : (~coll(sK20,sK21,X1) | coll(X1,sK24,sK20)) )),
  inference(resolution,[],[f361,f503])).
fof(f503,plain,(
  coll(sK20,sK21,sK24)),
  inference(resolution,[],[f496,f359])).
fof(f2925,plain,(
  coll(sK20,sK21,sK22) | ~spl26_242),
  inference(avatar_component_clause,[],[f2924])).
fof(f3262,plain,(
  spl26_242 | ~spl26_172),
  inference(avatar_split_clause,[],[f3259,f1868,f2924])).
fof(f1868,plain,(
  spl26_172 <=> coll(sK20,sK22,sK21)),
  introduced(avatar_definition,[new_symbols(naming,[spl26_172])])).
fof(f3259,plain,(
  coll(sK20,sK21,sK22) | ~spl26_172),
  inference(resolution,[],[f1869,f359])).
fof(f1869,plain,(
  coll(sK20,sK22,sK21) | ~spl26_172),
  inference(avatar_component_clause,[],[f1868])).
fof(f3244,plain,(
  spl26_172 | ~spl26_6),
  inference(avatar_split_clause,[],[f3242,f563,f1868])).
fof(f563,plain,(
  spl26_6 <=> coll(sK22,sK20,sK21)),
  introduced(avatar_definition,[new_symbols(naming,[spl26_6])])).
fof(f3242,plain,(
  coll(sK20,sK22,sK21) | ~spl26_6),
  inference(resolution,[],[f564,f360])).
fof(f564,plain,(
  coll(sK22,sK20,sK21) | ~spl26_6),
  inference(avatar_component_clause,[],[f563])).
fof(f3229,plain,(
  spl26_4 | ~spl26_7 | ~spl26_2),
  inference(avatar_split_clause,[],[f3039,f489,f566,f560])).
fof(f566,plain,(
  spl26_7 <=> ~coll(sK22,sK20,sK21)),
  introduced(avatar_definition,[new_symbols(naming,[spl26_7])])).
fof(f3039,plain,(
  ~coll(sK22,sK20,sK21) | midp(sK22,sK20,sK21) | ~spl26_2),
  inference(resolution,[],[f490,f426])).
fof(f3228,plain,(
  spl26_252 | ~spl26_7 | ~spl26_2),
  inference(avatar_split_clause,[],[f3094,f489,f566,f3034])).
fof(f3094,plain,(
  ~coll(sK22,sK20,sK21) | perp(sK20,sK21,sK21,sK21) | ~spl26_2),
  inference(resolution,[],[f3073,f412])).
fof(f412,plain,(
  ( ! [X2,X0,X3,X1] : (~circle(X3,X0,X1,X2) | ~coll(X3,X0,X2) | perp(X0,X1,X1,X2)) )),
  inference(cnf_transformation,[],[f245])).
fof(f245,plain,(
  ! [X0,X1,X2,X3] : (perp(X0,X1,X1,X2) | ~coll(X3,X0,X2) | ~circle(X3,X0,X1,X2))),
  inference(flattening,[],[f244])).
fof(f244,plain,(
  ! [X0,X1,X2,X3] : (perp(X0,X1,X1,X2) | (~coll(X3,X0,X2) | ~circle(X3,X0,X1,X2)))),
  inference(ennf_transformation,[],[f133])).
fof(f133,plain,(
  ! [X0,X1,X2,X3] : ((coll(X3,X0,X2) & circle(X3,X0,X1,X2)) => perp(X0,X1,X1,X2))),
  inference(rectify,[],[f54])).
fof(f54,axiom,(
  ! [X0,X1,X2,X7] : ((coll(X7,X0,X2) & circle(X7,X0,X1,X2)) => perp(X0,X1,X1,X2))),
  file('./6_GDD_FULL/41-60/56.p',ruleD53)).
fof(f3073,plain,(
  circle(sK22,sK20,sK21,sK21) | ~spl26_2),
  inference(resolution,[],[f3037,f490])).
fof(f3037,plain,(
  ( ! [X0] : (~cong(sK22,sK20,sK22,X0) | circle(sK22,sK20,X0,sK21)) ) | ~spl26_2),
  inference(resolution,[],[f490,f370])).
fof(f3227,plain,(
  spl26_6 | ~spl26_226),
  inference(avatar_split_clause,[],[f3220,f2730,f563])).
fof(f3220,plain,(
  coll(sK22,sK20,sK21) | ~spl26_226),
  inference(resolution,[],[f3216,f842])).
fof(f842,plain,(
  ( ! [X0] : (~coll(sK21,sK21,X0) | coll(X0,sK20,sK21)) )),
  inference(resolution,[],[f838,f361])).
fof(f3216,plain,(
  coll(sK21,sK21,sK22) | ~spl26_226),
  inference(resolution,[],[f3209,f359])).
fof(f3195,plain,(
  ~spl26_184),
  inference(avatar_contradiction_clause,[],[f3192])).
fof(f3192,plain,(
  $false | ~spl26_184),
  inference(resolution,[],[f2073,f1449])).
fof(f1449,plain,(
  para(sK20,sK25,sK20,sK25)),
  inference(resolution,[],[f731,f516])).
fof(f2073,plain,(
  ( ! [X0,X1] : (~para(X0,X1,X0,X1)) ) | ~spl26_184),
  inference(avatar_component_clause,[],[f2072])).
fof(f2072,plain,(
  spl26_184 <=> ! [X1,X0] : ~para(X0,X1,X0,X1)),
  introduced(avatar_definition,[new_symbols(naming,[spl26_184])])).
fof(f3171,plain,(
  ~spl26_182 | spl26_227),
  inference(avatar_contradiction_clause,[],[f3155])).
fof(f3155,plain,(
  $false | (~spl26_182 | ~spl26_227)),
  inference(resolution,[],[f3124,f2734])).
fof(f2734,plain,(
  ~coll(sK22,sK21,sK21) | ~spl26_227),
  inference(avatar_component_clause,[],[f2733])).
fof(f2733,plain,(
  spl26_227 <=> ~coll(sK22,sK21,sK21)),
  introduced(avatar_definition,[new_symbols(naming,[spl26_227])])).
fof(f3124,plain,(
  ( ! [X0,X1] : (coll(X0,X1,X1)) ) | ~spl26_182),
  inference(resolution,[],[f2070,f425])).
fof(f2942,plain,(
  ~spl26_245 | spl26_246 | ~spl26_0),
  inference(avatar_split_clause,[],[f2911,f483,f2940,f2934])).
fof(f2934,plain,(
  spl26_245 <=> ~para(sK21,sK22,sK20,sK22)),
  introduced(avatar_definition,[new_symbols(naming,[spl26_245])])).
fof(f2911,plain,(
  para(sK20,sK21,sK21,sK22) | ~para(sK21,sK22,sK20,sK22) | ~spl26_0),
  inference(resolution,[],[f2760,f1037])).
fof(f2760,plain,(
  cong(sK20,sK21,sK20,sK22) | ~spl26_0),
  inference(resolution,[],[f2751,f381])).
fof(f2751,plain,(
  cong(sK20,sK21,sK22,sK20) | ~spl26_0),
  inference(resolution,[],[f484,f382])).
fof(f2929,plain,(
  spl26_240 | ~spl26_243 | ~spl26_0),
  inference(avatar_split_clause,[],[f2905,f483,f2927,f2921])).
fof(f2927,plain,(
  spl26_243 <=> ~coll(sK20,sK21,sK22)),
  introduced(avatar_definition,[new_symbols(naming,[spl26_243])])).
fof(f2905,plain,(
  ~coll(sK20,sK21,sK22) | midp(sK20,sK21,sK22) | ~spl26_0),
  inference(resolution,[],[f2760,f426])).
fof(f2253,plain,(
  spl26_186 | spl26_188),
  inference(avatar_split_clause,[],[f2230,f2251,f2248])).
fof(f2230,plain,(
  ( ! [X47,X45,X48,X46,X44,X49] : (~para(X44,X45,X46,X47) | ~perp(X48,X49,X48,X49) | perp(X44,X45,X46,X47)) )),
  inference(resolution,[],[f824,f433])).
fof(f2074,plain,(
  spl26_182 | spl26_184),
  inference(avatar_split_clause,[],[f2045,f2072,f2069])).
fof(f2045,plain,(
  ( ! [X2,X0,X3,X1] : (~para(X0,X1,X0,X1) | para(X2,X3,X2,X3)) )),
  inference(resolution,[],[f811,f397])).
fof(f811,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (eqangle(X0,X1,X2,X3,X0,X1,X4,X5) | ~para(X2,X3,X4,X5)) )),
  inference(resolution,[],[f377,f398])).
fof(f1879,plain,(
  spl26_170 | ~spl26_173 | ~spl26_175),
  inference(avatar_split_clause,[],[f1836,f1877,f1871,f1865])).
fof(f1871,plain,(
  spl26_173 <=> ~coll(sK20,sK22,sK21)),
  introduced(avatar_definition,[new_symbols(naming,[spl26_173])])).
fof(f1877,plain,(
  spl26_175 <=> ~perp(sK23,sK20,sK20,sK20)),
  introduced(avatar_definition,[new_symbols(naming,[spl26_175])])).
fof(f1836,plain,(
  ~perp(sK23,sK20,sK20,sK20) | ~coll(sK20,sK22,sK21) | cyclic(sK20,sK21,sK20,sK22)),
  inference(resolution,[],[f999,f401])).
fof(f999,plain,(
  ( ! [X0] : (eqangle(sK20,X0,sK20,sK21,sK22,sK20,sK22,sK21) | ~perp(sK23,sK20,sK20,X0)) )),
  inference(resolution,[],[f407,f473])).
fof(f407,plain,(
  ( ! [X4,X2,X0,X3,X1] : (~circle(X3,X0,X1,X2) | ~perp(X3,X0,X0,X4) | eqangle(X0,X4,X0,X1,X2,X0,X2,X1)) )),
  inference(cnf_transformation,[],[f235])).
fof(f235,plain,(
  ! [X0,X1,X2,X3,X4] : (eqangle(X0,X4,X0,X1,X2,X0,X2,X1) | ~perp(X3,X0,X0,X4) | ~circle(X3,X0,X1,X2))),
  inference(flattening,[],[f234])).
fof(f234,plain,(
  ! [X0,X1,X2,X3,X4] : (eqangle(X0,X4,X0,X1,X2,X0,X2,X1) | (~perp(X3,X0,X0,X4) | ~circle(X3,X0,X1,X2)))),
  inference(ennf_transformation,[],[f128])).
fof(f128,plain,(
  ! [X0,X1,X2,X3,X4] : ((perp(X3,X0,X0,X4) & circle(X3,X0,X1,X2)) => eqangle(X0,X4,X0,X1,X2,X0,X2,X1))),
  inference(rectify,[],[f49])).
fof(f49,axiom,(
  ! [X0,X1,X2,X7,X15] : ((perp(X7,X0,X0,X15) & circle(X7,X0,X1,X2)) => eqangle(X0,X15,X0,X1,X2,X0,X2,X1))),
  file('./6_GDD_FULL/41-60/56.p',ruleD48)).
fof(f1060,plain,(
  spl26_80 | spl26_82),
  inference(avatar_split_clause,[],[f1053,f1058,f1055])).
fof(f1053,plain,(
  ( ! [X6,X10,X8,X7,X11,X9] : (~cong(X6,X7,X7,X8) | cong(X9,X10,X10,X11) | ~midp(X6,X7,X8) | ~midp(X9,X10,X11)) )),
  inference(resolution,[],[f434,f429])).
fof(f429,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (eqratio(X4,X0,X0,X1,X5,X2,X2,X3) | ~midp(X5,X2,X3) | ~midp(X4,X0,X1)) )),
  inference(cnf_transformation,[],[f273])).
fof(f273,plain,(
  ! [X0,X1,X2,X3,X4,X5] : (eqratio(X4,X0,X0,X1,X5,X2,X2,X3) | ~midp(X5,X2,X3) | ~midp(X4,X0,X1))),
  inference(flattening,[],[f272])).
fof(f272,plain,(
  ! [X0,X1,X2,X3,X4,X5] : (eqratio(X4,X0,X0,X1,X5,X2,X2,X3) | (~midp(X5,X2,X3) | ~midp(X4,X0,X1)))),
  inference(ennf_transformation,[],[f145])).
fof(f145,plain,(
  ! [X0,X1,X2,X3,X4,X5] : ((midp(X5,X2,X3) & midp(X4,X0,X1)) => eqratio(X4,X0,X0,X1,X5,X2,X2,X3))),
  inference(rectify,[],[f71])).
fof(f71,axiom,(
  ! [X0,X1,X2,X3,X6,X16] : ((midp(X16,X2,X3) & midp(X6,X0,X1)) => eqratio(X6,X0,X0,X1,X16,X2,X2,X3))),
  file('./6_GDD_FULL/41-60/56.p',ruleD70)).
fof(f620,plain,(
  spl26_12 | ~spl26_15),
  inference(avatar_split_clause,[],[f607,f618,f612])).
fof(f618,plain,(
  spl26_15 <=> ~coll(sK23,sK20,sK22)),
  introduced(avatar_definition,[new_symbols(naming,[spl26_15])])).
fof(f607,plain,(
  ~coll(sK23,sK20,sK22) | perp(sK20,sK21,sK21,sK22)),
  inference(resolution,[],[f412,f473])).
fof(f491,plain,(
  spl26_0 | spl26_2),
  inference(avatar_split_clause,[],[f472,f489,f483])).
fof(f472,plain,(
  cong(sK22,sK20,sK22,sK21) | cong(sK22,sK20,sK20,sK21)),
  inference(cnf_transformation,[],[f358])).
% SZS output end Proof for 56
% ------------------------------
% Version: Vampire 4.4.0 (commit 7916d27 on 2019-08-23 08:50:16 +0100)
% Termination reason: Refutation

% Memory used [KB]: 21620
% Time elapsed: 0.907 s
% ------------------------------
% ------------------------------
