% Refutation found. Thanks to Tanya!
% SZS status Theorem for 56
% SZS output start Proof for 56
7. ! [X0,X1,X2,X3] : (perp(X0,X1,X2,X3) => perp(X0,X1,X3,X2)) [input]
8. ! [X0,X1,X2,X3] : (perp(X0,X1,X2,X3) => perp(X2,X3,X0,X1)) [input]
24. ! [X0,X1,X2,X3] : (cong(X0,X1,X2,X3) => cong(X2,X3,X0,X1)) [input]
55. ! [X0,X1,X6,X7] : ((perp(X7,X6,X0,X1) & midp(X6,X0,X1)) => cong(X7,X0,X7,X1)) [input]
67. ! [X0,X1,X2] : ((coll(X0,X1,X2) & cong(X0,X1,X0,X2)) => midp(X0,X1,X2)) [input]
94. ! [X0,X1,X6,X7,X3,X4] : (para(X0,X7,X4,X6) & perp(X0,X7,X0,X4) & coll(X3,X0,X1) & coll(X3,X6,X7) & circle(X7,X0,X1,X6) & (cong(X6,X0,X6,X1) | cong(X6,X0,X0,X1))) => cong(X6,X4,X6,X3) [input]
95. ~! [X3,X6,X4] : (! [X0,X1,X6,X7,X3,X4] : (para(X0,X7,X4,X6) & perp(X0,X7,X0,X4) & coll(X3,X0,X1) & coll(X3,X6,X7) & circle(X7,X0,X1,X6) & (cong(X6,X0,X6,X1) | cong(X6,X0,X0,X1))) => cong(X6,X4,X6,X3)) [negated conjecture 94]
132. ! [X0,X1,X2,X3] : ((perp(X3,X2,X0,X1) & midp(X2,X0,X1)) => cong(X3,X0,X3,X1)) [rectify 55]
165. ~! [X0,X1,X2] : (! [X3,X4,X5,X6,X7,X8] : (para(X3,X6,X8,X5) & perp(X3,X6,X3,X8) & coll(X7,X3,X4) & coll(X7,X5,X6) & circle(X6,X3,X4,X5) & (cong(X5,X3,X5,X4) | cong(X5,X3,X3,X4))) => cong(X1,X2,X1,X0)) [rectify 95]
174. ! [X0,X1,X2,X3] : (perp(X0,X1,X3,X2) | ~perp(X0,X1,X2,X3)) [ennf transformation 7]
175. ! [X0,X1,X2,X3] : (perp(X2,X3,X0,X1) | ~perp(X0,X1,X2,X3)) [ennf transformation 8]
197. ! [X0,X1,X2,X3] : (cong(X2,X3,X0,X1) | ~cong(X0,X1,X2,X3)) [ennf transformation 24]
244. ! [X0,X1,X2,X3] : (cong(X3,X0,X3,X1) | (~perp(X3,X2,X0,X1) | ~midp(X2,X0,X1))) [ennf transformation 132]
245. ! [X0,X1,X2,X3] : (cong(X3,X0,X3,X1) | ~perp(X3,X2,X0,X1) | ~midp(X2,X0,X1)) [flattening 244]
264. ! [X0,X1,X2] : (midp(X0,X1,X2) | (~coll(X0,X1,X2) | ~cong(X0,X1,X0,X2))) [ennf transformation 67]
265. ! [X0,X1,X2] : (midp(X0,X1,X2) | ~coll(X0,X1,X2) | ~cong(X0,X1,X0,X2)) [flattening 264]
315. ? [X0,X1,X2] : (~cong(X1,X2,X1,X0) & ! [X3,X4,X5,X6,X7,X8] : (para(X3,X6,X8,X5) & perp(X3,X6,X3,X8) & coll(X7,X3,X4) & coll(X7,X5,X6) & circle(X6,X3,X4,X5) & (cong(X5,X3,X5,X4) | cong(X5,X3,X3,X4)))) [ennf transformation 165]
352. ? [X0,X1,X2] : (~cong(X1,X2,X1,X0) & ! [X3,X4,X5,X6,X7,X8] : (para(X3,X6,X8,X5) & perp(X3,X6,X3,X8) & coll(X7,X3,X4) & coll(X7,X5,X6) & circle(X6,X3,X4,X5) & (cong(X5,X3,X5,X4) | cong(X5,X3,X3,X4)))) => (~cong(sK21,sK22,sK21,sK20) & ! [X3,X4,X5,X6,X7,X8] : (para(X3,X6,X8,X5) & perp(X3,X6,X3,X8) & coll(X7,X3,X4) & coll(X7,X5,X6) & circle(X6,X3,X4,X5) & (cong(X5,X3,X5,X4) | cong(X5,X3,X3,X4)))) [choice axiom]
353. ~cong(sK21,sK22,sK21,sK20) & ! [X3,X4,X5,X6,X7,X8] : (para(X3,X6,X8,X5) & perp(X3,X6,X3,X8) & coll(X7,X3,X4) & coll(X7,X5,X6) & circle(X6,X3,X4,X5) & (cong(X5,X3,X5,X4) | cong(X5,X3,X3,X4))) [skolemisation 315,352]
360. ~perp(X0,X1,X2,X3) | perp(X0,X1,X3,X2) [cnf transformation 174]
361. ~perp(X0,X1,X2,X3) | perp(X2,X3,X0,X1) [cnf transformation 175]
377. ~cong(X0,X1,X2,X3) | cong(X2,X3,X0,X1) [cnf transformation 197]
409. ~perp(X3,X2,X0,X1) | cong(X3,X0,X3,X1) | ~midp(X2,X0,X1) [cnf transformation 245]
421. ~cong(X0,X1,X0,X2) | ~coll(X0,X1,X2) | midp(X0,X1,X2) [cnf transformation 265]
467. cong(X5,X3,X5,X4) | cong(X5,X3,X3,X4) [cnf transformation 353]
470. coll(X7,X3,X4) [cnf transformation 353]
471. perp(X3,X6,X3,X8) [cnf transformation 353]
473. ~cong(sK21,sK22,sK21,sK20) [cnf transformation 353]
479. cong(X0,X0,X0,X1) [factoring 467]
482. perp(X0,X1,X2,X0) [resolution 360,471]
485. perp(X3,X4,X4,X5) [resolution 361,482]
487. perp(X3,X4,X5,X4) [resolution 485,360]
499. cong(X9,X10,X9,X9) [resolution 377,479]
516. ~coll(X8,X9,X8) | midp(X8,X9,X8) [resolution 421,499]
521. midp(X8,X9,X8) [subsumption resolution 516,470]
575. cong(X6,X7,X6,X8) | ~midp(X8,X7,X8) [resolution 409,487]
577. cong(X6,X7,X6,X8) [subsumption resolution 575,521]
592. $false [resolution 577,473]
% SZS output end Proof for 56
% ------------------------------
% Version: Vampire 4.4.0 (commit 7916d27 on 2019-08-23 08:50:16 +0100)
% Termination reason: Refutation

% Memory used [KB]: 5117
% Time elapsed: 0.028 s
% ------------------------------
% ------------------------------
