% Refutation found. Thanks to Tanya!
% SZS status Theorem for E044-6
% SZS output start Proof for E044-6
23. ! [X0,X1,X2,X3] : (cong(X0,X1,X2,X3) => cong(X0,X1,X3,X2)) [input]
24. ! [X0,X1,X2,X3] : (cong(X0,X1,X2,X3) => cong(X2,X3,X0,X1)) [input]
25. ! [X0,X1,X2,X3,X4,X5] : ((cong(X2,X3,X4,X5) & cong(X0,X1,X2,X3)) => cong(X0,X1,X4,X5)) [input]
68. ! [X0,X1,X2] : (midp(X0,X1,X2) => cong(X0,X1,X0,X2)) [input]
94. ! [X0,X1,X2,X3,X4,X5,X12,X13,X18,X19,X20,X21,X22] : (midp(X13,X3,X5) & coll(X12,X1,X2) & perp(X12,X0,X1,X2) & coll(X5,X1,X4) & circle(X0,X1,X5,X22) & coll(X4,X2,X3) & perp(X4,X1,X2,X3) & circle(X0,X1,X3,X21) & circle(X0,X1,X2,X20) & circle(X0,X1,X18,X19)) => cong(X0,X12,X13,X4) [input]
95. ~! [X4,X13,X12,X0] : (! [X0,X1,X2,X3,X4,X5,X12,X13,X18,X19,X20,X21,X22] : (midp(X13,X3,X5) & coll(X12,X1,X2) & perp(X12,X0,X1,X2) & coll(X5,X1,X4) & circle(X0,X1,X5,X22) & coll(X4,X2,X3) & perp(X4,X1,X2,X3) & circle(X0,X1,X3,X21) & circle(X0,X1,X2,X20) & circle(X0,X1,X18,X19)) => cong(X0,X12,X13,X4)) [negated conjecture 94]
165. ~! [X0,X1,X2,X3] : (! [X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16] : (midp(X11,X7,X9) & coll(X10,X5,X6) & perp(X10,X4,X5,X6) & coll(X9,X5,X8) & circle(X4,X5,X9,X16) & coll(X8,X6,X7) & perp(X8,X5,X6,X7) & circle(X4,X5,X7,X15) & circle(X4,X5,X6,X14) & circle(X4,X5,X12,X13)) => cong(X3,X2,X1,X0)) [rectify 95]
196. ! [X0,X1,X2,X3] : (cong(X0,X1,X3,X2) | ~cong(X0,X1,X2,X3)) [ennf transformation 23]
197. ! [X0,X1,X2,X3] : (cong(X2,X3,X0,X1) | ~cong(X0,X1,X2,X3)) [ennf transformation 24]
198. ! [X0,X1,X2,X3,X4,X5] : (cong(X0,X1,X4,X5) | (~cong(X2,X3,X4,X5) | ~cong(X0,X1,X2,X3))) [ennf transformation 25]
199. ! [X0,X1,X2,X3,X4,X5] : (cong(X0,X1,X4,X5) | ~cong(X2,X3,X4,X5) | ~cong(X0,X1,X2,X3)) [flattening 198]
266. ! [X0,X1,X2] : (cong(X0,X1,X0,X2) | ~midp(X0,X1,X2)) [ennf transformation 68]
315. ? [X0,X1,X2,X3] : (~cong(X3,X2,X1,X0) & ! [X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16] : (midp(X11,X7,X9) & coll(X10,X5,X6) & perp(X10,X4,X5,X6) & coll(X9,X5,X8) & circle(X4,X5,X9,X16) & coll(X8,X6,X7) & perp(X8,X5,X6,X7) & circle(X4,X5,X7,X15) & circle(X4,X5,X6,X14) & circle(X4,X5,X12,X13))) [ennf transformation 165]
352. ? [X0,X1,X2,X3] : (~cong(X3,X2,X1,X0) & ! [X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16] : (midp(X11,X7,X9) & coll(X10,X5,X6) & perp(X10,X4,X5,X6) & coll(X9,X5,X8) & circle(X4,X5,X9,X16) & coll(X8,X6,X7) & perp(X8,X5,X6,X7) & circle(X4,X5,X7,X15) & circle(X4,X5,X6,X14) & circle(X4,X5,X12,X13))) => (~cong(sK23,sK22,sK21,sK20) & ! [X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16] : (midp(X11,X7,X9) & coll(X10,X5,X6) & perp(X10,X4,X5,X6) & coll(X9,X5,X8) & circle(X4,X5,X9,X16) & coll(X8,X6,X7) & perp(X8,X5,X6,X7) & circle(X4,X5,X7,X15) & circle(X4,X5,X6,X14) & circle(X4,X5,X12,X13))) [choice axiom]
353. ~cong(sK23,sK22,sK21,sK20) & ! [X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16] : (midp(X11,X7,X9) & coll(X10,X5,X6) & perp(X10,X4,X5,X6) & coll(X9,X5,X8) & circle(X4,X5,X9,X16) & coll(X8,X6,X7) & perp(X8,X5,X6,X7) & circle(X4,X5,X7,X15) & circle(X4,X5,X6,X14) & circle(X4,X5,X12,X13)) [skolemisation 315,352]
376. ~cong(X0,X1,X2,X3) | cong(X0,X1,X3,X2) [cnf transformation 196]
377. ~cong(X0,X1,X2,X3) | cong(X2,X3,X0,X1) [cnf transformation 197]
378. ~cong(X2,X3,X4,X5) | cong(X0,X1,X4,X5) | ~cong(X0,X1,X2,X3) [cnf transformation 199]
422. cong(X0,X1,X0,X2) | ~midp(X0,X1,X2) [cnf transformation 266]
476. midp(X11,X7,X9) [cnf transformation 353]
477. ~cong(sK23,sK22,sK21,sK20) [cnf transformation 353]
484. cong(X0,X1,X2,X0) | ~midp(X0,X1,X2) [resolution 376,422]
485. cong(X0,X1,X2,X0) [subsumption resolution 484,476]
486. cong(X0,X1,X0,X2) | ~midp(X0,X2,X1) [resolution 377,422]
487. cong(X0,X1,X0,X2) [subsumption resolution 486,476]
528. ~cong(X0,X1,X2,X4) | cong(X0,X1,X2,X3) [resolution 378,487]
535. cong(X7,X8,X9,X10) [resolution 528,485]
543. $false [resolution 535,477]
% SZS output end Proof for E044-6
% ------------------------------
% Version: Vampire 4.4.0 (commit 7916d27 on 2019-08-23 08:50:16 +0100)
% Termination reason: Refutation

% Memory used [KB]: 5117
% Time elapsed: 0.027 s
% ------------------------------
% ------------------------------
