% Refutation found. Thanks to Tanya!
% SZS status Theorem for E044-6
% SZS output start Proof for E044-6
fof(f20765,plain,(
  $false),
  inference(avatar_sat_refutation,[],[f633,f1999,f2197,f2220,f2703,f3253,f3505,f3738,f7631,f8577,f8596,f8599,f9091,f12212,f13124,f13864,f13865,f13902,f15358,f17668,f17771,f19149,f20764])).
fof(f20764,plain,(
  ~spl33_160 | ~spl33_176 | ~spl33_232 | ~spl33_492 | spl33_579 | ~spl33_662),
  inference(avatar_contradiction_clause,[],[f20763])).
fof(f20763,plain,(
  $false | (~spl33_160 | ~spl33_176 | ~spl33_232 | ~spl33_492 | ~spl33_579 | ~spl33_662)),
  inference(resolution,[],[f20728,f12211])).
fof(f12211,plain,(
  ~midp(sK27,sK20,sK24) | ~spl33_579),
  inference(avatar_component_clause,[],[f12210])).
fof(f12210,plain,(
  spl33_579 <=> ~midp(sK27,sK20,sK24)),
  introduced(avatar_definition,[new_symbols(naming,[spl33_579])])).
fof(f20728,plain,(
  ( ! [X6,X4,X5] : (midp(X4,X5,X6)) ) | (~spl33_160 | ~spl33_176 | ~spl33_232 | ~spl33_492 | ~spl33_662)),
  inference(subsumption_resolution,[],[f20727,f17900])).
fof(f17900,plain,(
  ( ! [X4,X2,X5,X3] : (para(X2,X3,X4,X5)) ) | (~spl33_160 | ~spl33_176 | ~spl33_232 | ~spl33_492 | ~spl33_662)),
  inference(subsumption_resolution,[],[f17877,f17831])).
fof(f17831,plain,(
  ( ! [X0,X1] : (midp(X0,sK21,X1)) ) | (~spl33_160 | ~spl33_176 | ~spl33_232 | ~spl33_492 | ~spl33_662)),
  inference(resolution,[],[f17820,f369])).
fof(f369,plain,(
  ( ! [X2,X0,X1] : (~midp(X2,X1,X0) | midp(X2,X0,X1)) )),
  inference(cnf_transformation,[],[f182])).
fof(f182,plain,(
  ! [X0,X1,X2] : (midp(X2,X0,X1) | ~midp(X2,X1,X0))),
  inference(ennf_transformation,[],[f97])).
fof(f97,plain,(
  ! [X0,X1,X2] : (midp(X2,X1,X0) => midp(X2,X0,X1))),
  inference(rectify,[],[f11])).
fof(f11,axiom,(
  ! [X0,X1,X6] : (midp(X6,X1,X0) => midp(X6,X0,X1))),
  file('./7_Book/00EE/04/E044-6.p',ruleD11)).
fof(f17820,plain,(
  ( ! [X8,X7] : (midp(X7,X8,sK21)) ) | (~spl33_160 | ~spl33_176 | ~spl33_232 | ~spl33_492 | ~spl33_662)),
  inference(subsumption_resolution,[],[f17819,f16966])).
fof(f16966,plain,(
  ( ! [X2] : (midp(sK21,X2,sK32)) ) | (~spl33_160 | ~spl33_176 | ~spl33_232 | ~spl33_492)),
  inference(subsumption_resolution,[],[f16965,f8610])).
fof(f8610,plain,(
  ( ! [X10] : (midp(sK32,X10,sK25)) ) | (~spl33_176 | ~spl33_232 | ~spl33_492)),
  inference(resolution,[],[f8595,f7736])).
fof(f7736,plain,(
  ( ! [X28,X29,X27] : (~midp(X27,X28,X29) | midp(X28,X27,X29)) ) | (~spl33_176 | ~spl33_232)),
  inference(subsumption_resolution,[],[f7686,f4064])).
fof(f4064,plain,(
  ( ! [X2,X0,X1] : (coll(X0,X1,X2)) ) | ~spl33_176),
  inference(subsumption_resolution,[],[f4061,f4048])).
fof(f4048,plain,(
  ( ! [X6,X5] : (coll(X5,X5,X6)) ) | ~spl33_176),
  inference(resolution,[],[f4032,f359])).
fof(f359,plain,(
  ( ! [X2,X0,X1] : (~coll(X0,X1,X2) | coll(X0,X2,X1)) )),
  inference(cnf_transformation,[],[f168])).
fof(f168,plain,(
  ! [X0,X1,X2] : (coll(X0,X2,X1) | ~coll(X0,X1,X2))),
  inference(ennf_transformation,[],[f1])).
fof(f1,axiom,(
  ! [X0,X1,X2] : (coll(X0,X1,X2) => coll(X0,X2,X1))),
  file('./7_Book/00EE/04/E044-6.p',ruleD1)).
fof(f4032,plain,(
  ( ! [X4,X3] : (coll(X3,X4,X3)) ) | ~spl33_176),
  inference(resolution,[],[f4013,f360])).
fof(f360,plain,(
  ( ! [X2,X0,X1] : (~coll(X0,X1,X2) | coll(X1,X0,X2)) )),
  inference(cnf_transformation,[],[f169])).
fof(f169,plain,(
  ! [X0,X1,X2] : (coll(X1,X0,X2) | ~coll(X0,X1,X2))),
  inference(ennf_transformation,[],[f2])).
fof(f2,axiom,(
  ! [X0,X1,X2] : (coll(X0,X1,X2) => coll(X1,X0,X2))),
  file('./7_Book/00EE/04/E044-6.p',ruleD2)).
fof(f4013,plain,(
  ( ! [X4,X3] : (coll(X3,X4,X4)) ) | ~spl33_176),
  inference(resolution,[],[f1995,f425])).
fof(f425,plain,(
  ( ! [X2,X0,X1] : (~para(X0,X1,X0,X2) | coll(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f267])).
fof(f267,plain,(
  ! [X0,X1,X2] : (coll(X0,X1,X2) | ~para(X0,X1,X0,X2))),
  inference(ennf_transformation,[],[f67])).
fof(f67,axiom,(
  ! [X0,X1,X2] : (para(X0,X1,X0,X2) => coll(X0,X1,X2))),
  file('./7_Book/00EE/04/E044-6.p',ruleD66)).
fof(f1995,plain,(
  ( ! [X2,X3] : (para(X2,X3,X2,X3)) ) | ~spl33_176),
  inference(avatar_component_clause,[],[f1994])).
fof(f1994,plain,(
  spl33_176 <=> ! [X3,X2] : para(X2,X3,X2,X3)),
  introduced(avatar_definition,[new_symbols(naming,[spl33_176])])).
fof(f4061,plain,(
  ( ! [X2,X0,X1] : (coll(X0,X1,X2) | ~coll(X2,X2,X0)) ) | ~spl33_176),
  inference(resolution,[],[f4048,f361])).
fof(f361,plain,(
  ( ! [X2,X0,X3,X1] : (~coll(X0,X1,X3) | coll(X2,X3,X0) | ~coll(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f171])).
fof(f171,plain,(
  ! [X0,X1,X2,X3] : (coll(X2,X3,X0) | ~coll(X0,X1,X3) | ~coll(X0,X1,X2))),
  inference(flattening,[],[f170])).
fof(f170,plain,(
  ! [X0,X1,X2,X3] : (coll(X2,X3,X0) | (~coll(X0,X1,X3) | ~coll(X0,X1,X2)))),
  inference(ennf_transformation,[],[f3])).
fof(f3,axiom,(
  ! [X0,X1,X2,X3] : ((coll(X0,X1,X3) & coll(X0,X1,X2)) => coll(X2,X3,X0))),
  file('./7_Book/00EE/04/E044-6.p',ruleD3)).
fof(f7686,plain,(
  ( ! [X28,X29,X27] : (~midp(X27,X28,X29) | ~coll(X28,X27,X29) | midp(X28,X27,X29)) ) | ~spl33_232),
  inference(resolution,[],[f2698,f426])).
fof(f426,plain,(
  ( ! [X2,X0,X1] : (~cong(X0,X1,X0,X2) | ~coll(X0,X1,X2) | midp(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f269])).
fof(f269,plain,(
  ! [X0,X1,X2] : (midp(X0,X1,X2) | ~coll(X0,X1,X2) | ~cong(X0,X1,X0,X2))),
  inference(flattening,[],[f268])).
fof(f268,plain,(
  ! [X0,X1,X2] : (midp(X0,X1,X2) | (~coll(X0,X1,X2) | ~cong(X0,X1,X0,X2)))),
  inference(ennf_transformation,[],[f68])).
fof(f68,axiom,(
  ! [X0,X1,X2] : ((coll(X0,X1,X2) & cong(X0,X1,X0,X2)) => midp(X0,X1,X2))),
  file('./7_Book/00EE/04/E044-6.p',ruleD67)).
fof(f2698,plain,(
  ( ! [X14,X15,X13] : (cong(X14,X13,X14,X15) | ~midp(X13,X14,X15)) ) | ~spl33_232),
  inference(avatar_component_clause,[],[f2697])).
fof(f2697,plain,(
  spl33_232 <=> ! [X13,X15,X14] : (~midp(X13,X14,X15) | cong(X14,X13,X14,X15))),
  introduced(avatar_definition,[new_symbols(naming,[spl33_232])])).
fof(f8595,plain,(
  ( ! [X1] : (midp(X1,sK32,sK25)) ) | ~spl33_492),
  inference(avatar_component_clause,[],[f8594])).
fof(f8594,plain,(
  spl33_492 <=> ! [X1] : midp(X1,sK32,sK25)),
  introduced(avatar_definition,[new_symbols(naming,[spl33_492])])).
fof(f16965,plain,(
  ( ! [X2] : (midp(sK21,X2,sK32) | ~midp(sK32,X2,sK25)) ) | (~spl33_160 | ~spl33_176)),
  inference(subsumption_resolution,[],[f16957,f4064])).
fof(f16957,plain,(
  ( ! [X2] : (~coll(sK21,X2,sK32) | midp(sK21,X2,sK32) | ~midp(sK32,X2,sK25)) ) | ~spl33_160),
  inference(resolution,[],[f15363,f404])).
fof(f404,plain,(
  ( ! [X4,X2,X0,X3,X1] : (~para(X3,X4,X1,X2) | ~coll(X4,X0,X2) | midp(X4,X0,X2) | ~midp(X3,X0,X1)) )),
  inference(cnf_transformation,[],[f230])).
fof(f230,plain,(
  ! [X0,X1,X2,X3,X4] : (midp(X4,X0,X2) | ~coll(X4,X0,X2) | ~para(X3,X4,X1,X2) | ~midp(X3,X0,X1))),
  inference(flattening,[],[f229])).
fof(f229,plain,(
  ! [X0,X1,X2,X3,X4] : (midp(X4,X0,X2) | (~coll(X4,X0,X2) | ~para(X3,X4,X1,X2) | ~midp(X3,X0,X1)))),
  inference(ennf_transformation,[],[f125])).
fof(f125,plain,(
  ! [X0,X1,X2,X3,X4] : ((coll(X4,X0,X2) & para(X3,X4,X1,X2) & midp(X3,X0,X1)) => midp(X4,X0,X2))),
  inference(rectify,[],[f46])).
fof(f46,axiom,(
  ! [X0,X1,X2,X4,X5] : ((coll(X5,X0,X2) & para(X4,X5,X1,X2) & midp(X4,X0,X1)) => midp(X5,X0,X2))),
  file('./7_Book/00EE/04/E044-6.p',ruleD45)).
fof(f15363,plain,(
  para(sK32,sK21,sK25,sK32) | ~spl33_160),
  inference(resolution,[],[f1815,f362])).
fof(f362,plain,(
  ( ! [X2,X0,X3,X1] : (~para(X0,X1,X2,X3) | para(X0,X1,X3,X2)) )),
  inference(cnf_transformation,[],[f172])).
fof(f172,plain,(
  ! [X0,X1,X2,X3] : (para(X0,X1,X3,X2) | ~para(X0,X1,X2,X3))),
  inference(ennf_transformation,[],[f4])).
fof(f4,axiom,(
  ! [X0,X1,X2,X3] : (para(X0,X1,X2,X3) => para(X0,X1,X3,X2))),
  file('./7_Book/00EE/04/E044-6.p',ruleD4)).
fof(f1815,plain,(
  para(sK32,sK21,sK32,sK25) | ~spl33_160),
  inference(avatar_component_clause,[],[f1814])).
fof(f1814,plain,(
  spl33_160 <=> para(sK32,sK21,sK32,sK25)),
  introduced(avatar_definition,[new_symbols(naming,[spl33_160])])).
fof(f17819,plain,(
  ( ! [X8,X7] : (midp(X7,X8,sK21) | ~midp(sK21,X8,sK32)) ) | (~spl33_176 | ~spl33_232 | ~spl33_662)),
  inference(subsumption_resolution,[],[f17809,f4064])).
fof(f17809,plain,(
  ( ! [X8,X7] : (~coll(X7,X8,sK21) | midp(X7,X8,sK21) | ~midp(sK21,X8,sK32)) ) | (~spl33_176 | ~spl33_232 | ~spl33_662)),
  inference(resolution,[],[f17772,f404])).
fof(f17772,plain,(
  ( ! [X9] : (para(sK21,X9,sK32,sK21)) ) | (~spl33_176 | ~spl33_232 | ~spl33_662)),
  inference(subsumption_resolution,[],[f13123,f17429])).
fof(f17429,plain,(
  ( ! [X6,X8,X7] : (perp(X6,X8,X8,X7)) ) | (~spl33_176 | ~spl33_232)),
  inference(subsumption_resolution,[],[f17404,f4064])).
fof(f17404,plain,(
  ( ! [X6,X8,X7,X5] : (~coll(X5,X6,X7) | perp(X6,X8,X8,X7)) ) | (~spl33_176 | ~spl33_232)),
  inference(resolution,[],[f12176,f412])).
fof(f412,plain,(
  ( ! [X2,X0,X3,X1] : (~circle(X3,X0,X1,X2) | ~coll(X3,X0,X2) | perp(X0,X1,X1,X2)) )),
  inference(cnf_transformation,[],[f245])).
fof(f245,plain,(
  ! [X0,X1,X2,X3] : (perp(X0,X1,X1,X2) | ~coll(X3,X0,X2) | ~circle(X3,X0,X1,X2))),
  inference(flattening,[],[f244])).
fof(f244,plain,(
  ! [X0,X1,X2,X3] : (perp(X0,X1,X1,X2) | (~coll(X3,X0,X2) | ~circle(X3,X0,X1,X2)))),
  inference(ennf_transformation,[],[f133])).
fof(f133,plain,(
  ! [X0,X1,X2,X3] : ((coll(X3,X0,X2) & circle(X3,X0,X1,X2)) => perp(X0,X1,X1,X2))),
  inference(rectify,[],[f54])).
fof(f54,axiom,(
  ! [X0,X1,X2,X7] : ((coll(X7,X0,X2) & circle(X7,X0,X1,X2)) => perp(X0,X1,X1,X2))),
  file('./7_Book/00EE/04/E044-6.p',ruleD53)).
fof(f12176,plain,(
  ( ! [X26,X24,X23,X25] : (circle(X23,X25,X26,X24)) ) | (~spl33_176 | ~spl33_232)),
  inference(subsumption_resolution,[],[f12175,f9110])).
fof(f9110,plain,(
  ( ! [X14,X12,X13] : (cong(X12,X13,X12,X14)) ) | ~spl33_176),
  inference(subsumption_resolution,[],[f9049,f8543])).
fof(f8543,plain,(
  ( ! [X12,X13] : (midp(X12,X12,X13)) ) | ~spl33_176),
  inference(resolution,[],[f8489,f4133])).
fof(f4133,plain,(
  ( ! [X12,X10,X11] : (~midp(X12,X11,X10) | midp(X10,X11,X12)) ) | ~spl33_176),
  inference(subsumption_resolution,[],[f4126,f4064])).
fof(f4126,plain,(
  ( ! [X12,X10,X11] : (~coll(X10,X11,X12) | midp(X10,X11,X12) | ~midp(X12,X11,X10)) ) | ~spl33_176),
  inference(resolution,[],[f4020,f404])).
fof(f4020,plain,(
  ( ! [X24,X23] : (para(X23,X24,X24,X23)) ) | ~spl33_176),
  inference(resolution,[],[f1995,f362])).
fof(f8489,plain,(
  ( ! [X12,X13] : (midp(X12,X13,X13)) ) | ~spl33_176),
  inference(resolution,[],[f8476,f4024])).
fof(f4024,plain,(
  ( ! [X12,X13,X11] : (~midp(X11,X13,X13) | midp(X11,X12,X12)) ) | ~spl33_176),
  inference(subsumption_resolution,[],[f4016,f1995])).
fof(f4016,plain,(
  ( ! [X12,X13,X11] : (midp(X11,X12,X12) | ~para(X13,X12,X13,X12) | ~midp(X11,X13,X13)) ) | ~spl33_176),
  inference(resolution,[],[f1995,f423])).
fof(f423,plain,(
  ( ! [X4,X2,X0,X3,X1] : (~para(X0,X3,X1,X2) | midp(X4,X2,X3) | ~para(X0,X2,X1,X3) | ~midp(X4,X0,X1)) )),
  inference(cnf_transformation,[],[f264])).
fof(f264,plain,(
  ! [X0,X1,X2,X3,X4] : (midp(X4,X2,X3) | ~para(X0,X3,X1,X2) | ~para(X0,X2,X1,X3) | ~midp(X4,X0,X1))),
  inference(flattening,[],[f263])).
fof(f263,plain,(
  ! [X0,X1,X2,X3,X4] : (midp(X4,X2,X3) | (~para(X0,X3,X1,X2) | ~para(X0,X2,X1,X3) | ~midp(X4,X0,X1)))),
  inference(ennf_transformation,[],[f143])).
fof(f143,plain,(
  ! [X0,X1,X2,X3,X4] : ((para(X0,X3,X1,X2) & para(X0,X2,X1,X3) & midp(X4,X0,X1)) => midp(X4,X2,X3))),
  inference(rectify,[],[f65])).
fof(f65,axiom,(
  ! [X0,X1,X2,X3,X6] : ((para(X0,X3,X1,X2) & para(X0,X2,X1,X3) & midp(X6,X0,X1)) => midp(X6,X2,X3))),
  file('./7_Book/00EE/04/E044-6.p',ruleD64)).
fof(f8476,plain,(
  ( ! [X14] : (midp(X14,X14,X14)) ) | ~spl33_176),
  inference(subsumption_resolution,[],[f8455,f4048])).
fof(f8455,plain,(
  ( ! [X14] : (~coll(X14,X14,X14) | midp(X14,X14,X14)) ) | ~spl33_176),
  inference(resolution,[],[f8375,f426])).
fof(f8375,plain,(
  ( ! [X35] : (cong(X35,X35,X35,X35)) ) | ~spl33_176),
  inference(subsumption_resolution,[],[f8371,f8333])).
fof(f8333,plain,(
  ( ! [X2,X0,X1] : (cyclic(X0,X0,X1,X2)) ) | ~spl33_176),
  inference(resolution,[],[f3364,f1995])).
fof(f3364,plain,(
  ( ! [X68,X66,X67] : (~para(X66,X67,X66,X67) | cyclic(X67,X67,X66,X68)) )),
  inference(subsumption_resolution,[],[f2180,f2181])).
fof(f2181,plain,(
  ( ! [X70,X71,X69] : (~para(X69,X70,X69,X70) | coll(X69,X71,X70) | cyclic(X70,X70,X69,X71)) )),
  inference(resolution,[],[f884,f400])).
fof(f400,plain,(
  ( ! [X2,X0,X3,X1] : (~eqangle(X2,X0,X2,X1,X3,X0,X3,X1) | coll(X2,X3,X0) | cyclic(X0,X1,X2,X3)) )),
  inference(cnf_transformation,[],[f222])).
fof(f222,plain,(
  ! [X0,X1,X2,X3] : (cyclic(X0,X1,X2,X3) | coll(X2,X3,X0) | ~eqangle(X2,X0,X2,X1,X3,X0,X3,X1))),
  inference(flattening,[],[f221])).
fof(f221,plain,(
  ! [X0,X1,X2,X3] : (cyclic(X0,X1,X2,X3) | (coll(X2,X3,X0) | ~eqangle(X2,X0,X2,X1,X3,X0,X3,X1)))),
  inference(ennf_transformation,[],[f121])).
fof(f121,plain,(
  ! [X0,X1,X2,X3] : ((~coll(X2,X3,X0) & eqangle(X2,X0,X2,X1,X3,X0,X3,X1)) => cyclic(X0,X1,X2,X3))),
  inference(rectify,[],[f42])).
fof(f42,axiom,(
  ! [X0,X1,X8,X9] : ((~coll(X8,X9,X0) & eqangle(X8,X0,X8,X1,X9,X0,X9,X1)) => cyclic(X0,X1,X8,X9))),
  file('./7_Book/00EE/04/E044-6.p',ruleD42a)).
fof(f884,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (eqangle(X0,X1,X2,X3,X4,X5,X4,X5) | ~para(X0,X1,X2,X3)) )),
  inference(resolution,[],[f379,f398])).
fof(f398,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (eqangle(X0,X1,X4,X5,X2,X3,X4,X5) | ~para(X0,X1,X2,X3)) )),
  inference(cnf_transformation,[],[f219])).
fof(f219,plain,(
  ! [X0,X1,X2,X3,X4,X5] : (eqangle(X0,X1,X4,X5,X2,X3,X4,X5) | ~para(X0,X1,X2,X3))),
  inference(ennf_transformation,[],[f119])).
fof(f119,plain,(
  ! [X0,X1,X2,X3,X4,X5] : (para(X0,X1,X2,X3) => eqangle(X0,X1,X4,X5,X2,X3,X4,X5))),
  inference(rectify,[],[f40])).
fof(f40,axiom,(
  ! [X0,X1,X2,X3,X8,X9] : (para(X0,X1,X2,X3) => eqangle(X0,X1,X8,X9,X2,X3,X8,X9))),
  file('./7_Book/00EE/04/E044-6.p',ruleD40)).
fof(f379,plain,(
  ( ! [X6,X4,X2,X0,X7,X5,X3,X1] : (~eqangle(X0,X1,X2,X3,X4,X5,X6,X7) | eqangle(X0,X1,X4,X5,X2,X3,X6,X7)) )),
  inference(cnf_transformation,[],[f195])).
fof(f195,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6,X7] : (eqangle(X0,X1,X4,X5,X2,X3,X6,X7) | ~eqangle(X0,X1,X2,X3,X4,X5,X6,X7))),
  inference(ennf_transformation,[],[f103])).
fof(f103,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6,X7] : (eqangle(X0,X1,X2,X3,X4,X5,X6,X7) => eqangle(X0,X1,X4,X5,X2,X3,X6,X7))),
  inference(rectify,[],[f21])).
fof(f21,axiom,(
  ! [X0,X1,X2,X3,X8,X9,X10,X11] : (eqangle(X0,X1,X2,X3,X8,X9,X10,X11) => eqangle(X0,X1,X8,X9,X2,X3,X10,X11))),
  file('./7_Book/00EE/04/E044-6.p',ruleD21)).
fof(f2180,plain,(
  ( ! [X68,X66,X67] : (~para(X66,X67,X66,X67) | ~coll(X66,X68,X67) | cyclic(X67,X67,X66,X68)) )),
  inference(resolution,[],[f884,f401])).
fof(f401,plain,(
  ( ! [X2,X0,X3,X1] : (~eqangle(X2,X0,X2,X1,X3,X0,X3,X1) | ~coll(X2,X3,X1) | cyclic(X0,X1,X2,X3)) )),
  inference(cnf_transformation,[],[f224])).
fof(f224,plain,(
  ! [X0,X1,X2,X3] : (cyclic(X0,X1,X2,X3) | ~coll(X2,X3,X1) | ~eqangle(X2,X0,X2,X1,X3,X0,X3,X1))),
  inference(flattening,[],[f223])).
fof(f223,plain,(
  ! [X0,X1,X2,X3] : (cyclic(X0,X1,X2,X3) | (~coll(X2,X3,X1) | ~eqangle(X2,X0,X2,X1,X3,X0,X3,X1)))),
  inference(ennf_transformation,[],[f122])).
fof(f122,plain,(
  ! [X0,X1,X2,X3] : ((coll(X2,X3,X1) & eqangle(X2,X0,X2,X1,X3,X0,X3,X1)) => cyclic(X0,X1,X2,X3))),
  inference(rectify,[],[f43])).
fof(f43,axiom,(
  ! [X0,X1,X8,X9] : ((coll(X8,X9,X1) & eqangle(X8,X0,X8,X1,X9,X0,X9,X1)) => cyclic(X0,X1,X8,X9))),
  file('./7_Book/00EE/04/E044-6.p',ruleD42b)).
fof(f8371,plain,(
  ( ! [X35,X36] : (cong(X35,X35,X35,X35) | ~cyclic(X35,X35,X36,X35)) ) | ~spl33_176),
  inference(duplicate_literal_removal,[],[f8358])).
fof(f8358,plain,(
  ( ! [X35,X36] : (cong(X35,X35,X35,X35) | ~cyclic(X35,X35,X36,X35) | ~cyclic(X35,X35,X36,X35)) ) | ~spl33_176),
  inference(resolution,[],[f8333,f1342])).
fof(f1342,plain,(
  ( ! [X10,X8,X7,X9] : (~cyclic(X7,X8,X9,X10) | cong(X7,X8,X7,X8) | ~cyclic(X7,X8,X9,X8) | ~cyclic(X7,X8,X9,X7)) )),
  inference(duplicate_literal_removal,[],[f1336])).
fof(f1336,plain,(
  ( ! [X10,X8,X7,X9] : (cong(X7,X8,X7,X8) | ~cyclic(X7,X8,X9,X10) | ~cyclic(X7,X8,X9,X8) | ~cyclic(X7,X8,X9,X7) | ~cyclic(X7,X8,X9,X10)) )),
  inference(resolution,[],[f402,f399])).
fof(f399,plain,(
  ( ! [X2,X0,X3,X1] : (eqangle(X2,X0,X2,X1,X3,X0,X3,X1) | ~cyclic(X0,X1,X2,X3)) )),
  inference(cnf_transformation,[],[f220])).
fof(f220,plain,(
  ! [X0,X1,X2,X3] : (eqangle(X2,X0,X2,X1,X3,X0,X3,X1) | ~cyclic(X0,X1,X2,X3))),
  inference(ennf_transformation,[],[f120])).
fof(f120,plain,(
  ! [X0,X1,X2,X3] : (cyclic(X0,X1,X2,X3) => eqangle(X2,X0,X2,X1,X3,X0,X3,X1))),
  inference(rectify,[],[f41])).
fof(f41,axiom,(
  ! [X0,X1,X8,X9] : (cyclic(X0,X1,X8,X9) => eqangle(X8,X0,X8,X1,X9,X0,X9,X1))),
  file('./7_Book/00EE/04/E044-6.p',ruleD41)).
fof(f402,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (~eqangle(X2,X0,X2,X1,X5,X3,X5,X4) | cong(X0,X1,X3,X4) | ~cyclic(X0,X1,X2,X5) | ~cyclic(X0,X1,X2,X4) | ~cyclic(X0,X1,X2,X3)) )),
  inference(cnf_transformation,[],[f226])).
fof(f226,plain,(
  ! [X0,X1,X2,X3,X4,X5] : (cong(X0,X1,X3,X4) | ~eqangle(X2,X0,X2,X1,X5,X3,X5,X4) | ~cyclic(X0,X1,X2,X5) | ~cyclic(X0,X1,X2,X4) | ~cyclic(X0,X1,X2,X3))),
  inference(flattening,[],[f225])).
fof(f225,plain,(
  ! [X0,X1,X2,X3,X4,X5] : (cong(X0,X1,X3,X4) | (~eqangle(X2,X0,X2,X1,X5,X3,X5,X4) | ~cyclic(X0,X1,X2,X5) | ~cyclic(X0,X1,X2,X4) | ~cyclic(X0,X1,X2,X3)))),
  inference(ennf_transformation,[],[f123])).
fof(f123,plain,(
  ! [X0,X1,X2,X3,X4,X5] : ((eqangle(X2,X0,X2,X1,X5,X3,X5,X4) & cyclic(X0,X1,X2,X5) & cyclic(X0,X1,X2,X4) & cyclic(X0,X1,X2,X3)) => cong(X0,X1,X3,X4))),
  inference(rectify,[],[f44])).
fof(f44,axiom,(
  ! [X0,X1,X2,X8,X9,X14] : ((eqangle(X2,X0,X2,X1,X14,X8,X14,X9) & cyclic(X0,X1,X2,X14) & cyclic(X0,X1,X2,X9) & cyclic(X0,X1,X2,X8)) => cong(X0,X1,X8,X9))),
  file('./7_Book/00EE/04/E044-6.p',ruleD43)).
fof(f9049,plain,(
  ( ! [X14,X12,X13] : (cong(X12,X13,X12,X14) | ~midp(X12,X12,X13)) ) | ~spl33_176),
  inference(resolution,[],[f8543,f1888])).
fof(f1888,plain,(
  ( ! [X2,X0,X3,X1] : (~midp(X0,X3,X2) | cong(X0,X1,X0,X2) | ~midp(X0,X3,X1)) )),
  inference(resolution,[],[f830,f524])).
fof(f524,plain,(
  ( ! [X2,X0,X1] : (cong(X0,X1,X0,X2) | ~midp(X0,X2,X1)) )),
  inference(resolution,[],[f382,f427])).
fof(f427,plain,(
  ( ! [X2,X0,X1] : (cong(X0,X1,X0,X2) | ~midp(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f270])).
fof(f270,plain,(
  ! [X0,X1,X2] : (cong(X0,X1,X0,X2) | ~midp(X0,X1,X2))),
  inference(ennf_transformation,[],[f69])).
fof(f69,axiom,(
  ! [X0,X1,X2] : (midp(X0,X1,X2) => cong(X0,X1,X0,X2))),
  file('./7_Book/00EE/04/E044-6.p',ruleD68)).
fof(f382,plain,(
  ( ! [X2,X0,X3,X1] : (~cong(X0,X1,X2,X3) | cong(X2,X3,X0,X1)) )),
  inference(cnf_transformation,[],[f199])).
fof(f199,plain,(
  ! [X0,X1,X2,X3] : (cong(X2,X3,X0,X1) | ~cong(X0,X1,X2,X3))),
  inference(ennf_transformation,[],[f24])).
fof(f24,axiom,(
  ! [X0,X1,X2,X3] : (cong(X0,X1,X2,X3) => cong(X2,X3,X0,X1))),
  file('./7_Book/00EE/04/E044-6.p',ruleD24)).
fof(f830,plain,(
  ( ! [X4,X2,X0,X3,X1] : (~cong(X0,X1,X2,X4) | cong(X0,X1,X2,X3) | ~midp(X2,X4,X3)) )),
  inference(resolution,[],[f383,f427])).
fof(f383,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (~cong(X2,X3,X4,X5) | cong(X0,X1,X4,X5) | ~cong(X0,X1,X2,X3)) )),
  inference(cnf_transformation,[],[f201])).
fof(f201,plain,(
  ! [X0,X1,X2,X3,X4,X5] : (cong(X0,X1,X4,X5) | ~cong(X2,X3,X4,X5) | ~cong(X0,X1,X2,X3))),
  inference(flattening,[],[f200])).
fof(f200,plain,(
  ! [X0,X1,X2,X3,X4,X5] : (cong(X0,X1,X4,X5) | (~cong(X2,X3,X4,X5) | ~cong(X0,X1,X2,X3)))),
  inference(ennf_transformation,[],[f25])).
fof(f25,axiom,(
  ! [X0,X1,X2,X3,X4,X5] : ((cong(X2,X3,X4,X5) & cong(X0,X1,X2,X3)) => cong(X0,X1,X4,X5))),
  file('./7_Book/00EE/04/E044-6.p',ruleD25)).
fof(f12175,plain,(
  ( ! [X26,X24,X23,X25] : (circle(X23,X25,X26,X24) | ~cong(X23,X25,X23,X26)) ) | (~spl33_176 | ~spl33_232)),
  inference(subsumption_resolution,[],[f12174,f8545])).
fof(f8545,plain,(
  ( ! [X17,X16] : (midp(X16,X17,X16)) ) | (~spl33_176 | ~spl33_232)),
  inference(resolution,[],[f8489,f7736])).
fof(f12174,plain,(
  ( ! [X26,X24,X23,X25] : (~midp(X23,X25,X23) | circle(X23,X25,X26,X24) | ~cong(X23,X25,X23,X26)) ) | ~spl33_176),
  inference(subsumption_resolution,[],[f12141,f8543])).
fof(f12141,plain,(
  ( ! [X26,X24,X23,X25] : (~midp(X23,X23,X24) | ~midp(X23,X25,X23) | circle(X23,X25,X26,X24) | ~cong(X23,X25,X23,X26)) )),
  inference(resolution,[],[f1891,f370])).
fof(f370,plain,(
  ( ! [X2,X0,X3,X1] : (~cong(X3,X0,X3,X2) | circle(X3,X0,X1,X2) | ~cong(X3,X0,X3,X1)) )),
  inference(cnf_transformation,[],[f184])).
fof(f184,plain,(
  ! [X0,X1,X2,X3] : (circle(X3,X0,X1,X2) | ~cong(X3,X0,X3,X2) | ~cong(X3,X0,X3,X1))),
  inference(flattening,[],[f183])).
fof(f183,plain,(
  ! [X0,X1,X2,X3] : (circle(X3,X0,X1,X2) | (~cong(X3,X0,X3,X2) | ~cong(X3,X0,X3,X1)))),
  inference(ennf_transformation,[],[f98])).
fof(f98,plain,(
  ! [X0,X1,X2,X3] : ((cong(X3,X0,X3,X2) & cong(X3,X0,X3,X1)) => circle(X3,X0,X1,X2))),
  inference(rectify,[],[f12])).
fof(f12,axiom,(
  ! [X0,X1,X2,X7] : ((cong(X7,X0,X7,X2) & cong(X7,X0,X7,X1)) => circle(X7,X0,X1,X2))),
  file('./7_Book/00EE/04/E044-6.p',ruleD12)).
fof(f1891,plain,(
  ( ! [X14,X12,X15,X13] : (cong(X12,X13,X14,X15) | ~midp(X14,X12,X15) | ~midp(X12,X13,X14)) )),
  inference(resolution,[],[f830,f521])).
fof(f521,plain,(
  ( ! [X2,X0,X1] : (cong(X0,X1,X2,X0) | ~midp(X0,X1,X2)) )),
  inference(resolution,[],[f381,f427])).
fof(f381,plain,(
  ( ! [X2,X0,X3,X1] : (~cong(X0,X1,X2,X3) | cong(X0,X1,X3,X2)) )),
  inference(cnf_transformation,[],[f198])).
fof(f198,plain,(
  ! [X0,X1,X2,X3] : (cong(X0,X1,X3,X2) | ~cong(X0,X1,X2,X3))),
  inference(ennf_transformation,[],[f23])).
fof(f23,axiom,(
  ! [X0,X1,X2,X3] : (cong(X0,X1,X2,X3) => cong(X0,X1,X3,X2))),
  file('./7_Book/00EE/04/E044-6.p',ruleD23)).
fof(f13123,plain,(
  ( ! [X9] : (~perp(sK20,sK21,sK21,X9) | para(sK21,X9,sK32,sK21)) ) | ~spl33_662),
  inference(avatar_component_clause,[],[f13122])).
fof(f13122,plain,(
  spl33_662 <=> ! [X9] : (~perp(sK20,sK21,sK21,X9) | para(sK21,X9,sK32,sK21))),
  introduced(avatar_definition,[new_symbols(naming,[spl33_662])])).
fof(f17877,plain,(
  ( ! [X4,X2,X5,X3] : (para(X2,X3,X4,X5) | ~midp(X2,sK21,X4)) ) | (~spl33_160 | ~spl33_176 | ~spl33_232 | ~spl33_492 | ~spl33_662)),
  inference(resolution,[],[f17831,f403])).
fof(f403,plain,(
  ( ! [X4,X2,X0,X3,X1] : (~midp(X4,X0,X2) | para(X3,X4,X1,X2) | ~midp(X3,X0,X1)) )),
  inference(cnf_transformation,[],[f228])).
fof(f228,plain,(
  ! [X0,X1,X2,X3,X4] : (para(X3,X4,X1,X2) | ~midp(X4,X0,X2) | ~midp(X3,X0,X1))),
  inference(flattening,[],[f227])).
fof(f227,plain,(
  ! [X0,X1,X2,X3,X4] : (para(X3,X4,X1,X2) | (~midp(X4,X0,X2) | ~midp(X3,X0,X1)))),
  inference(ennf_transformation,[],[f124])).
fof(f124,plain,(
  ! [X0,X1,X2,X3,X4] : ((midp(X4,X0,X2) & midp(X3,X0,X1)) => para(X3,X4,X1,X2))),
  inference(rectify,[],[f45])).
fof(f45,axiom,(
  ! [X0,X1,X2,X4,X5] : ((midp(X5,X0,X2) & midp(X4,X0,X1)) => para(X4,X5,X1,X2))),
  file('./7_Book/00EE/04/E044-6.p',ruleD44)).
fof(f20727,plain,(
  ( ! [X6,X4,X7,X5] : (midp(X4,X5,X6) | ~para(X7,X6,X7,X4)) ) | (~spl33_176 | ~spl33_232)),
  inference(subsumption_resolution,[],[f20719,f4064])).
fof(f20719,plain,(
  ( ! [X6,X4,X7,X5] : (midp(X4,X5,X6) | ~coll(X4,X5,X6) | ~para(X7,X6,X7,X4)) ) | (~spl33_176 | ~spl33_232)),
  inference(resolution,[],[f1986,f12176])).
fof(f1986,plain,(
  ( ! [X78,X76,X77,X75] : (~circle(X75,X75,X78,X76) | midp(X77,X78,X76) | ~coll(X77,X78,X76) | ~para(X75,X76,X75,X77)) )),
  inference(resolution,[],[f858,f410])).
fof(f410,plain,(
  ( ! [X4,X2,X0,X3,X1] : (~eqangle(X0,X1,X0,X2,X3,X1,X3,X4) | midp(X4,X1,X2) | ~coll(X4,X1,X2) | ~circle(X3,X0,X1,X2)) )),
  inference(cnf_transformation,[],[f241])).
fof(f241,plain,(
  ! [X0,X1,X2,X3,X4] : (midp(X4,X1,X2) | ~eqangle(X0,X1,X0,X2,X3,X1,X3,X4) | ~coll(X4,X1,X2) | ~circle(X3,X0,X1,X2))),
  inference(flattening,[],[f240])).
fof(f240,plain,(
  ! [X0,X1,X2,X3,X4] : (midp(X4,X1,X2) | (~eqangle(X0,X1,X0,X2,X3,X1,X3,X4) | ~coll(X4,X1,X2) | ~circle(X3,X0,X1,X2)))),
  inference(ennf_transformation,[],[f131])).
fof(f131,plain,(
  ! [X0,X1,X2,X3,X4] : ((eqangle(X0,X1,X0,X2,X3,X1,X3,X4) & coll(X4,X1,X2) & circle(X3,X0,X1,X2)) => midp(X4,X1,X2))),
  inference(rectify,[],[f52])).
fof(f52,axiom,(
  ! [X0,X1,X2,X7,X6] : ((eqangle(X0,X1,X0,X2,X7,X1,X7,X6) & coll(X6,X1,X2) & circle(X7,X0,X1,X2)) => midp(X6,X1,X2))),
  file('./7_Book/00EE/04/E044-6.p',ruleD51)).
fof(f858,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (eqangle(X0,X1,X2,X3,X0,X1,X4,X5) | ~para(X2,X3,X4,X5)) )),
  inference(resolution,[],[f377,f398])).
fof(f377,plain,(
  ( ! [X6,X4,X2,X0,X7,X5,X3,X1] : (~eqangle(X0,X1,X2,X3,X4,X5,X6,X7) | eqangle(X2,X3,X0,X1,X6,X7,X4,X5)) )),
  inference(cnf_transformation,[],[f193])).
fof(f193,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6,X7] : (eqangle(X2,X3,X0,X1,X6,X7,X4,X5) | ~eqangle(X0,X1,X2,X3,X4,X5,X6,X7))),
  inference(ennf_transformation,[],[f101])).
fof(f101,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6,X7] : (eqangle(X0,X1,X2,X3,X4,X5,X6,X7) => eqangle(X2,X3,X0,X1,X6,X7,X4,X5))),
  inference(rectify,[],[f19])).
fof(f19,axiom,(
  ! [X0,X1,X2,X3,X8,X9,X10,X11] : (eqangle(X0,X1,X2,X3,X8,X9,X10,X11) => eqangle(X2,X3,X0,X1,X10,X11,X8,X9))),
  file('./7_Book/00EE/04/E044-6.p',ruleD19)).
fof(f19149,plain,(
  ~spl33_176 | ~spl33_232 | ~spl33_434 | spl33_577),
  inference(avatar_contradiction_clause,[],[f19133])).
fof(f19133,plain,(
  $false | (~spl33_176 | ~spl33_232 | ~spl33_434 | ~spl33_577)),
  inference(resolution,[],[f18979,f12205])).
fof(f12205,plain,(
  ~midp(sK20,sK26,sK27) | ~spl33_577),
  inference(avatar_component_clause,[],[f12204])).
fof(f12204,plain,(
  spl33_577 <=> ~midp(sK20,sK26,sK27)),
  introduced(avatar_definition,[new_symbols(naming,[spl33_577])])).
fof(f18979,plain,(
  ( ! [X23] : (midp(sK20,sK26,X23)) ) | (~spl33_176 | ~spl33_232 | ~spl33_434)),
  inference(resolution,[],[f18952,f7736])).
fof(f18952,plain,(
  ( ! [X0] : (midp(sK26,sK20,X0)) ) | (~spl33_176 | ~spl33_232 | ~spl33_434)),
  inference(resolution,[],[f18951,f369])).
fof(f18951,plain,(
  ( ! [X3] : (midp(sK26,X3,sK20)) ) | (~spl33_176 | ~spl33_232 | ~spl33_434)),
  inference(subsumption_resolution,[],[f18950,f8545])).
fof(f18950,plain,(
  ( ! [X3] : (midp(sK26,X3,sK20) | ~midp(sK26,X3,sK26)) ) | (~spl33_176 | ~spl33_434)),
  inference(subsumption_resolution,[],[f18939,f4064])).
fof(f18939,plain,(
  ( ! [X3] : (~coll(sK26,X3,sK20) | midp(sK26,X3,sK20) | ~midp(sK26,X3,sK26)) ) | ~spl33_434),
  inference(resolution,[],[f12245,f404])).
fof(f12245,plain,(
  para(sK26,sK26,sK26,sK20) | ~spl33_434),
  inference(resolution,[],[f7600,f806])).
fof(f806,plain,(
  ( ! [X6,X7] : (~perp(X6,X7,sK21,sK22) | para(X6,X7,sK26,sK20)) )),
  inference(resolution,[],[f367,f512])).
fof(f512,plain,(
  perp(sK21,sK22,sK26,sK20)),
  inference(resolution,[],[f366,f479])).
fof(f479,plain,(
  perp(sK26,sK20,sK21,sK22)),
  inference(cnf_transformation,[],[f358])).
fof(f358,plain,(
  ~cong(sK20,sK26,sK27,sK24) & midp(sK27,sK23,sK25) & coll(sK26,sK21,sK22) & perp(sK26,sK20,sK21,sK22) & coll(sK25,sK21,sK24) & circle(sK20,sK21,sK25,sK32) & coll(sK24,sK22,sK23) & perp(sK24,sK21,sK22,sK23) & circle(sK20,sK21,sK23,sK31) & circle(sK20,sK21,sK22,sK30) & circle(sK20,sK21,sK28,sK29)),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK20,sK21,sK22,sK23,sK24,sK25,sK26,sK27,sK28,sK29,sK30,sK31,sK32])],[f320,f357])).
fof(f357,plain,(
  ? [X0,X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12] : (~cong(X0,X6,X7,X4) & midp(X7,X3,X5) & coll(X6,X1,X2) & perp(X6,X0,X1,X2) & coll(X5,X1,X4) & circle(X0,X1,X5,X12) & coll(X4,X2,X3) & perp(X4,X1,X2,X3) & circle(X0,X1,X3,X11) & circle(X0,X1,X2,X10) & circle(X0,X1,X8,X9)) => (~cong(sK20,sK26,sK27,sK24) & midp(sK27,sK23,sK25) & coll(sK26,sK21,sK22) & perp(sK26,sK20,sK21,sK22) & coll(sK25,sK21,sK24) & circle(sK20,sK21,sK25,sK32) & coll(sK24,sK22,sK23) & perp(sK24,sK21,sK22,sK23) & circle(sK20,sK21,sK23,sK31) & circle(sK20,sK21,sK22,sK30) & circle(sK20,sK21,sK28,sK29))),
  introduced(choice_axiom,[])).
fof(f320,plain,(
  ? [X0,X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12] : (~cong(X0,X6,X7,X4) & midp(X7,X3,X5) & coll(X6,X1,X2) & perp(X6,X0,X1,X2) & coll(X5,X1,X4) & circle(X0,X1,X5,X12) & coll(X4,X2,X3) & perp(X4,X1,X2,X3) & circle(X0,X1,X3,X11) & circle(X0,X1,X2,X10) & circle(X0,X1,X8,X9))),
  inference(flattening,[],[f319])).
fof(f319,plain,(
  ? [X0,X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12] : (~cong(X0,X6,X7,X4) & (midp(X7,X3,X5) & coll(X6,X1,X2) & perp(X6,X0,X1,X2) & coll(X5,X1,X4) & circle(X0,X1,X5,X12) & coll(X4,X2,X3) & perp(X4,X1,X2,X3) & circle(X0,X1,X3,X11) & circle(X0,X1,X2,X10) & circle(X0,X1,X8,X9)))),
  inference(ennf_transformation,[],[f167])).
fof(f167,plain,(
  ~! [X0,X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12] : ((midp(X7,X3,X5) & coll(X6,X1,X2) & perp(X6,X0,X1,X2) & coll(X5,X1,X4) & circle(X0,X1,X5,X12) & coll(X4,X2,X3) & perp(X4,X1,X2,X3) & circle(X0,X1,X3,X11) & circle(X0,X1,X2,X10) & circle(X0,X1,X8,X9)) => cong(X0,X6,X7,X4))),
  inference(rectify,[],[f96])).
fof(f96,negated_conjecture,(
  ~! [X0,X1,X2,X3,X4,X5,X12,X13,X18,X19,X20,X21,X22] : ((midp(X13,X3,X5) & coll(X12,X1,X2) & perp(X12,X0,X1,X2) & coll(X5,X1,X4) & circle(X0,X1,X5,X22) & coll(X4,X2,X3) & perp(X4,X1,X2,X3) & circle(X0,X1,X3,X21) & circle(X0,X1,X2,X20) & circle(X0,X1,X18,X19)) => cong(X0,X12,X13,X4))),
  inference(negated_conjecture,[],[f95])).
fof(f95,conjecture,(
  ! [X0,X1,X2,X3,X4,X5,X12,X13,X18,X19,X20,X21,X22] : ((midp(X13,X3,X5) & coll(X12,X1,X2) & perp(X12,X0,X1,X2) & coll(X5,X1,X4) & circle(X0,X1,X5,X22) & coll(X4,X2,X3) & perp(X4,X1,X2,X3) & circle(X0,X1,X3,X21) & circle(X0,X1,X2,X20) & circle(X0,X1,X18,X19)) => cong(X0,X12,X13,X4))),
  file('./7_Book/00EE/04/E044-6.p',exemplo7Book00EE04E0446)).
fof(f366,plain,(
  ( ! [X2,X0,X3,X1] : (~perp(X0,X1,X2,X3) | perp(X2,X3,X0,X1)) )),
  inference(cnf_transformation,[],[f177])).
fof(f177,plain,(
  ! [X0,X1,X2,X3] : (perp(X2,X3,X0,X1) | ~perp(X0,X1,X2,X3))),
  inference(ennf_transformation,[],[f8])).
fof(f8,axiom,(
  ! [X0,X1,X2,X3] : (perp(X0,X1,X2,X3) => perp(X2,X3,X0,X1))),
  file('./7_Book/00EE/04/E044-6.p',ruleD8)).
fof(f367,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (~perp(X2,X3,X4,X5) | para(X0,X1,X4,X5) | ~perp(X0,X1,X2,X3)) )),
  inference(cnf_transformation,[],[f179])).
fof(f179,plain,(
  ! [X0,X1,X2,X3,X4,X5] : (para(X0,X1,X4,X5) | ~perp(X2,X3,X4,X5) | ~perp(X0,X1,X2,X3))),
  inference(flattening,[],[f178])).
fof(f178,plain,(
  ! [X0,X1,X2,X3,X4,X5] : (para(X0,X1,X4,X5) | (~perp(X2,X3,X4,X5) | ~perp(X0,X1,X2,X3)))),
  inference(ennf_transformation,[],[f9])).
fof(f9,axiom,(
  ! [X0,X1,X2,X3,X4,X5] : ((perp(X2,X3,X4,X5) & perp(X0,X1,X2,X3)) => para(X0,X1,X4,X5))),
  file('./7_Book/00EE/04/E044-6.p',ruleD9)).
fof(f7600,plain,(
  perp(sK26,sK26,sK21,sK22) | ~spl33_434),
  inference(avatar_component_clause,[],[f7599])).
fof(f7599,plain,(
  spl33_434 <=> perp(sK26,sK26,sK21,sK22)),
  introduced(avatar_definition,[new_symbols(naming,[spl33_434])])).
fof(f17771,plain,(
  spl33_660 | ~spl33_502),
  inference(avatar_split_clause,[],[f17756,f8685,f13116])).
fof(f13116,plain,(
  spl33_660 <=> para(sK21,sK25,sK32,sK25)),
  introduced(avatar_definition,[new_symbols(naming,[spl33_660])])).
fof(f8685,plain,(
  spl33_502 <=> para(sK21,sK25,sK25,sK32)),
  introduced(avatar_definition,[new_symbols(naming,[spl33_502])])).
fof(f17756,plain,(
  para(sK21,sK25,sK32,sK25) | ~spl33_502),
  inference(resolution,[],[f8686,f362])).
fof(f8686,plain,(
  para(sK21,sK25,sK25,sK32) | ~spl33_502),
  inference(avatar_component_clause,[],[f8685])).
fof(f17668,plain,(
  spl33_502 | ~spl33_8 | ~spl33_176 | ~spl33_182),
  inference(avatar_split_clause,[],[f17659,f2195,f1994,f625,f8685])).
fof(f625,plain,(
  spl33_8 <=> perp(sK21,sK25,sK25,sK32)),
  introduced(avatar_definition,[new_symbols(naming,[spl33_8])])).
fof(f2195,plain,(
  spl33_182 <=> ! [X44,X46,X45,X47] : (~para(X44,X45,X46,X47) | perp(X44,X45,X46,X47))),
  introduced(avatar_definition,[new_symbols(naming,[spl33_182])])).
fof(f17659,plain,(
  para(sK21,sK25,sK25,sK32) | (~spl33_8 | ~spl33_176 | ~spl33_182)),
  inference(resolution,[],[f13971,f2272])).
fof(f2272,plain,(
  ( ! [X8,X7] : (~perp(X7,X8,sK21,sK25) | para(X7,X8,sK25,sK32)) ) | ~spl33_8),
  inference(resolution,[],[f626,f367])).
fof(f626,plain,(
  perp(sK21,sK25,sK25,sK32) | ~spl33_8),
  inference(avatar_component_clause,[],[f625])).
fof(f13971,plain,(
  ( ! [X0,X1] : (perp(X0,X1,X0,X1)) ) | (~spl33_176 | ~spl33_182)),
  inference(resolution,[],[f2196,f1995])).
fof(f2196,plain,(
  ( ! [X47,X45,X46,X44] : (~para(X44,X45,X46,X47) | perp(X44,X45,X46,X47)) ) | ~spl33_182),
  inference(avatar_component_clause,[],[f2195])).
fof(f15358,plain,(
  ~spl33_176 | ~spl33_732),
  inference(avatar_contradiction_clause,[],[f15357])).
fof(f15357,plain,(
  $false | (~spl33_176 | ~spl33_732)),
  inference(subsumption_resolution,[],[f15355,f8543])).
fof(f15355,plain,(
  ~midp(sK21,sK21,sK25) | ~spl33_732),
  inference(resolution,[],[f13901,f1486])).
fof(f1486,plain,(
  ( ! [X2,X3] : (para(X2,X2,X2,X3) | ~midp(X2,X2,X3)) )),
  inference(resolution,[],[f586,f427])).
fof(f586,plain,(
  ( ! [X0,X1] : (~cong(X0,X0,X0,X1) | para(X0,X0,X0,X1)) )),
  inference(resolution,[],[f405,f397])).
fof(f397,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (~eqangle(X0,X1,X4,X5,X2,X3,X4,X5) | para(X0,X1,X2,X3)) )),
  inference(cnf_transformation,[],[f218])).
fof(f218,plain,(
  ! [X0,X1,X2,X3,X4,X5] : (para(X0,X1,X2,X3) | ~eqangle(X0,X1,X4,X5,X2,X3,X4,X5))),
  inference(ennf_transformation,[],[f118])).
fof(f118,plain,(
  ! [X0,X1,X2,X3,X4,X5] : (eqangle(X0,X1,X4,X5,X2,X3,X4,X5) => para(X0,X1,X2,X3))),
  inference(rectify,[],[f39])).
fof(f39,axiom,(
  ! [X0,X1,X2,X3,X8,X9] : (eqangle(X0,X1,X8,X9,X2,X3,X8,X9) => para(X0,X1,X2,X3))),
  file('./7_Book/00EE/04/E044-6.p',ruleD39)).
fof(f405,plain,(
  ( ! [X2,X0,X1] : (eqangle(X2,X0,X0,X1,X0,X1,X2,X1) | ~cong(X2,X0,X2,X1)) )),
  inference(cnf_transformation,[],[f231])).
fof(f231,plain,(
  ! [X0,X1,X2] : (eqangle(X2,X0,X0,X1,X0,X1,X2,X1) | ~cong(X2,X0,X2,X1))),
  inference(ennf_transformation,[],[f126])).
fof(f126,plain,(
  ! [X0,X1,X2] : (cong(X2,X0,X2,X1) => eqangle(X2,X0,X0,X1,X0,X1,X2,X1))),
  inference(rectify,[],[f47])).
fof(f47,axiom,(
  ! [X0,X1,X7] : (cong(X7,X0,X7,X1) => eqangle(X7,X0,X0,X1,X0,X1,X7,X1))),
  file('./7_Book/00EE/04/E044-6.p',ruleD46)).
fof(f13901,plain,(
  ( ! [X28] : (~para(X28,sK21,X28,sK25)) ) | ~spl33_732),
  inference(avatar_component_clause,[],[f13900])).
fof(f13900,plain,(
  spl33_732 <=> ! [X28] : ~para(X28,sK21,X28,sK25)),
  introduced(avatar_definition,[new_symbols(naming,[spl33_732])])).
fof(f13902,plain,(
  spl33_732 | spl33_710 | spl33_161 | ~spl33_176),
  inference(avatar_split_clause,[],[f13898,f1994,f1817,f13565,f13900])).
fof(f13565,plain,(
  spl33_710 <=> ! [X27] : ~perp(sK20,sK21,sK21,X27)),
  introduced(avatar_definition,[new_symbols(naming,[spl33_710])])).
fof(f1817,plain,(
  spl33_161 <=> ~para(sK32,sK21,sK32,sK25)),
  introduced(avatar_definition,[new_symbols(naming,[spl33_161])])).
fof(f13898,plain,(
  ( ! [X28,X27] : (~perp(sK20,sK21,sK21,X27) | ~para(X28,sK21,X28,sK25)) ) | (~spl33_161 | ~spl33_176)),
  inference(subsumption_resolution,[],[f13614,f13058])).
fof(f13058,plain,(
  ( ! [X17] : (~perp(sK20,sK21,sK21,X17) | ~para(sK21,X17,sK21,sK25)) ) | ~spl33_161),
  inference(subsumption_resolution,[],[f13044,f1818])).
fof(f1818,plain,(
  ~para(sK32,sK21,sK32,sK25) | ~spl33_161),
  inference(avatar_component_clause,[],[f1817])).
fof(f13044,plain,(
  ( ! [X17] : (~perp(sK20,sK21,sK21,X17) | ~para(sK21,X17,sK21,sK25) | para(sK32,sK21,sK32,sK25)) )),
  inference(resolution,[],[f1803,f432])).
fof(f432,plain,(
  ( ! [X6,X4,X2,X0,X7,X5,X3,X1] : (~eqangle(X0,X1,X2,X3,X4,X5,X6,X7) | ~para(X4,X5,X6,X7) | para(X0,X1,X2,X3)) )),
  inference(cnf_transformation,[],[f279])).
fof(f279,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6,X7] : (para(X0,X1,X2,X3) | ~para(X4,X5,X6,X7) | ~eqangle(X0,X1,X2,X3,X4,X5,X6,X7))),
  inference(flattening,[],[f278])).
fof(f278,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6,X7] : (para(X0,X1,X2,X3) | (~para(X4,X5,X6,X7) | ~eqangle(X0,X1,X2,X3,X4,X5,X6,X7)))),
  inference(ennf_transformation,[],[f146])).
fof(f146,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6,X7] : ((para(X4,X5,X6,X7) & eqangle(X0,X1,X2,X3,X4,X5,X6,X7)) => para(X0,X1,X2,X3))),
  inference(rectify,[],[f74])).
fof(f74,axiom,(
  ! [X0,X1,X2,X3,X8,X9,X10,X11] : ((para(X8,X9,X10,X11) & eqangle(X0,X1,X2,X3,X8,X9,X10,X11)) => para(X0,X1,X2,X3))),
  file('./7_Book/00EE/04/E044-6.p',ruleD73)).
fof(f1803,plain,(
  ( ! [X2] : (eqangle(sK32,sK21,sK32,sK25,sK21,X2,sK21,sK25) | ~perp(sK20,sK21,sK21,X2)) )),
  inference(resolution,[],[f1022,f378])).
fof(f378,plain,(
  ( ! [X6,X4,X2,X0,X7,X5,X3,X1] : (~eqangle(X0,X1,X2,X3,X4,X5,X6,X7) | eqangle(X4,X5,X6,X7,X0,X1,X2,X3)) )),
  inference(cnf_transformation,[],[f194])).
fof(f194,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6,X7] : (eqangle(X4,X5,X6,X7,X0,X1,X2,X3) | ~eqangle(X0,X1,X2,X3,X4,X5,X6,X7))),
  inference(ennf_transformation,[],[f102])).
fof(f102,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6,X7] : (eqangle(X0,X1,X2,X3,X4,X5,X6,X7) => eqangle(X4,X5,X6,X7,X0,X1,X2,X3))),
  inference(rectify,[],[f20])).
fof(f20,axiom,(
  ! [X0,X1,X2,X3,X8,X9,X10,X11] : (eqangle(X0,X1,X2,X3,X8,X9,X10,X11) => eqangle(X8,X9,X10,X11,X0,X1,X2,X3))),
  file('./7_Book/00EE/04/E044-6.p',ruleD20)).
fof(f1022,plain,(
  ( ! [X2] : (eqangle(sK21,X2,sK21,sK25,sK32,sK21,sK32,sK25) | ~perp(sK20,sK21,sK21,X2)) )),
  inference(resolution,[],[f407,f477])).
fof(f477,plain,(
  circle(sK20,sK21,sK25,sK32)),
  inference(cnf_transformation,[],[f358])).
fof(f407,plain,(
  ( ! [X4,X2,X0,X3,X1] : (~circle(X3,X0,X1,X2) | ~perp(X3,X0,X0,X4) | eqangle(X0,X4,X0,X1,X2,X0,X2,X1)) )),
  inference(cnf_transformation,[],[f235])).
fof(f235,plain,(
  ! [X0,X1,X2,X3,X4] : (eqangle(X0,X4,X0,X1,X2,X0,X2,X1) | ~perp(X3,X0,X0,X4) | ~circle(X3,X0,X1,X2))),
  inference(flattening,[],[f234])).
fof(f234,plain,(
  ! [X0,X1,X2,X3,X4] : (eqangle(X0,X4,X0,X1,X2,X0,X2,X1) | (~perp(X3,X0,X0,X4) | ~circle(X3,X0,X1,X2)))),
  inference(ennf_transformation,[],[f128])).
fof(f128,plain,(
  ! [X0,X1,X2,X3,X4] : ((perp(X3,X0,X0,X4) & circle(X3,X0,X1,X2)) => eqangle(X0,X4,X0,X1,X2,X0,X2,X1))),
  inference(rectify,[],[f49])).
fof(f49,axiom,(
  ! [X0,X1,X2,X7,X15] : ((perp(X7,X0,X0,X15) & circle(X7,X0,X1,X2)) => eqangle(X0,X15,X0,X1,X2,X0,X2,X1))),
  file('./7_Book/00EE/04/E044-6.p',ruleD48)).
fof(f13614,plain,(
  ( ! [X28,X27] : (~perp(sK20,sK21,sK21,X27) | ~para(X28,sK21,X28,sK25) | para(sK21,X27,sK21,sK25)) ) | ~spl33_176),
  inference(resolution,[],[f10829,f432])).
fof(f10829,plain,(
  ( ! [X154,X155] : (eqangle(sK21,X154,sK21,sK25,X155,sK21,X155,sK25) | ~perp(sK20,sK21,sK21,X154)) ) | ~spl33_176),
  inference(subsumption_resolution,[],[f4650,f8372])).
fof(f8372,plain,(
  ( ! [X12,X10,X11,X9] : (cyclic(X9,X10,X11,X12)) ) | ~spl33_176),
  inference(subsumption_resolution,[],[f8350,f8333])).
fof(f8350,plain,(
  ( ! [X12,X10,X11,X9] : (cyclic(X9,X10,X11,X12) | ~cyclic(X9,X9,X10,X11)) ) | ~spl33_176),
  inference(resolution,[],[f8333,f375])).
fof(f375,plain,(
  ( ! [X4,X2,X0,X3,X1] : (~cyclic(X0,X1,X2,X4) | cyclic(X1,X2,X3,X4) | ~cyclic(X0,X1,X2,X3)) )),
  inference(cnf_transformation,[],[f191])).
fof(f191,plain,(
  ! [X0,X1,X2,X3,X4] : (cyclic(X1,X2,X3,X4) | ~cyclic(X0,X1,X2,X4) | ~cyclic(X0,X1,X2,X3))),
  inference(flattening,[],[f190])).
fof(f190,plain,(
  ! [X0,X1,X2,X3,X4] : (cyclic(X1,X2,X3,X4) | (~cyclic(X0,X1,X2,X4) | ~cyclic(X0,X1,X2,X3)))),
  inference(ennf_transformation,[],[f17])).
fof(f17,axiom,(
  ! [X0,X1,X2,X3,X4] : ((cyclic(X0,X1,X2,X4) & cyclic(X0,X1,X2,X3)) => cyclic(X1,X2,X3,X4))),
  file('./7_Book/00EE/04/E044-6.p',ruleD17)).
fof(f4650,plain,(
  ( ! [X154,X155] : (eqangle(sK21,X154,sK21,sK25,X155,sK21,X155,sK25) | ~cyclic(sK21,sK25,sK32,X155) | ~perp(sK20,sK21,sK21,X154)) )),
  inference(resolution,[],[f1260,f1022])).
fof(f1260,plain,(
  ( ! [X24,X23,X21,X19,X25,X22,X20,X18] : (~eqangle(X18,X19,X20,X21,X25,X23,X25,X24) | eqangle(X18,X19,X20,X21,X22,X23,X22,X24) | ~cyclic(X23,X24,X25,X22)) )),
  inference(resolution,[],[f380,f399])).
fof(f380,plain,(
  ( ! [X6,X4,X2,X0,X10,X8,X7,X5,X3,X1,X11,X9] : (~eqangle(X4,X5,X6,X7,X8,X9,X10,X11) | eqangle(X0,X1,X2,X3,X8,X9,X10,X11) | ~eqangle(X0,X1,X2,X3,X4,X5,X6,X7)) )),
  inference(cnf_transformation,[],[f197])).
fof(f197,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11] : (eqangle(X0,X1,X2,X3,X8,X9,X10,X11) | ~eqangle(X4,X5,X6,X7,X8,X9,X10,X11) | ~eqangle(X0,X1,X2,X3,X4,X5,X6,X7))),
  inference(flattening,[],[f196])).
fof(f196,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11] : (eqangle(X0,X1,X2,X3,X8,X9,X10,X11) | (~eqangle(X4,X5,X6,X7,X8,X9,X10,X11) | ~eqangle(X0,X1,X2,X3,X4,X5,X6,X7)))),
  inference(ennf_transformation,[],[f104])).
fof(f104,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11] : ((eqangle(X4,X5,X6,X7,X8,X9,X10,X11) & eqangle(X0,X1,X2,X3,X4,X5,X6,X7)) => eqangle(X0,X1,X2,X3,X8,X9,X10,X11))),
  inference(rectify,[],[f22])).
fof(f22,axiom,(
  ! [X0,X1,X2,X3,X8,X9,X10,X11,X4,X5,X12,X13] : ((eqangle(X8,X9,X10,X11,X4,X5,X12,X13) & eqangle(X0,X1,X2,X3,X8,X9,X10,X11)) => eqangle(X0,X1,X2,X3,X4,X5,X12,X13))),
  file('./7_Book/00EE/04/E044-6.p',ruleD22)).
fof(f13865,plain,(
  ~spl33_176 | ~spl33_180),
  inference(avatar_contradiction_clause,[],[f13841])).
fof(f13841,plain,(
  $false | (~spl33_176 | ~spl33_180)),
  inference(resolution,[],[f8969,f2193])).
fof(f2193,plain,(
  ( ! [X48,X49] : (~perp(X48,X49,X48,X49)) ) | ~spl33_180),
  inference(avatar_component_clause,[],[f2192])).
fof(f2192,plain,(
  spl33_180 <=> ! [X49,X48] : ~perp(X48,X49,X48,X49)),
  introduced(avatar_definition,[new_symbols(naming,[spl33_180])])).
fof(f8969,plain,(
  ( ! [X59,X60] : (perp(X59,X60,X60,X60)) ) | ~spl33_176),
  inference(subsumption_resolution,[],[f8968,f8375])).
fof(f8968,plain,(
  ( ! [X59,X60] : (perp(X59,X60,X60,X60) | ~cong(X60,X60,X60,X60)) ) | ~spl33_176),
  inference(subsumption_resolution,[],[f8964,f8489])).
fof(f8964,plain,(
  ( ! [X59,X60] : (~midp(X59,X60,X60) | perp(X59,X60,X60,X60) | ~cong(X60,X60,X60,X60)) )),
  inference(duplicate_literal_removal,[],[f8958])).
fof(f8958,plain,(
  ( ! [X59,X60] : (~midp(X59,X60,X60) | ~midp(X59,X60,X60) | perp(X59,X60,X60,X60) | ~cong(X60,X60,X60,X60)) )),
  inference(resolution,[],[f1883,f1030])).
fof(f1030,plain,(
  ( ! [X12,X13] : (~circle(X12,X13,X13,X13) | perp(X12,X13,X13,X13) | ~cong(X13,X13,X13,X13)) )),
  inference(resolution,[],[f408,f405])).
fof(f408,plain,(
  ( ! [X4,X2,X0,X3,X1] : (~eqangle(X0,X4,X0,X1,X2,X0,X2,X1) | perp(X3,X0,X0,X4) | ~circle(X3,X0,X1,X2)) )),
  inference(cnf_transformation,[],[f237])).
fof(f237,plain,(
  ! [X0,X1,X2,X3,X4] : (perp(X3,X0,X0,X4) | ~eqangle(X0,X4,X0,X1,X2,X0,X2,X1) | ~circle(X3,X0,X1,X2))),
  inference(flattening,[],[f236])).
fof(f236,plain,(
  ! [X0,X1,X2,X3,X4] : (perp(X3,X0,X0,X4) | (~eqangle(X0,X4,X0,X1,X2,X0,X2,X1) | ~circle(X3,X0,X1,X2)))),
  inference(ennf_transformation,[],[f129])).
fof(f129,plain,(
  ! [X0,X1,X2,X3,X4] : ((eqangle(X0,X4,X0,X1,X2,X0,X2,X1) & circle(X3,X0,X1,X2)) => perp(X3,X0,X0,X4))),
  inference(rectify,[],[f50])).
fof(f50,axiom,(
  ! [X0,X1,X2,X7,X15] : ((eqangle(X0,X15,X0,X1,X2,X0,X2,X1) & circle(X7,X0,X1,X2)) => perp(X7,X0,X0,X15))),
  file('./7_Book/00EE/04/E044-6.p',ruleD49)).
fof(f1883,plain,(
  ( ! [X2,X0,X3,X1] : (circle(X0,X1,X2,X3) | ~midp(X0,X1,X3) | ~midp(X0,X2,X1)) )),
  inference(resolution,[],[f829,f524])).
fof(f829,plain,(
  ( ! [X2,X0,X3,X1] : (~cong(X0,X1,X0,X2) | circle(X0,X1,X2,X3) | ~midp(X0,X1,X3)) )),
  inference(resolution,[],[f370,f427])).
fof(f13864,plain,(
  ~spl33_176 | ~spl33_710),
  inference(avatar_contradiction_clause,[],[f13863])).
fof(f13863,plain,(
  $false | (~spl33_176 | ~spl33_710)),
  inference(resolution,[],[f8969,f13566])).
fof(f13566,plain,(
  ( ! [X27] : (~perp(sK20,sK21,sK21,X27)) ) | ~spl33_710),
  inference(avatar_component_clause,[],[f13565])).
fof(f13124,plain,(
  ~spl33_661 | spl33_662),
  inference(avatar_split_clause,[],[f13109,f13122,f13119])).
fof(f13119,plain,(
  spl33_661 <=> ~para(sK21,sK25,sK32,sK25)),
  introduced(avatar_definition,[new_symbols(naming,[spl33_661])])).
fof(f13109,plain,(
  ( ! [X9] : (~perp(sK20,sK21,sK21,X9) | ~para(sK21,sK25,sK32,sK25) | para(sK21,X9,sK32,sK21)) )),
  inference(resolution,[],[f1804,f432])).
fof(f1804,plain,(
  ( ! [X3] : (eqangle(sK21,X3,sK32,sK21,sK21,sK25,sK32,sK25) | ~perp(sK20,sK21,sK21,X3)) )),
  inference(resolution,[],[f1022,f379])).
fof(f12212,plain,(
  ~spl33_577 | ~spl33_579),
  inference(avatar_split_clause,[],[f12171,f12210,f12204])).
fof(f12171,plain,(
  ~midp(sK27,sK20,sK24) | ~midp(sK20,sK26,sK27)),
  inference(resolution,[],[f1891,f482])).
fof(f482,plain,(
  ~cong(sK20,sK26,sK27,sK24)),
  inference(cnf_transformation,[],[f358])).
fof(f9091,plain,(
  ~spl33_176 | spl33_439),
  inference(avatar_contradiction_clause,[],[f9084])).
fof(f9084,plain,(
  $false | (~spl33_176 | ~spl33_439)),
  inference(resolution,[],[f8543,f7630])).
fof(f7630,plain,(
  ~midp(sK26,sK26,sK20) | ~spl33_439),
  inference(avatar_component_clause,[],[f7629])).
fof(f7629,plain,(
  spl33_439 <=> ~midp(sK26,sK26,sK20)),
  introduced(avatar_definition,[new_symbols(naming,[spl33_439])])).
fof(f8599,plain,(
  spl33_490 | ~spl33_306),
  inference(avatar_split_clause,[],[f8584,f3503,f8588])).
fof(f8588,plain,(
  spl33_490 <=> para(sK21,sK32,sK21,sK25)),
  introduced(avatar_definition,[new_symbols(naming,[spl33_490])])).
fof(f3503,plain,(
  spl33_306 <=> para(sK21,sK25,sK21,sK32)),
  introduced(avatar_definition,[new_symbols(naming,[spl33_306])])).
fof(f8584,plain,(
  para(sK21,sK32,sK21,sK25) | ~spl33_306),
  inference(resolution,[],[f3504,f363])).
fof(f363,plain,(
  ( ! [X2,X0,X3,X1] : (~para(X0,X1,X2,X3) | para(X2,X3,X0,X1)) )),
  inference(cnf_transformation,[],[f173])).
fof(f173,plain,(
  ! [X0,X1,X2,X3] : (para(X2,X3,X0,X1) | ~para(X0,X1,X2,X3))),
  inference(ennf_transformation,[],[f5])).
fof(f5,axiom,(
  ! [X0,X1,X2,X3] : (para(X0,X1,X2,X3) => para(X2,X3,X0,X1))),
  file('./7_Book/00EE/04/E044-6.p',ruleD5)).
fof(f3504,plain,(
  para(sK21,sK25,sK21,sK32) | ~spl33_306),
  inference(avatar_component_clause,[],[f3503])).
fof(f8596,plain,(
  ~spl33_491 | spl33_492 | ~spl33_176 | ~spl33_306),
  inference(avatar_split_clause,[],[f8586,f3503,f1994,f8594,f8591])).
fof(f8591,plain,(
  spl33_491 <=> ~para(sK21,sK32,sK21,sK25)),
  introduced(avatar_definition,[new_symbols(naming,[spl33_491])])).
fof(f8586,plain,(
  ( ! [X1] : (midp(X1,sK32,sK25) | ~para(sK21,sK32,sK21,sK25)) ) | (~spl33_176 | ~spl33_306)),
  inference(subsumption_resolution,[],[f8581,f8489])).
fof(f8581,plain,(
  ( ! [X1] : (midp(X1,sK32,sK25) | ~para(sK21,sK32,sK21,sK25) | ~midp(X1,sK21,sK21)) ) | ~spl33_306),
  inference(resolution,[],[f3504,f423])).
fof(f8577,plain,(
  ~spl33_176 | ~spl33_304),
  inference(avatar_contradiction_clause,[],[f8576])).
fof(f8576,plain,(
  $false | (~spl33_176 | ~spl33_304)),
  inference(subsumption_resolution,[],[f8575,f8489])).
fof(f8575,plain,(
  ~midp(sK27,sK20,sK20) | (~spl33_176 | ~spl33_304)),
  inference(subsumption_resolution,[],[f7347,f8508])).
fof(f8508,plain,(
  ( ! [X8] : (midp(sK23,sK25,X8)) ) | ~spl33_176),
  inference(resolution,[],[f8503,f4133])).
fof(f8503,plain,(
  ( ! [X33] : (midp(X33,sK25,sK23)) ) | ~spl33_176),
  inference(subsumption_resolution,[],[f8500,f8489])).
fof(f8500,plain,(
  ( ! [X33] : (midp(X33,sK25,sK23) | ~midp(X33,sK27,sK27)) ) | ~spl33_176),
  inference(resolution,[],[f8476,f1404])).
fof(f1404,plain,(
  ( ! [X2,X3,X1] : (~midp(sK27,X1,X2) | midp(X3,sK25,sK23) | ~midp(X3,X1,X2)) )),
  inference(subsumption_resolution,[],[f1399,f556])).
fof(f556,plain,(
  ( ! [X2,X3] : (para(X2,sK25,X3,sK23) | ~midp(sK27,X2,X3)) )),
  inference(resolution,[],[f422,f493])).
fof(f493,plain,(
  midp(sK27,sK25,sK23)),
  inference(resolution,[],[f369,f481])).
fof(f481,plain,(
  midp(sK27,sK23,sK25)),
  inference(cnf_transformation,[],[f358])).
fof(f422,plain,(
  ( ! [X4,X2,X0,X3,X1] : (~midp(X4,X2,X3) | para(X0,X2,X1,X3) | ~midp(X4,X0,X1)) )),
  inference(cnf_transformation,[],[f262])).
fof(f262,plain,(
  ! [X0,X1,X2,X3,X4] : (para(X0,X2,X1,X3) | ~midp(X4,X2,X3) | ~midp(X4,X0,X1))),
  inference(flattening,[],[f261])).
fof(f261,plain,(
  ! [X0,X1,X2,X3,X4] : (para(X0,X2,X1,X3) | (~midp(X4,X2,X3) | ~midp(X4,X0,X1)))),
  inference(ennf_transformation,[],[f142])).
fof(f142,plain,(
  ! [X0,X1,X2,X3,X4] : ((midp(X4,X2,X3) & midp(X4,X0,X1)) => para(X0,X2,X1,X3))),
  inference(rectify,[],[f64])).
fof(f64,axiom,(
  ! [X0,X1,X2,X3,X6] : ((midp(X6,X2,X3) & midp(X6,X0,X1)) => para(X0,X2,X1,X3))),
  file('./7_Book/00EE/04/E044-6.p',ruleD63)).
fof(f1399,plain,(
  ( ! [X2,X3,X1] : (~midp(sK27,X1,X2) | midp(X3,sK25,sK23) | ~para(X1,sK25,X2,sK23) | ~midp(X3,X1,X2)) )),
  inference(resolution,[],[f555,f423])).
fof(f555,plain,(
  ( ! [X0,X1] : (para(X0,sK23,X1,sK25) | ~midp(sK27,X0,X1)) )),
  inference(resolution,[],[f422,f481])).
fof(f7347,plain,(
  ~midp(sK23,sK25,sK32) | ~midp(sK27,sK20,sK20) | ~spl33_304),
  inference(resolution,[],[f3498,f1402])).
fof(f1402,plain,(
  ( ! [X12,X11] : (para(X12,sK25,X11,sK23) | ~midp(sK27,X11,X12)) )),
  inference(resolution,[],[f555,f363])).
fof(f3498,plain,(
  ( ! [X2] : (~para(sK20,sK25,sK20,X2) | ~midp(X2,sK25,sK32)) ) | ~spl33_304),
  inference(avatar_component_clause,[],[f3497])).
fof(f3497,plain,(
  spl33_304 <=> ! [X2] : (~midp(X2,sK25,sK32) | ~para(sK20,sK25,sK20,X2))),
  introduced(avatar_definition,[new_symbols(naming,[spl33_304])])).
fof(f7631,plain,(
  spl33_434 | ~spl33_439),
  inference(avatar_split_clause,[],[f7623,f7629,f7599])).
fof(f7623,plain,(
  ~midp(sK26,sK26,sK20) | perp(sK26,sK26,sK21,sK22)),
  inference(resolution,[],[f1486,f827])).
fof(f827,plain,(
  ( ! [X24,X25] : (~para(X24,X25,sK26,sK20) | perp(X24,X25,sK21,sK22)) )),
  inference(resolution,[],[f368,f479])).
fof(f368,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (~perp(X2,X3,X4,X5) | perp(X0,X1,X4,X5) | ~para(X0,X1,X2,X3)) )),
  inference(cnf_transformation,[],[f181])).
fof(f181,plain,(
  ! [X0,X1,X2,X3,X4,X5] : (perp(X0,X1,X4,X5) | ~perp(X2,X3,X4,X5) | ~para(X0,X1,X2,X3))),
  inference(flattening,[],[f180])).
fof(f180,plain,(
  ! [X0,X1,X2,X3,X4,X5] : (perp(X0,X1,X4,X5) | (~perp(X2,X3,X4,X5) | ~para(X0,X1,X2,X3)))),
  inference(ennf_transformation,[],[f10])).
fof(f10,axiom,(
  ! [X0,X1,X2,X3,X4,X5] : ((perp(X2,X3,X4,X5) & para(X0,X1,X2,X3)) => perp(X0,X1,X4,X5))),
  file('./7_Book/00EE/04/E044-6.p',ruleD10)).
fof(f3738,plain,(
  ~spl33_178),
  inference(avatar_contradiction_clause,[],[f3731])).
fof(f3731,plain,(
  $false | ~spl33_178),
  inference(resolution,[],[f1998,f1462])).
fof(f1462,plain,(
  para(sK21,sK22,sK21,sK22)),
  inference(resolution,[],[f815,f512])).
fof(f815,plain,(
  ( ! [X24,X25] : (~perp(X24,X25,sK26,sK20) | para(X24,X25,sK21,sK22)) )),
  inference(resolution,[],[f367,f479])).
fof(f1998,plain,(
  ( ! [X0,X1] : (~para(X0,X1,X0,X1)) ) | ~spl33_178),
  inference(avatar_component_clause,[],[f1997])).
fof(f1997,plain,(
  spl33_178 <=> ! [X1,X0] : ~para(X0,X1,X0,X1)),
  introduced(avatar_definition,[new_symbols(naming,[spl33_178])])).
fof(f3505,plain,(
  spl33_304 | spl33_306),
  inference(avatar_split_clause,[],[f3472,f3503,f3497])).
fof(f3472,plain,(
  ( ! [X2] : (para(sK21,sK25,sK21,sK32) | ~midp(X2,sK25,sK32) | ~para(sK20,sK25,sK20,X2)) )),
  inference(resolution,[],[f1061,f477])).
fof(f1061,plain,(
  ( ! [X14,X17,X15,X18,X16] : (~circle(X14,X17,X15,X18) | para(X17,X15,X17,X18) | ~midp(X16,X15,X18) | ~para(X14,X15,X14,X16)) )),
  inference(resolution,[],[f432,f409])).
fof(f409,plain,(
  ( ! [X4,X2,X0,X3,X1] : (eqangle(X0,X1,X0,X2,X3,X1,X3,X4) | ~midp(X4,X1,X2) | ~circle(X3,X0,X1,X2)) )),
  inference(cnf_transformation,[],[f239])).
fof(f239,plain,(
  ! [X0,X1,X2,X3,X4] : (eqangle(X0,X1,X0,X2,X3,X1,X3,X4) | ~midp(X4,X1,X2) | ~circle(X3,X0,X1,X2))),
  inference(flattening,[],[f238])).
fof(f238,plain,(
  ! [X0,X1,X2,X3,X4] : (eqangle(X0,X1,X0,X2,X3,X1,X3,X4) | (~midp(X4,X1,X2) | ~circle(X3,X0,X1,X2)))),
  inference(ennf_transformation,[],[f130])).
fof(f130,plain,(
  ! [X0,X1,X2,X3,X4] : ((midp(X4,X1,X2) & circle(X3,X0,X1,X2)) => eqangle(X0,X1,X0,X2,X3,X1,X3,X4))),
  inference(rectify,[],[f51])).
fof(f51,axiom,(
  ! [X0,X1,X2,X7,X6] : ((midp(X6,X1,X2) & circle(X7,X0,X1,X2)) => eqangle(X0,X1,X0,X2,X7,X1,X7,X6))),
  file('./7_Book/00EE/04/E044-6.p',ruleD50)).
fof(f3253,plain,(
  ~spl33_82 | ~spl33_176),
  inference(avatar_contradiction_clause,[],[f3244])).
fof(f3244,plain,(
  $false | (~spl33_82 | ~spl33_176)),
  inference(resolution,[],[f3058,f473])).
fof(f473,plain,(
  circle(sK20,sK21,sK22,sK30)),
  inference(cnf_transformation,[],[f358])).
fof(f3058,plain,(
  ( ! [X12,X10,X11,X9] : (~circle(X11,X12,X10,X9)) ) | (~spl33_82 | ~spl33_176)),
  inference(subsumption_resolution,[],[f3057,f2709])).
fof(f2709,plain,(
  ( ! [X0,X1] : (~midp(X0,X1,X0)) ) | ~spl33_82),
  inference(subsumption_resolution,[],[f2704,f369])).
fof(f2704,plain,(
  ( ! [X0,X1] : (~midp(X0,X0,X1) | ~midp(X0,X1,X0)) ) | ~spl33_82),
  inference(resolution,[],[f1082,f524])).
fof(f1082,plain,(
  ( ! [X6,X8,X7] : (~cong(X6,X7,X7,X8) | ~midp(X6,X7,X8)) ) | ~spl33_82),
  inference(avatar_component_clause,[],[f1081])).
fof(f1081,plain,(
  spl33_82 <=> ! [X7,X8,X6] : (~cong(X6,X7,X7,X8) | ~midp(X6,X7,X8))),
  introduced(avatar_definition,[new_symbols(naming,[spl33_82])])).
fof(f3057,plain,(
  ( ! [X12,X10,X11,X9] : (~circle(X11,X12,X10,X9) | midp(X9,X10,X9)) ) | ~spl33_176),
  inference(subsumption_resolution,[],[f3050,f2084])).
fof(f2084,plain,(
  ( ! [X4,X3] : (coll(X3,X4,X3)) ) | ~spl33_176),
  inference(resolution,[],[f2051,f360])).
fof(f2051,plain,(
  ( ! [X0,X1] : (coll(X0,X1,X1)) ) | ~spl33_176),
  inference(resolution,[],[f1995,f425])).
fof(f3050,plain,(
  ( ! [X12,X10,X11,X9] : (~coll(X9,X10,X9) | ~circle(X11,X12,X10,X9) | midp(X9,X10,X9)) ) | ~spl33_176),
  inference(resolution,[],[f1160,f2625])).
fof(f2625,plain,(
  ( ! [X12,X10,X11,X9] : (cyclic(X9,X10,X11,X12)) ) | ~spl33_176),
  inference(subsumption_resolution,[],[f2618,f2199])).
fof(f2199,plain,(
  ( ! [X68,X66,X67] : (cyclic(X67,X67,X66,X68)) ) | ~spl33_176),
  inference(subsumption_resolution,[],[f2198,f2166])).
fof(f2166,plain,(
  ( ! [X2,X0,X1] : (coll(X0,X1,X2)) ) | ~spl33_176),
  inference(subsumption_resolution,[],[f2163,f2124])).
fof(f2124,plain,(
  ( ! [X6,X5] : (coll(X5,X5,X6)) ) | ~spl33_176),
  inference(resolution,[],[f2084,f359])).
fof(f2163,plain,(
  ( ! [X2,X0,X1] : (coll(X0,X1,X2) | ~coll(X2,X2,X0)) ) | ~spl33_176),
  inference(resolution,[],[f2124,f361])).
fof(f2198,plain,(
  ( ! [X68,X66,X67] : (~coll(X66,X68,X67) | cyclic(X67,X67,X66,X68)) ) | ~spl33_176),
  inference(subsumption_resolution,[],[f2180,f1995])).
fof(f2618,plain,(
  ( ! [X12,X10,X11,X9] : (cyclic(X9,X10,X11,X12) | ~cyclic(X9,X9,X10,X11)) ) | ~spl33_176),
  inference(resolution,[],[f2199,f375])).
fof(f1160,plain,(
  ( ! [X6,X8,X7,X9] : (~cyclic(X7,X6,X9,X8) | ~coll(X6,X7,X6) | ~circle(X8,X9,X7,X6) | midp(X6,X7,X6)) )),
  inference(resolution,[],[f410,f399])).
fof(f2703,plain,(
  spl33_232 | spl33_82),
  inference(avatar_split_clause,[],[f2690,f1081,f2697])).
fof(f2690,plain,(
  ( ! [X14,X12,X10,X15,X13,X11] : (~midp(X10,X11,X12) | ~midp(X13,X14,X15) | ~cong(X10,X11,X11,X12) | cong(X14,X13,X14,X15)) )),
  inference(resolution,[],[f889,f434])).
fof(f434,plain,(
  ( ! [X6,X4,X2,X0,X7,X5,X3,X1] : (~eqratio(X0,X1,X2,X3,X4,X5,X6,X7) | ~cong(X4,X5,X6,X7) | cong(X0,X1,X2,X3)) )),
  inference(cnf_transformation,[],[f283])).
fof(f283,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6,X7] : (cong(X0,X1,X2,X3) | ~cong(X4,X5,X6,X7) | ~eqratio(X0,X1,X2,X3,X4,X5,X6,X7))),
  inference(flattening,[],[f282])).
fof(f282,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6,X7] : (cong(X0,X1,X2,X3) | (~cong(X4,X5,X6,X7) | ~eqratio(X0,X1,X2,X3,X4,X5,X6,X7)))),
  inference(ennf_transformation,[],[f148])).
fof(f148,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6,X7] : ((cong(X4,X5,X6,X7) & eqratio(X0,X1,X2,X3,X4,X5,X6,X7)) => cong(X0,X1,X2,X3))),
  inference(rectify,[],[f76])).
fof(f76,axiom,(
  ! [X0,X1,X2,X3,X8,X9,X10,X11] : ((cong(X8,X9,X10,X11) & eqratio(X0,X1,X2,X3,X8,X9,X10,X11)) => cong(X0,X1,X2,X3))),
  file('./7_Book/00EE/04/E044-6.p',ruleD75)).
fof(f889,plain,(
  ( ! [X6,X10,X8,X7,X11,X9] : (eqratio(X6,X7,X6,X8,X9,X10,X10,X11) | ~midp(X9,X10,X11) | ~midp(X7,X6,X8)) )),
  inference(resolution,[],[f384,f429])).
fof(f429,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (eqratio(X4,X0,X0,X1,X5,X2,X2,X3) | ~midp(X5,X2,X3) | ~midp(X4,X0,X1)) )),
  inference(cnf_transformation,[],[f273])).
fof(f273,plain,(
  ! [X0,X1,X2,X3,X4,X5] : (eqratio(X4,X0,X0,X1,X5,X2,X2,X3) | ~midp(X5,X2,X3) | ~midp(X4,X0,X1))),
  inference(flattening,[],[f272])).
fof(f272,plain,(
  ! [X0,X1,X2,X3,X4,X5] : (eqratio(X4,X0,X0,X1,X5,X2,X2,X3) | (~midp(X5,X2,X3) | ~midp(X4,X0,X1)))),
  inference(ennf_transformation,[],[f145])).
fof(f145,plain,(
  ! [X0,X1,X2,X3,X4,X5] : ((midp(X5,X2,X3) & midp(X4,X0,X1)) => eqratio(X4,X0,X0,X1,X5,X2,X2,X3))),
  inference(rectify,[],[f71])).
fof(f71,axiom,(
  ! [X0,X1,X2,X3,X6,X16] : ((midp(X16,X2,X3) & midp(X6,X0,X1)) => eqratio(X6,X0,X0,X1,X16,X2,X2,X3))),
  file('./7_Book/00EE/04/E044-6.p',ruleD70)).
fof(f384,plain,(
  ( ! [X6,X4,X2,X0,X7,X5,X3,X1] : (~eqratio(X0,X1,X2,X3,X4,X5,X6,X7) | eqratio(X1,X0,X2,X3,X4,X5,X6,X7)) )),
  inference(cnf_transformation,[],[f202])).
fof(f202,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6,X7] : (eqratio(X1,X0,X2,X3,X4,X5,X6,X7) | ~eqratio(X0,X1,X2,X3,X4,X5,X6,X7))),
  inference(ennf_transformation,[],[f105])).
fof(f105,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6,X7] : (eqratio(X0,X1,X2,X3,X4,X5,X6,X7) => eqratio(X1,X0,X2,X3,X4,X5,X6,X7))),
  inference(rectify,[],[f26])).
fof(f26,axiom,(
  ! [X0,X1,X2,X3,X8,X9,X10,X11] : (eqratio(X0,X1,X2,X3,X8,X9,X10,X11) => eqratio(X1,X0,X2,X3,X8,X9,X10,X11))),
  file('./7_Book/00EE/04/E044-6.p',ruleD26)).
fof(f2220,plain,(
  spl33_11 | ~spl33_176),
  inference(avatar_contradiction_clause,[],[f2207])).
fof(f2207,plain,(
  $false | (~spl33_11 | ~spl33_176)),
  inference(resolution,[],[f2166,f632])).
fof(f632,plain,(
  ~coll(sK20,sK21,sK32) | ~spl33_11),
  inference(avatar_component_clause,[],[f631])).
fof(f631,plain,(
  spl33_11 <=> ~coll(sK20,sK21,sK32)),
  introduced(avatar_definition,[new_symbols(naming,[spl33_11])])).
fof(f2197,plain,(
  spl33_180 | spl33_182),
  inference(avatar_split_clause,[],[f2174,f2195,f2192])).
fof(f2174,plain,(
  ( ! [X47,X45,X48,X46,X44,X49] : (~para(X44,X45,X46,X47) | ~perp(X48,X49,X48,X49) | perp(X44,X45,X46,X47)) )),
  inference(resolution,[],[f884,f433])).
fof(f433,plain,(
  ( ! [X6,X4,X2,X0,X7,X5,X3,X1] : (~eqangle(X0,X1,X2,X3,X4,X5,X6,X7) | ~perp(X4,X5,X6,X7) | perp(X0,X1,X2,X3)) )),
  inference(cnf_transformation,[],[f281])).
fof(f281,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6,X7] : (perp(X0,X1,X2,X3) | ~perp(X4,X5,X6,X7) | ~eqangle(X0,X1,X2,X3,X4,X5,X6,X7))),
  inference(flattening,[],[f280])).
fof(f280,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6,X7] : (perp(X0,X1,X2,X3) | (~perp(X4,X5,X6,X7) | ~eqangle(X0,X1,X2,X3,X4,X5,X6,X7)))),
  inference(ennf_transformation,[],[f147])).
fof(f147,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6,X7] : ((perp(X4,X5,X6,X7) & eqangle(X0,X1,X2,X3,X4,X5,X6,X7)) => perp(X0,X1,X2,X3))),
  inference(rectify,[],[f75])).
fof(f75,axiom,(
  ! [X0,X1,X2,X3,X8,X9,X10,X11] : ((perp(X8,X9,X10,X11) & eqangle(X0,X1,X2,X3,X8,X9,X10,X11)) => perp(X0,X1,X2,X3))),
  file('./7_Book/00EE/04/E044-6.p',ruleD74)).
fof(f1999,plain,(
  spl33_176 | spl33_178),
  inference(avatar_split_clause,[],[f1970,f1997,f1994])).
fof(f1970,plain,(
  ( ! [X2,X0,X3,X1] : (~para(X0,X1,X0,X1) | para(X2,X3,X2,X3)) )),
  inference(resolution,[],[f858,f397])).
fof(f633,plain,(
  spl33_8 | ~spl33_11),
  inference(avatar_split_clause,[],[f593,f631,f625])).
fof(f593,plain,(
  ~coll(sK20,sK21,sK32) | perp(sK21,sK25,sK25,sK32)),
  inference(resolution,[],[f412,f477])).
% SZS output end Proof for E044-6
% ------------------------------
% Version: Vampire 4.4.0 (commit 7916d27 on 2019-08-23 08:50:16 +0100)
% Termination reason: Refutation

% Memory used [KB]: 15351
% Time elapsed: 0.347 s
% ------------------------------
% ------------------------------
