% Refutation found. Thanks to Tanya!
% SZS status Theorem for E051-5
% SZS output start Proof for E051-5
fof(f24011,plain,(
  $false),
  inference(avatar_sat_refutation,[],[f2155,f4902,f5857,f11224,f12061,f12094,f23977])).
fof(f23977,plain,(
  ~spl29_114 | ~spl29_180),
  inference(avatar_contradiction_clause,[],[f23971])).
fof(f23971,plain,(
  $false | (~spl29_114 | ~spl29_180)),
  inference(resolution,[],[f23862,f7660])).
fof(f7660,plain,(
  ~midp(sK27,sK28,sK23)),
  inference(resolution,[],[f1797,f485])).
fof(f485,plain,(
  ~cong(sK23,sK27,sK27,sK28)),
  inference(cnf_transformation,[],[f359])).
fof(f359,plain,(
  ~cong(sK23,sK27,sK27,sK28) & coll(sK28,sK23,sK24) & coll(sK27,sK23,sK24) & para(sK24,sK25,sK28,sK26) & para(sK23,sK25,sK27,sK26) & perp(sK23,sK24,sK20,sK26) & midp(sK20,sK23,sK24) & perp(sK25,sK24,sK21,sK26) & midp(sK21,sK25,sK24) & perp(sK25,sK23,sK22,sK26) & midp(sK22,sK25,sK23) & eqangle(sK24,sK23,sK23,sK25,sK23,sK25,sK25,sK24) & eqangle(sK24,sK23,sK23,sK25,sK23,sK24,sK24,sK25)),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK20,sK21,sK22,sK23,sK24,sK25,sK26,sK27,sK28])],[f321,f358])).
fof(f358,plain,(
  ? [X0,X1,X2,X3,X4,X5,X6,X7,X8] : (~cong(X3,X7,X7,X8) & coll(X8,X3,X4) & coll(X7,X3,X4) & para(X4,X5,X8,X6) & para(X3,X5,X7,X6) & perp(X3,X4,X0,X6) & midp(X0,X3,X4) & perp(X5,X4,X1,X6) & midp(X1,X5,X4) & perp(X5,X3,X2,X6) & midp(X2,X5,X3) & eqangle(X4,X3,X3,X5,X3,X5,X5,X4) & eqangle(X4,X3,X3,X5,X3,X4,X4,X5)) => (~cong(sK23,sK27,sK27,sK28) & coll(sK28,sK23,sK24) & coll(sK27,sK23,sK24) & para(sK24,sK25,sK28,sK26) & para(sK23,sK25,sK27,sK26) & perp(sK23,sK24,sK20,sK26) & midp(sK20,sK23,sK24) & perp(sK25,sK24,sK21,sK26) & midp(sK21,sK25,sK24) & perp(sK25,sK23,sK22,sK26) & midp(sK22,sK25,sK23) & eqangle(sK24,sK23,sK23,sK25,sK23,sK25,sK25,sK24) & eqangle(sK24,sK23,sK23,sK25,sK23,sK24,sK24,sK25))),
  introduced(choice_axiom,[])).
fof(f321,plain,(
  ? [X0,X1,X2,X3,X4,X5,X6,X7,X8] : (~cong(X3,X7,X7,X8) & coll(X8,X3,X4) & coll(X7,X3,X4) & para(X4,X5,X8,X6) & para(X3,X5,X7,X6) & perp(X3,X4,X0,X6) & midp(X0,X3,X4) & perp(X5,X4,X1,X6) & midp(X1,X5,X4) & perp(X5,X3,X2,X6) & midp(X2,X5,X3) & eqangle(X4,X3,X3,X5,X3,X5,X5,X4) & eqangle(X4,X3,X3,X5,X3,X4,X4,X5))),
  inference(flattening,[],[f320])).
fof(f320,plain,(
  ? [X0,X1,X2,X3,X4,X5,X6,X7,X8] : (~cong(X3,X7,X7,X8) & (coll(X8,X3,X4) & coll(X7,X3,X4) & para(X4,X5,X8,X6) & para(X3,X5,X7,X6) & perp(X3,X4,X0,X6) & midp(X0,X3,X4) & perp(X5,X4,X1,X6) & midp(X1,X5,X4) & perp(X5,X3,X2,X6) & midp(X2,X5,X3) & eqangle(X4,X3,X3,X5,X3,X5,X5,X4) & eqangle(X4,X3,X3,X5,X3,X4,X4,X5)))),
  inference(ennf_transformation,[],[f168])).
fof(f168,plain,(
  ~! [X0,X1,X2,X3,X4,X5,X6,X7,X8] : ((coll(X8,X3,X4) & coll(X7,X3,X4) & para(X4,X5,X8,X6) & para(X3,X5,X7,X6) & perp(X3,X4,X0,X6) & midp(X0,X3,X4) & perp(X5,X4,X1,X6) & midp(X1,X5,X4) & perp(X5,X3,X2,X6) & midp(X2,X5,X3) & eqangle(X4,X3,X3,X5,X3,X5,X5,X4) & eqangle(X4,X3,X3,X5,X3,X4,X4,X5)) => cong(X3,X7,X7,X8))),
  inference(flattening,[],[f167])).
fof(f167,plain,(
  ~! [X0,X1,X2] : ! [X3,X4,X5,X6,X7,X8] : ((coll(X8,X3,X4) & coll(X7,X3,X4) & para(X4,X5,X8,X6) & para(X3,X5,X7,X6) & perp(X3,X4,X0,X6) & midp(X0,X3,X4) & perp(X5,X4,X1,X6) & midp(X1,X5,X4) & perp(X5,X3,X2,X6) & midp(X2,X5,X3) & eqangle(X4,X3,X3,X5,X3,X5,X5,X4) & eqangle(X4,X3,X3,X5,X3,X4,X4,X5)) => cong(X3,X7,X7,X8))),
  inference(rectify,[],[f96])).
fof(f96,negated_conjecture,(
  ~! [X20,X19,X18] : ! [X0,X1,X2,X3,X4,X5] : ((coll(X5,X0,X1) & coll(X4,X0,X1) & para(X1,X2,X5,X3) & para(X0,X2,X4,X3) & perp(X0,X1,X20,X3) & midp(X20,X0,X1) & perp(X2,X1,X19,X3) & midp(X19,X2,X1) & perp(X2,X0,X18,X3) & midp(X18,X2,X0) & eqangle(X1,X0,X0,X2,X0,X2,X2,X1) & eqangle(X1,X0,X0,X2,X0,X1,X1,X2)) => cong(X0,X4,X4,X5))),
  inference(negated_conjecture,[],[f95])).
fof(f95,conjecture,(
  ( ! [X19,X20,X18] : (! [X0,X1,X2,X3,X4,X5] : ((coll(X5,X0,X1) & coll(X4,X0,X1) & para(X1,X2,X5,X3) & para(X0,X2,X4,X3) & perp(X0,X1,X20,X3) & midp(X20,X0,X1) & perp(X2,X1,X19,X3) & midp(X19,X2,X1) & perp(X2,X0,X18,X3) & midp(X18,X2,X0) & eqangle(X1,X0,X0,X2,X0,X2,X2,X1) & eqangle(X1,X0,X0,X2,X0,X1,X1,X2)) => cong(X0,X4,X4,X5))) )),
  file('./7_Book/00EE/04/E051-5.p',exemplo7Book00EE04E0515)).
fof(f1797,plain,(
  ( ! [X6,X7,X5] : (cong(X7,X5,X5,X6) | ~midp(X5,X6,X7)) )),
  inference(resolution,[],[f536,f383])).
fof(f383,plain,(
  ( ! [X2,X0,X3,X1] : (~cong(X0,X1,X2,X3) | cong(X2,X3,X0,X1)) )),
  inference(cnf_transformation,[],[f200])).
fof(f200,plain,(
  ! [X0,X1,X2,X3] : (cong(X2,X3,X0,X1) | ~cong(X0,X1,X2,X3))),
  inference(ennf_transformation,[],[f24])).
fof(f24,axiom,(
  ! [X0,X1,X2,X3] : (cong(X0,X1,X2,X3) => cong(X2,X3,X0,X1))),
  file('./7_Book/00EE/04/E051-5.p',ruleD24)).
fof(f536,plain,(
  ( ! [X2,X0,X1] : (cong(X0,X1,X2,X0) | ~midp(X0,X1,X2)) )),
  inference(resolution,[],[f382,f428])).
fof(f428,plain,(
  ( ! [X2,X0,X1] : (cong(X0,X1,X0,X2) | ~midp(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f271])).
fof(f271,plain,(
  ! [X0,X1,X2] : (cong(X0,X1,X0,X2) | ~midp(X0,X1,X2))),
  inference(ennf_transformation,[],[f69])).
fof(f69,axiom,(
  ! [X0,X1,X2] : (midp(X0,X1,X2) => cong(X0,X1,X0,X2))),
  file('./7_Book/00EE/04/E051-5.p',ruleD68)).
fof(f382,plain,(
  ( ! [X2,X0,X3,X1] : (~cong(X0,X1,X2,X3) | cong(X0,X1,X3,X2)) )),
  inference(cnf_transformation,[],[f199])).
fof(f199,plain,(
  ! [X0,X1,X2,X3] : (cong(X0,X1,X3,X2) | ~cong(X0,X1,X2,X3))),
  inference(ennf_transformation,[],[f23])).
fof(f23,axiom,(
  ! [X0,X1,X2,X3] : (cong(X0,X1,X2,X3) => cong(X0,X1,X3,X2))),
  file('./7_Book/00EE/04/E051-5.p',ruleD23)).
fof(f23862,plain,(
  ( ! [X50,X48,X49] : (midp(X48,X49,X50)) ) | (~spl29_114 | ~spl29_180)),
  inference(subsumption_resolution,[],[f23818,f5990])).
fof(f5990,plain,(
  ( ! [X2,X0,X1] : (coll(X0,X1,X2)) ) | ~spl29_180),
  inference(subsumption_resolution,[],[f5987,f5958])).
fof(f5958,plain,(
  ( ! [X6,X5] : (coll(X5,X5,X6)) ) | ~spl29_180),
  inference(resolution,[],[f5930,f360])).
fof(f360,plain,(
  ( ! [X2,X0,X1] : (~coll(X0,X1,X2) | coll(X0,X2,X1)) )),
  inference(cnf_transformation,[],[f169])).
fof(f169,plain,(
  ! [X0,X1,X2] : (coll(X0,X2,X1) | ~coll(X0,X1,X2))),
  inference(ennf_transformation,[],[f1])).
fof(f1,axiom,(
  ! [X0,X1,X2] : (coll(X0,X1,X2) => coll(X0,X2,X1))),
  file('./7_Book/00EE/04/E051-5.p',ruleD1)).
fof(f5930,plain,(
  ( ! [X4,X3] : (coll(X3,X4,X3)) ) | ~spl29_180),
  inference(resolution,[],[f5893,f361])).
fof(f361,plain,(
  ( ! [X2,X0,X1] : (~coll(X0,X1,X2) | coll(X1,X0,X2)) )),
  inference(cnf_transformation,[],[f170])).
fof(f170,plain,(
  ! [X0,X1,X2] : (coll(X1,X0,X2) | ~coll(X0,X1,X2))),
  inference(ennf_transformation,[],[f2])).
fof(f2,axiom,(
  ! [X0,X1,X2] : (coll(X0,X1,X2) => coll(X1,X0,X2))),
  file('./7_Book/00EE/04/E051-5.p',ruleD2)).
fof(f5893,plain,(
  ( ! [X4,X3] : (coll(X3,X4,X4)) ) | ~spl29_180),
  inference(resolution,[],[f2151,f426])).
fof(f426,plain,(
  ( ! [X2,X0,X1] : (~para(X0,X1,X0,X2) | coll(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f268])).
fof(f268,plain,(
  ! [X0,X1,X2] : (coll(X0,X1,X2) | ~para(X0,X1,X0,X2))),
  inference(ennf_transformation,[],[f67])).
fof(f67,axiom,(
  ! [X0,X1,X2] : (para(X0,X1,X0,X2) => coll(X0,X1,X2))),
  file('./7_Book/00EE/04/E051-5.p',ruleD66)).
fof(f2151,plain,(
  ( ! [X2,X3] : (para(X2,X3,X2,X3)) ) | ~spl29_180),
  inference(avatar_component_clause,[],[f2150])).
fof(f2150,plain,(
  spl29_180 <=> ! [X3,X2] : para(X2,X3,X2,X3)),
  introduced(avatar_definition,[new_symbols(naming,[spl29_180])])).
fof(f5987,plain,(
  ( ! [X2,X0,X1] : (coll(X0,X1,X2) | ~coll(X2,X2,X0)) ) | ~spl29_180),
  inference(resolution,[],[f5958,f362])).
fof(f362,plain,(
  ( ! [X2,X0,X3,X1] : (~coll(X0,X1,X3) | coll(X2,X3,X0) | ~coll(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f172])).
fof(f172,plain,(
  ! [X0,X1,X2,X3] : (coll(X2,X3,X0) | ~coll(X0,X1,X3) | ~coll(X0,X1,X2))),
  inference(flattening,[],[f171])).
fof(f171,plain,(
  ! [X0,X1,X2,X3] : (coll(X2,X3,X0) | (~coll(X0,X1,X3) | ~coll(X0,X1,X2)))),
  inference(ennf_transformation,[],[f3])).
fof(f3,axiom,(
  ! [X0,X1,X2,X3] : ((coll(X0,X1,X3) & coll(X0,X1,X2)) => coll(X2,X3,X0))),
  file('./7_Book/00EE/04/E051-5.p',ruleD3)).
fof(f23818,plain,(
  ( ! [X50,X48,X49] : (~coll(X48,X49,X50) | midp(X48,X49,X50)) ) | (~spl29_114 | ~spl29_180)),
  inference(resolution,[],[f20365,f427])).
fof(f427,plain,(
  ( ! [X2,X0,X1] : (~cong(X0,X1,X0,X2) | ~coll(X0,X1,X2) | midp(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f270])).
fof(f270,plain,(
  ! [X0,X1,X2] : (midp(X0,X1,X2) | ~coll(X0,X1,X2) | ~cong(X0,X1,X0,X2))),
  inference(flattening,[],[f269])).
fof(f269,plain,(
  ! [X0,X1,X2] : (midp(X0,X1,X2) | (~coll(X0,X1,X2) | ~cong(X0,X1,X0,X2)))),
  inference(ennf_transformation,[],[f68])).
fof(f68,axiom,(
  ! [X0,X1,X2] : ((coll(X0,X1,X2) & cong(X0,X1,X0,X2)) => midp(X0,X1,X2))),
  file('./7_Book/00EE/04/E051-5.p',ruleD67)).
fof(f20365,plain,(
  ( ! [X14,X12,X13] : (cong(X12,X13,X12,X14)) ) | (~spl29_114 | ~spl29_180)),
  inference(subsumption_resolution,[],[f20266,f20261])).
fof(f20261,plain,(
  ( ! [X0,X1] : (midp(X0,X0,X1)) ) | (~spl29_114 | ~spl29_180)),
  inference(subsumption_resolution,[],[f20260,f15162])).
fof(f15162,plain,(
  ( ! [X30,X31,X29,X32] : (perp(X29,X30,X31,X32)) ) | (~spl29_114 | ~spl29_180)),
  inference(subsumption_resolution,[],[f15139,f15161])).
fof(f15161,plain,(
  ( ! [X28,X26,X27,X25] : (para(X25,X26,X27,X28)) ) | (~spl29_114 | ~spl29_180)),
  inference(subsumption_resolution,[],[f15138,f15137])).
fof(f15137,plain,(
  ( ! [X24,X23] : (perp(X23,X24,sK23,sK25)) ) | (~spl29_114 | ~spl29_180)),
  inference(resolution,[],[f12328,f367])).
fof(f367,plain,(
  ( ! [X2,X0,X3,X1] : (~perp(X0,X1,X2,X3) | perp(X2,X3,X0,X1)) )),
  inference(cnf_transformation,[],[f178])).
fof(f178,plain,(
  ! [X0,X1,X2,X3] : (perp(X2,X3,X0,X1) | ~perp(X0,X1,X2,X3))),
  inference(ennf_transformation,[],[f8])).
fof(f8,axiom,(
  ! [X0,X1,X2,X3] : (perp(X0,X1,X2,X3) => perp(X2,X3,X0,X1))),
  file('./7_Book/00EE/04/E051-5.p',ruleD8)).
fof(f12328,plain,(
  ( ! [X12,X13] : (perp(sK23,sK25,X12,X13)) ) | (~spl29_114 | ~spl29_180)),
  inference(subsumption_resolution,[],[f12325,f12215])).
fof(f12215,plain,(
  ( ! [X6] : (cong(sK23,X6,sK25,X6)) ) | (~spl29_114 | ~spl29_180)),
  inference(subsumption_resolution,[],[f12208,f11106])).
fof(f11106,plain,(
  ( ! [X0] : (midp(X0,sK23,sK24)) ) | ~spl29_180),
  inference(resolution,[],[f10956,f370])).
fof(f370,plain,(
  ( ! [X2,X0,X1] : (~midp(X2,X1,X0) | midp(X2,X0,X1)) )),
  inference(cnf_transformation,[],[f183])).
fof(f183,plain,(
  ! [X0,X1,X2] : (midp(X2,X0,X1) | ~midp(X2,X1,X0))),
  inference(ennf_transformation,[],[f97])).
fof(f97,plain,(
  ! [X0,X1,X2] : (midp(X2,X1,X0) => midp(X2,X0,X1))),
  inference(rectify,[],[f11])).
fof(f11,axiom,(
  ! [X0,X1,X6] : (midp(X6,X1,X0) => midp(X6,X0,X1))),
  file('./7_Book/00EE/04/E051-5.p',ruleD11)).
fof(f10956,plain,(
  ( ! [X87] : (midp(X87,sK24,sK23)) ) | ~spl29_180),
  inference(subsumption_resolution,[],[f10920,f10843])).
fof(f10843,plain,(
  ( ! [X30,X29] : (midp(X29,X30,X30)) ) | ~spl29_180),
  inference(subsumption_resolution,[],[f10813,f5893])).
fof(f10813,plain,(
  ( ! [X30,X29] : (~coll(X29,X30,X30) | midp(X29,X30,X30)) ) | ~spl29_180),
  inference(resolution,[],[f10802,f427])).
fof(f10802,plain,(
  ( ! [X85,X86] : (cong(X85,X86,X85,X86)) ) | ~spl29_180),
  inference(subsumption_resolution,[],[f10801,f9231])).
fof(f9231,plain,(
  ( ! [X12,X10,X11,X9] : (cyclic(X9,X10,X11,X12)) ) | ~spl29_180),
  inference(subsumption_resolution,[],[f9209,f8701])).
fof(f8701,plain,(
  ( ! [X6,X4,X5] : (cyclic(X4,X4,X5,X6)) ) | ~spl29_180),
  inference(resolution,[],[f4268,f2151])).
fof(f4268,plain,(
  ( ! [X68,X66,X67] : (~para(X66,X67,X66,X67) | cyclic(X67,X67,X66,X68)) )),
  inference(subsumption_resolution,[],[f2352,f2353])).
fof(f2353,plain,(
  ( ! [X70,X71,X69] : (~para(X69,X70,X69,X70) | coll(X69,X71,X70) | cyclic(X70,X70,X69,X71)) )),
  inference(resolution,[],[f829,f401])).
fof(f401,plain,(
  ( ! [X2,X0,X3,X1] : (~eqangle(X2,X0,X2,X1,X3,X0,X3,X1) | coll(X2,X3,X0) | cyclic(X0,X1,X2,X3)) )),
  inference(cnf_transformation,[],[f223])).
fof(f223,plain,(
  ! [X0,X1,X2,X3] : (cyclic(X0,X1,X2,X3) | coll(X2,X3,X0) | ~eqangle(X2,X0,X2,X1,X3,X0,X3,X1))),
  inference(flattening,[],[f222])).
fof(f222,plain,(
  ! [X0,X1,X2,X3] : (cyclic(X0,X1,X2,X3) | (coll(X2,X3,X0) | ~eqangle(X2,X0,X2,X1,X3,X0,X3,X1)))),
  inference(ennf_transformation,[],[f121])).
fof(f121,plain,(
  ! [X0,X1,X2,X3] : ((~coll(X2,X3,X0) & eqangle(X2,X0,X2,X1,X3,X0,X3,X1)) => cyclic(X0,X1,X2,X3))),
  inference(rectify,[],[f42])).
fof(f42,axiom,(
  ! [X0,X1,X8,X9] : ((~coll(X8,X9,X0) & eqangle(X8,X0,X8,X1,X9,X0,X9,X1)) => cyclic(X0,X1,X8,X9))),
  file('./7_Book/00EE/04/E051-5.p',ruleD42a)).
fof(f829,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (eqangle(X0,X1,X2,X3,X4,X5,X4,X5) | ~para(X0,X1,X2,X3)) )),
  inference(resolution,[],[f380,f399])).
fof(f399,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (eqangle(X0,X1,X4,X5,X2,X3,X4,X5) | ~para(X0,X1,X2,X3)) )),
  inference(cnf_transformation,[],[f220])).
fof(f220,plain,(
  ! [X0,X1,X2,X3,X4,X5] : (eqangle(X0,X1,X4,X5,X2,X3,X4,X5) | ~para(X0,X1,X2,X3))),
  inference(ennf_transformation,[],[f119])).
fof(f119,plain,(
  ! [X0,X1,X2,X3,X4,X5] : (para(X0,X1,X2,X3) => eqangle(X0,X1,X4,X5,X2,X3,X4,X5))),
  inference(rectify,[],[f40])).
fof(f40,axiom,(
  ! [X0,X1,X2,X3,X8,X9] : (para(X0,X1,X2,X3) => eqangle(X0,X1,X8,X9,X2,X3,X8,X9))),
  file('./7_Book/00EE/04/E051-5.p',ruleD40)).
fof(f380,plain,(
  ( ! [X6,X4,X2,X0,X7,X5,X3,X1] : (~eqangle(X0,X1,X2,X3,X4,X5,X6,X7) | eqangle(X0,X1,X4,X5,X2,X3,X6,X7)) )),
  inference(cnf_transformation,[],[f196])).
fof(f196,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6,X7] : (eqangle(X0,X1,X4,X5,X2,X3,X6,X7) | ~eqangle(X0,X1,X2,X3,X4,X5,X6,X7))),
  inference(ennf_transformation,[],[f103])).
fof(f103,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6,X7] : (eqangle(X0,X1,X2,X3,X4,X5,X6,X7) => eqangle(X0,X1,X4,X5,X2,X3,X6,X7))),
  inference(rectify,[],[f21])).
fof(f21,axiom,(
  ! [X0,X1,X2,X3,X8,X9,X10,X11] : (eqangle(X0,X1,X2,X3,X8,X9,X10,X11) => eqangle(X0,X1,X8,X9,X2,X3,X10,X11))),
  file('./7_Book/00EE/04/E051-5.p',ruleD21)).
fof(f2352,plain,(
  ( ! [X68,X66,X67] : (~para(X66,X67,X66,X67) | ~coll(X66,X68,X67) | cyclic(X67,X67,X66,X68)) )),
  inference(resolution,[],[f829,f402])).
fof(f402,plain,(
  ( ! [X2,X0,X3,X1] : (~eqangle(X2,X0,X2,X1,X3,X0,X3,X1) | ~coll(X2,X3,X1) | cyclic(X0,X1,X2,X3)) )),
  inference(cnf_transformation,[],[f225])).
fof(f225,plain,(
  ! [X0,X1,X2,X3] : (cyclic(X0,X1,X2,X3) | ~coll(X2,X3,X1) | ~eqangle(X2,X0,X2,X1,X3,X0,X3,X1))),
  inference(flattening,[],[f224])).
fof(f224,plain,(
  ! [X0,X1,X2,X3] : (cyclic(X0,X1,X2,X3) | (~coll(X2,X3,X1) | ~eqangle(X2,X0,X2,X1,X3,X0,X3,X1)))),
  inference(ennf_transformation,[],[f122])).
fof(f122,plain,(
  ! [X0,X1,X2,X3] : ((coll(X2,X3,X1) & eqangle(X2,X0,X2,X1,X3,X0,X3,X1)) => cyclic(X0,X1,X2,X3))),
  inference(rectify,[],[f43])).
fof(f43,axiom,(
  ! [X0,X1,X8,X9] : ((coll(X8,X9,X1) & eqangle(X8,X0,X8,X1,X9,X0,X9,X1)) => cyclic(X0,X1,X8,X9))),
  file('./7_Book/00EE/04/E051-5.p',ruleD42b)).
fof(f9209,plain,(
  ( ! [X12,X10,X11,X9] : (cyclic(X9,X10,X11,X12) | ~cyclic(X9,X9,X10,X11)) ) | ~spl29_180),
  inference(resolution,[],[f8701,f376])).
fof(f376,plain,(
  ( ! [X4,X2,X0,X3,X1] : (~cyclic(X0,X1,X2,X4) | cyclic(X1,X2,X3,X4) | ~cyclic(X0,X1,X2,X3)) )),
  inference(cnf_transformation,[],[f192])).
fof(f192,plain,(
  ! [X0,X1,X2,X3,X4] : (cyclic(X1,X2,X3,X4) | ~cyclic(X0,X1,X2,X4) | ~cyclic(X0,X1,X2,X3))),
  inference(flattening,[],[f191])).
fof(f191,plain,(
  ! [X0,X1,X2,X3,X4] : (cyclic(X1,X2,X3,X4) | (~cyclic(X0,X1,X2,X4) | ~cyclic(X0,X1,X2,X3)))),
  inference(ennf_transformation,[],[f17])).
fof(f17,axiom,(
  ! [X0,X1,X2,X3,X4] : ((cyclic(X0,X1,X2,X4) & cyclic(X0,X1,X2,X3)) => cyclic(X1,X2,X3,X4))),
  file('./7_Book/00EE/04/E051-5.p',ruleD17)).
fof(f10801,plain,(
  ( ! [X87,X85,X86] : (cong(X85,X86,X85,X86) | ~cyclic(X85,X86,X87,X86)) ) | ~spl29_180),
  inference(subsumption_resolution,[],[f10800,f9231])).
fof(f10800,plain,(
  ( ! [X87,X85,X86] : (cong(X85,X86,X85,X86) | ~cyclic(X85,X86,X87,X85) | ~cyclic(X85,X86,X87,X86)) ) | ~spl29_180),
  inference(subsumption_resolution,[],[f10787,f9207])).
fof(f9207,plain,(
  ( ! [X4,X5,X3] : (cyclic(X3,X4,X3,X5)) ) | ~spl29_180),
  inference(resolution,[],[f8701,f374])).
fof(f374,plain,(
  ( ! [X2,X0,X3,X1] : (~cyclic(X0,X1,X2,X3) | cyclic(X0,X2,X1,X3)) )),
  inference(cnf_transformation,[],[f189])).
fof(f189,plain,(
  ! [X0,X1,X2,X3] : (cyclic(X0,X2,X1,X3) | ~cyclic(X0,X1,X2,X3))),
  inference(ennf_transformation,[],[f15])).
fof(f15,axiom,(
  ! [X0,X1,X2,X3] : (cyclic(X0,X1,X2,X3) => cyclic(X0,X2,X1,X3))),
  file('./7_Book/00EE/04/E051-5.p',ruleD15)).
fof(f10787,plain,(
  ( ! [X87,X85,X86] : (~cyclic(X85,X86,X85,X87) | cong(X85,X86,X85,X86) | ~cyclic(X85,X86,X87,X85) | ~cyclic(X85,X86,X87,X86)) )),
  inference(duplicate_literal_removal,[],[f10777])).
fof(f10777,plain,(
  ( ! [X87,X85,X86] : (~cyclic(X85,X86,X85,X87) | cong(X85,X86,X85,X86) | ~cyclic(X85,X86,X87,X85) | ~cyclic(X85,X86,X87,X86) | ~cyclic(X85,X86,X87,X85)) )),
  inference(resolution,[],[f2079,f403])).
fof(f403,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (~eqangle(X2,X0,X2,X1,X5,X3,X5,X4) | cong(X0,X1,X3,X4) | ~cyclic(X0,X1,X2,X5) | ~cyclic(X0,X1,X2,X4) | ~cyclic(X0,X1,X2,X3)) )),
  inference(cnf_transformation,[],[f227])).
fof(f227,plain,(
  ! [X0,X1,X2,X3,X4,X5] : (cong(X0,X1,X3,X4) | ~eqangle(X2,X0,X2,X1,X5,X3,X5,X4) | ~cyclic(X0,X1,X2,X5) | ~cyclic(X0,X1,X2,X4) | ~cyclic(X0,X1,X2,X3))),
  inference(flattening,[],[f226])).
fof(f226,plain,(
  ! [X0,X1,X2,X3,X4,X5] : (cong(X0,X1,X3,X4) | (~eqangle(X2,X0,X2,X1,X5,X3,X5,X4) | ~cyclic(X0,X1,X2,X5) | ~cyclic(X0,X1,X2,X4) | ~cyclic(X0,X1,X2,X3)))),
  inference(ennf_transformation,[],[f123])).
fof(f123,plain,(
  ! [X0,X1,X2,X3,X4,X5] : ((eqangle(X2,X0,X2,X1,X5,X3,X5,X4) & cyclic(X0,X1,X2,X5) & cyclic(X0,X1,X2,X4) & cyclic(X0,X1,X2,X3)) => cong(X0,X1,X3,X4))),
  inference(rectify,[],[f44])).
fof(f44,axiom,(
  ! [X0,X1,X2,X8,X9,X14] : ((eqangle(X2,X0,X2,X1,X14,X8,X14,X9) & cyclic(X0,X1,X2,X14) & cyclic(X0,X1,X2,X9) & cyclic(X0,X1,X2,X8)) => cong(X0,X1,X8,X9))),
  file('./7_Book/00EE/04/E051-5.p',ruleD43)).
fof(f2079,plain,(
  ( ! [X14,X12,X13,X11] : (eqangle(X14,X11,X14,X12,X11,X13,X13,X12) | ~cyclic(X11,X12,X13,X14)) )),
  inference(resolution,[],[f790,f379])).
fof(f379,plain,(
  ( ! [X6,X4,X2,X0,X7,X5,X3,X1] : (~eqangle(X0,X1,X2,X3,X4,X5,X6,X7) | eqangle(X4,X5,X6,X7,X0,X1,X2,X3)) )),
  inference(cnf_transformation,[],[f195])).
fof(f195,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6,X7] : (eqangle(X4,X5,X6,X7,X0,X1,X2,X3) | ~eqangle(X0,X1,X2,X3,X4,X5,X6,X7))),
  inference(ennf_transformation,[],[f102])).
fof(f102,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6,X7] : (eqangle(X0,X1,X2,X3,X4,X5,X6,X7) => eqangle(X4,X5,X6,X7,X0,X1,X2,X3))),
  inference(rectify,[],[f20])).
fof(f20,axiom,(
  ! [X0,X1,X2,X3,X8,X9,X10,X11] : (eqangle(X0,X1,X2,X3,X8,X9,X10,X11) => eqangle(X8,X9,X10,X11,X0,X1,X2,X3))),
  file('./7_Book/00EE/04/E051-5.p',ruleD20)).
fof(f790,plain,(
  ( ! [X6,X8,X7,X9] : (eqangle(X6,X7,X7,X8,X9,X6,X9,X8) | ~cyclic(X6,X8,X7,X9)) )),
  inference(resolution,[],[f377,f400])).
fof(f400,plain,(
  ( ! [X2,X0,X3,X1] : (eqangle(X2,X0,X2,X1,X3,X0,X3,X1) | ~cyclic(X0,X1,X2,X3)) )),
  inference(cnf_transformation,[],[f221])).
fof(f221,plain,(
  ! [X0,X1,X2,X3] : (eqangle(X2,X0,X2,X1,X3,X0,X3,X1) | ~cyclic(X0,X1,X2,X3))),
  inference(ennf_transformation,[],[f120])).
fof(f120,plain,(
  ! [X0,X1,X2,X3] : (cyclic(X0,X1,X2,X3) => eqangle(X2,X0,X2,X1,X3,X0,X3,X1))),
  inference(rectify,[],[f41])).
fof(f41,axiom,(
  ! [X0,X1,X8,X9] : (cyclic(X0,X1,X8,X9) => eqangle(X8,X0,X8,X1,X9,X0,X9,X1))),
  file('./7_Book/00EE/04/E051-5.p',ruleD41)).
fof(f377,plain,(
  ( ! [X6,X4,X2,X0,X7,X5,X3,X1] : (~eqangle(X0,X1,X2,X3,X4,X5,X6,X7) | eqangle(X1,X0,X2,X3,X4,X5,X6,X7)) )),
  inference(cnf_transformation,[],[f193])).
fof(f193,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6,X7] : (eqangle(X1,X0,X2,X3,X4,X5,X6,X7) | ~eqangle(X0,X1,X2,X3,X4,X5,X6,X7))),
  inference(ennf_transformation,[],[f100])).
fof(f100,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6,X7] : (eqangle(X0,X1,X2,X3,X4,X5,X6,X7) => eqangle(X1,X0,X2,X3,X4,X5,X6,X7))),
  inference(rectify,[],[f18])).
fof(f18,axiom,(
  ! [X0,X1,X2,X3,X8,X9,X10,X11] : (eqangle(X0,X1,X2,X3,X8,X9,X10,X11) => eqangle(X1,X0,X2,X3,X8,X9,X10,X11))),
  file('./7_Book/00EE/04/E051-5.p',ruleD18)).
fof(f10920,plain,(
  ( ! [X88,X87] : (midp(X87,sK24,sK23) | ~midp(X87,X88,X88)) ) | ~spl29_180),
  inference(resolution,[],[f10843,f1661])).
fof(f1661,plain,(
  ( ! [X2,X3,X1] : (~midp(sK20,X1,X2) | midp(X3,sK24,sK23) | ~midp(X3,X1,X2)) )),
  inference(subsumption_resolution,[],[f1656,f578])).
fof(f578,plain,(
  ( ! [X10,X11] : (para(X10,sK24,X11,sK23) | ~midp(sK20,X10,X11)) )),
  inference(resolution,[],[f423,f496])).
fof(f496,plain,(
  midp(sK20,sK24,sK23)),
  inference(resolution,[],[f370,f479])).
fof(f479,plain,(
  midp(sK20,sK23,sK24)),
  inference(cnf_transformation,[],[f359])).
fof(f423,plain,(
  ( ! [X4,X2,X0,X3,X1] : (~midp(X4,X2,X3) | para(X0,X2,X1,X3) | ~midp(X4,X0,X1)) )),
  inference(cnf_transformation,[],[f263])).
fof(f263,plain,(
  ! [X0,X1,X2,X3,X4] : (para(X0,X2,X1,X3) | ~midp(X4,X2,X3) | ~midp(X4,X0,X1))),
  inference(flattening,[],[f262])).
fof(f262,plain,(
  ! [X0,X1,X2,X3,X4] : (para(X0,X2,X1,X3) | (~midp(X4,X2,X3) | ~midp(X4,X0,X1)))),
  inference(ennf_transformation,[],[f142])).
fof(f142,plain,(
  ! [X0,X1,X2,X3,X4] : ((midp(X4,X2,X3) & midp(X4,X0,X1)) => para(X0,X2,X1,X3))),
  inference(rectify,[],[f64])).
fof(f64,axiom,(
  ! [X0,X1,X2,X3,X6] : ((midp(X6,X2,X3) & midp(X6,X0,X1)) => para(X0,X2,X1,X3))),
  file('./7_Book/00EE/04/E051-5.p',ruleD63)).
fof(f1656,plain,(
  ( ! [X2,X3,X1] : (~midp(sK20,X1,X2) | midp(X3,sK24,sK23) | ~para(X1,sK24,X2,sK23) | ~midp(X3,X1,X2)) )),
  inference(resolution,[],[f575,f424])).
fof(f424,plain,(
  ( ! [X4,X2,X0,X3,X1] : (~para(X0,X3,X1,X2) | midp(X4,X2,X3) | ~para(X0,X2,X1,X3) | ~midp(X4,X0,X1)) )),
  inference(cnf_transformation,[],[f265])).
fof(f265,plain,(
  ! [X0,X1,X2,X3,X4] : (midp(X4,X2,X3) | ~para(X0,X3,X1,X2) | ~para(X0,X2,X1,X3) | ~midp(X4,X0,X1))),
  inference(flattening,[],[f264])).
fof(f264,plain,(
  ! [X0,X1,X2,X3,X4] : (midp(X4,X2,X3) | (~para(X0,X3,X1,X2) | ~para(X0,X2,X1,X3) | ~midp(X4,X0,X1)))),
  inference(ennf_transformation,[],[f143])).
fof(f143,plain,(
  ! [X0,X1,X2,X3,X4] : ((para(X0,X3,X1,X2) & para(X0,X2,X1,X3) & midp(X4,X0,X1)) => midp(X4,X2,X3))),
  inference(rectify,[],[f65])).
fof(f65,axiom,(
  ! [X0,X1,X2,X3,X6] : ((para(X0,X3,X1,X2) & para(X0,X2,X1,X3) & midp(X6,X0,X1)) => midp(X6,X2,X3))),
  file('./7_Book/00EE/04/E051-5.p',ruleD64)).
fof(f575,plain,(
  ( ! [X4,X5] : (para(X4,sK23,X5,sK24) | ~midp(sK20,X4,X5)) )),
  inference(resolution,[],[f423,f479])).
fof(f12208,plain,(
  ( ! [X6] : (~midp(X6,sK23,sK24) | cong(sK23,X6,sK25,X6)) ) | ~spl29_114),
  inference(resolution,[],[f1209,f412])).
fof(f412,plain,(
  ( ! [X2,X0,X3,X1] : (~perp(X0,X1,X1,X2) | ~midp(X3,X0,X2) | cong(X0,X3,X1,X3)) )),
  inference(cnf_transformation,[],[f244])).
fof(f244,plain,(
  ! [X0,X1,X2,X3] : (cong(X0,X3,X1,X3) | ~midp(X3,X0,X2) | ~perp(X0,X1,X1,X2))),
  inference(flattening,[],[f243])).
fof(f243,plain,(
  ! [X0,X1,X2,X3] : (cong(X0,X3,X1,X3) | (~midp(X3,X0,X2) | ~perp(X0,X1,X1,X2)))),
  inference(ennf_transformation,[],[f132])).
fof(f132,plain,(
  ! [X0,X1,X2,X3] : ((midp(X3,X0,X2) & perp(X0,X1,X1,X2)) => cong(X0,X3,X1,X3))),
  inference(rectify,[],[f53])).
fof(f53,axiom,(
  ! [X0,X1,X2,X6] : ((midp(X6,X0,X2) & perp(X0,X1,X1,X2)) => cong(X0,X6,X1,X6))),
  file('./7_Book/00EE/04/E051-5.p',ruleD52)).
fof(f1209,plain,(
  perp(sK23,sK25,sK25,sK24) | ~spl29_114),
  inference(avatar_component_clause,[],[f1208])).
fof(f1208,plain,(
  spl29_114 <=> perp(sK23,sK25,sK25,sK24)),
  introduced(avatar_definition,[new_symbols(naming,[spl29_114])])).
fof(f12325,plain,(
  ( ! [X12,X13] : (perp(sK23,sK25,X12,X13) | ~cong(sK23,X12,sK25,X12)) ) | (~spl29_114 | ~spl29_180)),
  inference(resolution,[],[f12215,f416])).
fof(f416,plain,(
  ( ! [X2,X0,X3,X1] : (~cong(X0,X3,X1,X3) | perp(X0,X1,X2,X3) | ~cong(X0,X2,X1,X2)) )),
  inference(cnf_transformation,[],[f252])).
fof(f252,plain,(
  ! [X0,X1,X2,X3] : (perp(X0,X1,X2,X3) | ~cong(X0,X3,X1,X3) | ~cong(X0,X2,X1,X2))),
  inference(flattening,[],[f251])).
fof(f251,plain,(
  ! [X0,X1,X2,X3] : (perp(X0,X1,X2,X3) | (~cong(X0,X3,X1,X3) | ~cong(X0,X2,X1,X2)))),
  inference(ennf_transformation,[],[f135])).
fof(f135,plain,(
  ! [X0,X1,X2,X3] : ((cong(X0,X3,X1,X3) & cong(X0,X2,X1,X2)) => perp(X0,X1,X2,X3))),
  inference(rectify,[],[f57])).
fof(f57,axiom,(
  ! [X0,X1,X8,X9] : ((cong(X0,X9,X1,X9) & cong(X0,X8,X1,X8)) => perp(X0,X1,X8,X9))),
  file('./7_Book/00EE/04/E051-5.p',ruleD56)).
fof(f15138,plain,(
  ( ! [X28,X26,X27,X25] : (para(X25,X26,X27,X28) | ~perp(X25,X26,sK23,sK25)) ) | (~spl29_114 | ~spl29_180)),
  inference(resolution,[],[f12328,f368])).
fof(f368,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (~perp(X2,X3,X4,X5) | para(X0,X1,X4,X5) | ~perp(X0,X1,X2,X3)) )),
  inference(cnf_transformation,[],[f180])).
fof(f180,plain,(
  ! [X0,X1,X2,X3,X4,X5] : (para(X0,X1,X4,X5) | ~perp(X2,X3,X4,X5) | ~perp(X0,X1,X2,X3))),
  inference(flattening,[],[f179])).
fof(f179,plain,(
  ! [X0,X1,X2,X3,X4,X5] : (para(X0,X1,X4,X5) | (~perp(X2,X3,X4,X5) | ~perp(X0,X1,X2,X3)))),
  inference(ennf_transformation,[],[f9])).
fof(f9,axiom,(
  ! [X0,X1,X2,X3,X4,X5] : ((perp(X2,X3,X4,X5) & perp(X0,X1,X2,X3)) => para(X0,X1,X4,X5))),
  file('./7_Book/00EE/04/E051-5.p',ruleD9)).
fof(f15139,plain,(
  ( ! [X30,X31,X29,X32] : (perp(X29,X30,X31,X32) | ~para(X29,X30,sK23,sK25)) ) | (~spl29_114 | ~spl29_180)),
  inference(resolution,[],[f12328,f369])).
fof(f369,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (~perp(X2,X3,X4,X5) | perp(X0,X1,X4,X5) | ~para(X0,X1,X2,X3)) )),
  inference(cnf_transformation,[],[f182])).
fof(f182,plain,(
  ! [X0,X1,X2,X3,X4,X5] : (perp(X0,X1,X4,X5) | ~perp(X2,X3,X4,X5) | ~para(X0,X1,X2,X3))),
  inference(flattening,[],[f181])).
fof(f181,plain,(
  ! [X0,X1,X2,X3,X4,X5] : (perp(X0,X1,X4,X5) | (~perp(X2,X3,X4,X5) | ~para(X0,X1,X2,X3)))),
  inference(ennf_transformation,[],[f10])).
fof(f10,axiom,(
  ! [X0,X1,X2,X3,X4,X5] : ((perp(X2,X3,X4,X5) & para(X0,X1,X2,X3)) => perp(X0,X1,X4,X5))),
  file('./7_Book/00EE/04/E051-5.p',ruleD10)).
fof(f20260,plain,(
  ( ! [X2,X0,X1] : (midp(X0,X0,X1) | ~perp(X2,X1,X1,X0)) ) | (~spl29_114 | ~spl29_180)),
  inference(subsumption_resolution,[],[f20259,f9231])).
fof(f20259,plain,(
  ( ! [X2,X0,X3,X1] : (midp(X0,X0,X1) | ~cyclic(X2,X0,X1,X3) | ~perp(X2,X1,X1,X0)) ) | (~spl29_114 | ~spl29_180)),
  inference(subsumption_resolution,[],[f20258,f15161])).
fof(f20258,plain,(
  ( ! [X2,X0,X3,X1] : (midp(X0,X0,X1) | ~para(X2,X0,X2,X1) | ~cyclic(X2,X0,X1,X3) | ~perp(X2,X1,X1,X0)) ) | ~spl29_180),
  inference(subsumption_resolution,[],[f20251,f5958])).
fof(f20251,plain,(
  ( ! [X2,X0,X3,X1] : (midp(X0,X0,X1) | ~coll(X0,X0,X1) | ~para(X2,X0,X2,X1) | ~cyclic(X2,X0,X1,X3) | ~perp(X2,X1,X1,X0)) )),
  inference(resolution,[],[f2355,f465])).
fof(f465,plain,(
  ( ! [X2,X0,X3,X1] : (circle(sK15(X0,X1,X2,X3),X0,X1,X2) | ~cyclic(X0,X1,X2,X3) | ~perp(X0,X2,X2,X1)) )),
  inference(cnf_transformation,[],[f349])).
fof(f349,plain,(
  ! [X0,X1,X2,X3] : (circle(sK15(X0,X1,X2,X3),X0,X1,X2) | ~cyclic(X0,X1,X2,X3) | ~perp(X0,X2,X2,X1))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK15])],[f311,f348])).
fof(f348,plain,(
  ! [X3,X2,X1,X0] : (? [X4] : (circle(X4,X0,X1,X2) | ~cyclic(X0,X1,X2,X3) | ~perp(X0,X2,X2,X1)) => (circle(sK15(X0,X1,X2,X3),X0,X1,X2) | ~cyclic(X0,X1,X2,X3) | ~perp(X0,X2,X2,X1)))),
  introduced(choice_axiom,[])).
fof(f311,plain,(
  ! [X0,X1,X2,X3] : ? [X4] : (circle(X4,X0,X1,X2) | ~cyclic(X0,X1,X2,X3) | ~perp(X0,X2,X2,X1))),
  inference(flattening,[],[f310])).
fof(f310,plain,(
  ! [X0,X1,X2,X3] : ? [X4] : (circle(X4,X0,X1,X2) | (~cyclic(X0,X1,X2,X3) | ~perp(X0,X2,X2,X1)))),
  inference(ennf_transformation,[],[f162])).
fof(f162,plain,(
  ! [X0,X1,X2,X3] : ? [X4] : ((cyclic(X0,X1,X2,X3) & perp(X0,X2,X2,X1)) => circle(X4,X0,X1,X2))),
  inference(rectify,[],[f90])).
fof(f90,axiom,(
  ! [X0,X1,X2,X3] : ? [X7] : ((cyclic(X0,X1,X2,X3) & perp(X0,X2,X2,X1)) => circle(X7,X0,X1,X2))),
  file('./7_Book/00EE/04/E051-5.p',ruleX14)).
fof(f2355,plain,(
  ( ! [X78,X76,X79,X77] : (~circle(X79,X76,X77,X78) | midp(X77,X77,X78) | ~coll(X77,X77,X78) | ~para(X76,X77,X76,X78)) )),
  inference(resolution,[],[f829,f411])).
fof(f411,plain,(
  ( ! [X4,X2,X0,X3,X1] : (~eqangle(X0,X1,X0,X2,X3,X1,X3,X4) | midp(X4,X1,X2) | ~coll(X4,X1,X2) | ~circle(X3,X0,X1,X2)) )),
  inference(cnf_transformation,[],[f242])).
fof(f242,plain,(
  ! [X0,X1,X2,X3,X4] : (midp(X4,X1,X2) | ~eqangle(X0,X1,X0,X2,X3,X1,X3,X4) | ~coll(X4,X1,X2) | ~circle(X3,X0,X1,X2))),
  inference(flattening,[],[f241])).
fof(f241,plain,(
  ! [X0,X1,X2,X3,X4] : (midp(X4,X1,X2) | (~eqangle(X0,X1,X0,X2,X3,X1,X3,X4) | ~coll(X4,X1,X2) | ~circle(X3,X0,X1,X2)))),
  inference(ennf_transformation,[],[f131])).
fof(f131,plain,(
  ! [X0,X1,X2,X3,X4] : ((eqangle(X0,X1,X0,X2,X3,X1,X3,X4) & coll(X4,X1,X2) & circle(X3,X0,X1,X2)) => midp(X4,X1,X2))),
  inference(rectify,[],[f52])).
fof(f52,axiom,(
  ! [X0,X1,X2,X7,X6] : ((eqangle(X0,X1,X0,X2,X7,X1,X7,X6) & coll(X6,X1,X2) & circle(X7,X0,X1,X2)) => midp(X6,X1,X2))),
  file('./7_Book/00EE/04/E051-5.p',ruleD51)).
fof(f20266,plain,(
  ( ! [X14,X12,X13] : (cong(X12,X13,X12,X14) | ~midp(X12,X12,X13)) ) | (~spl29_114 | ~spl29_180)),
  inference(resolution,[],[f20261,f2032])).
fof(f2032,plain,(
  ( ! [X2,X0,X3,X1] : (~midp(X0,X3,X2) | cong(X0,X1,X0,X2) | ~midp(X0,X3,X1)) )),
  inference(resolution,[],[f721,f539])).
fof(f539,plain,(
  ( ! [X2,X0,X1] : (cong(X0,X1,X0,X2) | ~midp(X0,X2,X1)) )),
  inference(resolution,[],[f383,f428])).
fof(f721,plain,(
  ( ! [X4,X2,X0,X3,X1] : (~cong(X0,X1,X2,X4) | cong(X0,X1,X2,X3) | ~midp(X2,X4,X3)) )),
  inference(resolution,[],[f384,f428])).
fof(f384,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (~cong(X2,X3,X4,X5) | cong(X0,X1,X4,X5) | ~cong(X0,X1,X2,X3)) )),
  inference(cnf_transformation,[],[f202])).
fof(f202,plain,(
  ! [X0,X1,X2,X3,X4,X5] : (cong(X0,X1,X4,X5) | ~cong(X2,X3,X4,X5) | ~cong(X0,X1,X2,X3))),
  inference(flattening,[],[f201])).
fof(f201,plain,(
  ! [X0,X1,X2,X3,X4,X5] : (cong(X0,X1,X4,X5) | (~cong(X2,X3,X4,X5) | ~cong(X0,X1,X2,X3)))),
  inference(ennf_transformation,[],[f25])).
fof(f25,axiom,(
  ! [X0,X1,X2,X3,X4,X5] : ((cong(X2,X3,X4,X5) & cong(X0,X1,X2,X3)) => cong(X0,X1,X4,X5))),
  file('./7_Book/00EE/04/E051-5.p',ruleD25)).
fof(f12094,plain,(
  spl29_114 | ~spl29_903),
  inference(avatar_split_clause,[],[f7216,f12028,f1208])).
fof(f12028,plain,(
  spl29_903 <=> ~perp(sK23,sK24,sK23,sK25)),
  introduced(avatar_definition,[new_symbols(naming,[spl29_903])])).
fof(f7216,plain,(
  ~perp(sK23,sK24,sK23,sK25) | perp(sK23,sK25,sK25,sK24)),
  inference(resolution,[],[f1682,f434])).
fof(f434,plain,(
  ( ! [X6,X4,X2,X0,X7,X5,X3,X1] : (~eqangle(X0,X1,X2,X3,X4,X5,X6,X7) | ~perp(X4,X5,X6,X7) | perp(X0,X1,X2,X3)) )),
  inference(cnf_transformation,[],[f282])).
fof(f282,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6,X7] : (perp(X0,X1,X2,X3) | ~perp(X4,X5,X6,X7) | ~eqangle(X0,X1,X2,X3,X4,X5,X6,X7))),
  inference(flattening,[],[f281])).
fof(f281,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6,X7] : (perp(X0,X1,X2,X3) | (~perp(X4,X5,X6,X7) | ~eqangle(X0,X1,X2,X3,X4,X5,X6,X7)))),
  inference(ennf_transformation,[],[f147])).
fof(f147,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6,X7] : ((perp(X4,X5,X6,X7) & eqangle(X0,X1,X2,X3,X4,X5,X6,X7)) => perp(X0,X1,X2,X3))),
  inference(rectify,[],[f75])).
fof(f75,axiom,(
  ! [X0,X1,X2,X3,X8,X9,X10,X11] : ((perp(X8,X9,X10,X11) & eqangle(X0,X1,X2,X3,X8,X9,X10,X11)) => perp(X0,X1,X2,X3))),
  file('./7_Book/00EE/04/E051-5.p',ruleD74)).
fof(f1682,plain,(
  eqangle(sK23,sK25,sK25,sK24,sK23,sK24,sK23,sK25)),
  inference(resolution,[],[f788,f379])).
fof(f788,plain,(
  eqangle(sK23,sK24,sK23,sK25,sK23,sK25,sK25,sK24)),
  inference(resolution,[],[f377,f474])).
fof(f474,plain,(
  eqangle(sK24,sK23,sK23,sK25,sK23,sK25,sK25,sK24)),
  inference(cnf_transformation,[],[f359])).
fof(f12061,plain,(
  spl29_902 | ~spl29_180 | ~spl29_802),
  inference(avatar_split_clause,[],[f12060,f10054,f2150,f12031])).
fof(f12031,plain,(
  spl29_902 <=> perp(sK23,sK24,sK23,sK25)),
  introduced(avatar_definition,[new_symbols(naming,[spl29_902])])).
fof(f10054,plain,(
  spl29_802 <=> circle(sK24,sK23,sK24,sK25)),
  introduced(avatar_definition,[new_symbols(naming,[spl29_802])])).
fof(f12060,plain,(
  perp(sK23,sK24,sK23,sK25) | (~spl29_180 | ~spl29_802)),
  inference(subsumption_resolution,[],[f12059,f10842])).
fof(f10842,plain,(
  ( ! [X21,X19,X20] : (perp(X19,X19,X20,X21)) ) | ~spl29_180),
  inference(subsumption_resolution,[],[f10810,f10802])).
fof(f10810,plain,(
  ( ! [X21,X19,X20] : (perp(X19,X19,X20,X21) | ~cong(X19,X20,X19,X20)) ) | ~spl29_180),
  inference(resolution,[],[f10802,f416])).
fof(f12059,plain,(
  ( ! [X10] : (perp(sK23,sK24,sK23,sK25) | ~perp(sK24,sK24,sK24,X10)) ) | (~spl29_180 | ~spl29_802)),
  inference(subsumption_resolution,[],[f12014,f10957])).
fof(f10957,plain,(
  ( ! [X90] : (midp(X90,sK24,sK25)) ) | ~spl29_180),
  inference(subsumption_resolution,[],[f10923,f10843])).
fof(f10923,plain,(
  ( ! [X90,X91] : (midp(X90,sK24,sK25) | ~midp(X90,X91,X91)) ) | ~spl29_180),
  inference(resolution,[],[f10843,f1654])).
fof(f1654,plain,(
  ( ! [X2,X3,X1] : (~midp(sK21,X1,X2) | midp(X3,sK24,sK25) | ~midp(X3,X1,X2)) )),
  inference(subsumption_resolution,[],[f1648,f577])).
fof(f577,plain,(
  ( ! [X8,X9] : (para(X8,sK24,X9,sK25) | ~midp(sK21,X8,X9)) )),
  inference(resolution,[],[f423,f495])).
fof(f495,plain,(
  midp(sK21,sK24,sK25)),
  inference(resolution,[],[f370,f477])).
fof(f477,plain,(
  midp(sK21,sK25,sK24)),
  inference(cnf_transformation,[],[f359])).
fof(f1648,plain,(
  ( ! [X2,X3,X1] : (~midp(sK21,X1,X2) | midp(X3,sK24,sK25) | ~para(X1,sK24,X2,sK25) | ~midp(X3,X1,X2)) )),
  inference(resolution,[],[f574,f424])).
fof(f574,plain,(
  ( ! [X2,X3] : (para(X2,sK25,X3,sK24) | ~midp(sK21,X2,X3)) )),
  inference(resolution,[],[f423,f477])).
fof(f12014,plain,(
  ( ! [X10] : (perp(sK23,sK24,sK23,sK25) | ~midp(X10,sK24,sK25) | ~perp(sK24,sK24,sK24,X10)) ) | ~spl29_802),
  inference(resolution,[],[f10055,f1191])).
fof(f1191,plain,(
  ( ! [X14,X17,X15,X18,X16] : (~circle(X14,X17,X15,X18) | perp(X17,X15,X17,X18) | ~midp(X16,X15,X18) | ~perp(X14,X15,X14,X16)) )),
  inference(resolution,[],[f434,f410])).
fof(f410,plain,(
  ( ! [X4,X2,X0,X3,X1] : (eqangle(X0,X1,X0,X2,X3,X1,X3,X4) | ~midp(X4,X1,X2) | ~circle(X3,X0,X1,X2)) )),
  inference(cnf_transformation,[],[f240])).
fof(f240,plain,(
  ! [X0,X1,X2,X3,X4] : (eqangle(X0,X1,X0,X2,X3,X1,X3,X4) | ~midp(X4,X1,X2) | ~circle(X3,X0,X1,X2))),
  inference(flattening,[],[f239])).
fof(f239,plain,(
  ! [X0,X1,X2,X3,X4] : (eqangle(X0,X1,X0,X2,X3,X1,X3,X4) | (~midp(X4,X1,X2) | ~circle(X3,X0,X1,X2)))),
  inference(ennf_transformation,[],[f130])).
fof(f130,plain,(
  ! [X0,X1,X2,X3,X4] : ((midp(X4,X1,X2) & circle(X3,X0,X1,X2)) => eqangle(X0,X1,X0,X2,X3,X1,X3,X4))),
  inference(rectify,[],[f51])).
fof(f51,axiom,(
  ! [X0,X1,X2,X7,X6] : ((midp(X6,X1,X2) & circle(X7,X0,X1,X2)) => eqangle(X0,X1,X0,X2,X7,X1,X7,X6))),
  file('./7_Book/00EE/04/E051-5.p',ruleD50)).
fof(f10055,plain,(
  circle(sK24,sK23,sK24,sK25) | ~spl29_802),
  inference(avatar_component_clause,[],[f10054])).
fof(f11224,plain,(
  ~spl29_180 | spl29_803),
  inference(avatar_contradiction_clause,[],[f11223])).
fof(f11223,plain,(
  $false | (~spl29_180 | ~spl29_803)),
  inference(subsumption_resolution,[],[f11222,f11106])).
fof(f11222,plain,(
  ~midp(sK24,sK23,sK24) | (~spl29_180 | ~spl29_803)),
  inference(subsumption_resolution,[],[f10092,f10958])).
fof(f10958,plain,(
  ( ! [X93] : (midp(X93,sK23,sK25)) ) | ~spl29_180),
  inference(subsumption_resolution,[],[f10926,f10843])).
fof(f10926,plain,(
  ( ! [X94,X93] : (midp(X93,sK23,sK25) | ~midp(X93,X94,X94)) ) | ~spl29_180),
  inference(resolution,[],[f10843,f1643])).
fof(f1643,plain,(
  ( ! [X2,X3,X1] : (~midp(sK22,X1,X2) | midp(X3,sK23,sK25) | ~midp(X3,X1,X2)) )),
  inference(subsumption_resolution,[],[f1638,f576])).
fof(f576,plain,(
  ( ! [X6,X7] : (para(X6,sK23,X7,sK25) | ~midp(sK22,X6,X7)) )),
  inference(resolution,[],[f423,f494])).
fof(f494,plain,(
  midp(sK22,sK23,sK25)),
  inference(resolution,[],[f370,f475])).
fof(f475,plain,(
  midp(sK22,sK25,sK23)),
  inference(cnf_transformation,[],[f359])).
fof(f1638,plain,(
  ( ! [X2,X3,X1] : (~midp(sK22,X1,X2) | midp(X3,sK23,sK25) | ~para(X1,sK23,X2,sK25) | ~midp(X3,X1,X2)) )),
  inference(resolution,[],[f573,f424])).
fof(f573,plain,(
  ( ! [X0,X1] : (para(X0,sK25,X1,sK23) | ~midp(sK22,X0,X1)) )),
  inference(resolution,[],[f423,f475])).
fof(f10092,plain,(
  ~midp(sK24,sK23,sK25) | ~midp(sK24,sK23,sK24) | ~spl29_803),
  inference(resolution,[],[f10058,f2019])).
fof(f2019,plain,(
  ( ! [X6,X4,X7,X5] : (circle(X4,X5,X6,X7) | ~midp(X4,X5,X7) | ~midp(X4,X5,X6)) )),
  inference(resolution,[],[f702,f428])).
fof(f702,plain,(
  ( ! [X2,X0,X3,X1] : (~cong(X0,X1,X0,X2) | circle(X0,X1,X2,X3) | ~midp(X0,X1,X3)) )),
  inference(resolution,[],[f371,f428])).
fof(f371,plain,(
  ( ! [X2,X0,X3,X1] : (~cong(X3,X0,X3,X2) | circle(X3,X0,X1,X2) | ~cong(X3,X0,X3,X1)) )),
  inference(cnf_transformation,[],[f185])).
fof(f185,plain,(
  ! [X0,X1,X2,X3] : (circle(X3,X0,X1,X2) | ~cong(X3,X0,X3,X2) | ~cong(X3,X0,X3,X1))),
  inference(flattening,[],[f184])).
fof(f184,plain,(
  ! [X0,X1,X2,X3] : (circle(X3,X0,X1,X2) | (~cong(X3,X0,X3,X2) | ~cong(X3,X0,X3,X1)))),
  inference(ennf_transformation,[],[f98])).
fof(f98,plain,(
  ! [X0,X1,X2,X3] : ((cong(X3,X0,X3,X2) & cong(X3,X0,X3,X1)) => circle(X3,X0,X1,X2))),
  inference(rectify,[],[f12])).
fof(f12,axiom,(
  ! [X0,X1,X2,X7] : ((cong(X7,X0,X7,X2) & cong(X7,X0,X7,X1)) => circle(X7,X0,X1,X2))),
  file('./7_Book/00EE/04/E051-5.p',ruleD12)).
fof(f10058,plain,(
  ~circle(sK24,sK23,sK24,sK25) | ~spl29_803),
  inference(avatar_component_clause,[],[f10057])).
fof(f10057,plain,(
  spl29_803 <=> ~circle(sK24,sK23,sK24,sK25)),
  introduced(avatar_definition,[new_symbols(naming,[spl29_803])])).
fof(f5857,plain,(
  ~spl29_294),
  inference(avatar_contradiction_clause,[],[f5834])).
fof(f5834,plain,(
  $false | ~spl29_294),
  inference(resolution,[],[f3126,f479])).
fof(f3126,plain,(
  ( ! [X12,X10,X11] : (~midp(X10,X11,X12)) ) | ~spl29_294),
  inference(avatar_component_clause,[],[f3125])).
fof(f3125,plain,(
  spl29_294 <=> ! [X11,X10,X12] : ~midp(X10,X11,X12)),
  introduced(avatar_definition,[new_symbols(naming,[spl29_294])])).
fof(f4902,plain,(
  spl29_294 | spl29_294 | ~spl29_182),
  inference(avatar_split_clause,[],[f4901,f2153,f3125,f3125])).
fof(f2153,plain,(
  spl29_182 <=> ! [X1,X0] : ~para(X0,X1,X0,X1)),
  introduced(avatar_definition,[new_symbols(naming,[spl29_182])])).
fof(f4901,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (~midp(X0,X1,X2) | ~midp(X4,X5,X3)) ) | ~spl29_182),
  inference(subsumption_resolution,[],[f4900,f2154])).
fof(f2154,plain,(
  ( ! [X0,X1] : (~para(X0,X1,X0,X1)) ) | ~spl29_182),
  inference(avatar_component_clause,[],[f2153])).
fof(f4900,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (~midp(X4,X5,X3) | para(X2,X3,X2,X3) | ~midp(X0,X1,X2)) )),
  inference(duplicate_literal_removal,[],[f4885])).
fof(f4885,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (~midp(X0,X1,X2) | para(X2,X3,X2,X3) | ~midp(X4,X5,X3) | ~midp(X0,X1,X2) | ~midp(X4,X5,X3)) )),
  inference(resolution,[],[f1225,f1226])).
fof(f1226,plain,(
  ( ! [X26,X24,X23,X27,X25,X22] : (para(X27,X24,sK5(X26,X27,X23,X24,X25,X22),X25) | ~midp(X25,X26,X27) | ~midp(X22,X23,X24)) )),
  inference(resolution,[],[f447,f364])).
fof(f364,plain,(
  ( ! [X2,X0,X3,X1] : (~para(X0,X1,X2,X3) | para(X2,X3,X0,X1)) )),
  inference(cnf_transformation,[],[f174])).
fof(f174,plain,(
  ! [X0,X1,X2,X3] : (para(X2,X3,X0,X1) | ~para(X0,X1,X2,X3))),
  inference(ennf_transformation,[],[f5])).
fof(f5,axiom,(
  ! [X0,X1,X2,X3] : (para(X0,X1,X2,X3) => para(X2,X3,X0,X1))),
  file('./7_Book/00EE/04/E051-5.p',ruleD5)).
fof(f447,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (para(sK5(X0,X1,X2,X3,X4,X5),X4,X1,X3) | ~midp(X5,X2,X3) | ~midp(X4,X0,X1)) )),
  inference(cnf_transformation,[],[f333])).
fof(f333,plain,(
  ! [X0,X1,X2,X3,X4,X5] : ((para(sK5(X0,X1,X2,X3,X4,X5),X5,X0,X2) & para(sK5(X0,X1,X2,X3,X4,X5),X4,X1,X3) & midp(sK5(X0,X1,X2,X3,X4,X5),X0,X3)) | ~midp(X5,X2,X3) | ~midp(X4,X0,X1))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK5])],[f296,f332])).
fof(f332,plain,(
  ! [X5,X4,X3,X2,X1,X0] : (? [X6] : ((para(X6,X5,X0,X2) & para(X6,X4,X1,X3) & midp(X6,X0,X3)) | ~midp(X5,X2,X3) | ~midp(X4,X0,X1)) => ((para(sK5(X0,X1,X2,X3,X4,X5),X5,X0,X2) & para(sK5(X0,X1,X2,X3,X4,X5),X4,X1,X3) & midp(sK5(X0,X1,X2,X3,X4,X5),X0,X3)) | ~midp(X5,X2,X3) | ~midp(X4,X0,X1)))),
  introduced(choice_axiom,[])).
fof(f296,plain,(
  ! [X0,X1,X2,X3,X4,X5] : ? [X6] : ((para(X6,X5,X0,X2) & para(X6,X4,X1,X3) & midp(X6,X0,X3)) | ~midp(X5,X2,X3) | ~midp(X4,X0,X1))),
  inference(flattening,[],[f295])).
fof(f295,plain,(
  ! [X0,X1,X2,X3,X4,X5] : ? [X6] : ((para(X6,X5,X0,X2) & para(X6,X4,X1,X3) & midp(X6,X0,X3)) | (~midp(X5,X2,X3) | ~midp(X4,X0,X1)))),
  inference(ennf_transformation,[],[f154])).
fof(f154,plain,(
  ! [X0,X1,X2,X3,X4,X5] : ? [X6] : ((midp(X5,X2,X3) & midp(X4,X0,X1)) => (para(X6,X5,X0,X2) & para(X6,X4,X1,X3) & midp(X6,X0,X3)))),
  inference(rectify,[],[f82])).
fof(f82,axiom,(
  ! [X0,X1,X2,X3,X6,X16] : ? [X8] : ((midp(X16,X2,X3) & midp(X6,X0,X1)) => (para(X8,X16,X0,X2) & para(X8,X6,X1,X3) & midp(X8,X0,X3)))),
  file('./7_Book/00EE/04/E051-5.p',ruleX6)).
fof(f1225,plain,(
  ( ! [X14,X21,X19,X17,X15,X20,X18,X16] : (~para(X20,X21,sK5(X18,X19,X15,X16,X17,X14),X17) | ~midp(X17,X18,X19) | para(X20,X21,X19,X16) | ~midp(X14,X15,X16)) )),
  inference(resolution,[],[f447,f365])).
fof(f365,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (~para(X2,X3,X4,X5) | para(X0,X1,X4,X5) | ~para(X0,X1,X2,X3)) )),
  inference(cnf_transformation,[],[f176])).
fof(f176,plain,(
  ! [X0,X1,X2,X3,X4,X5] : (para(X0,X1,X4,X5) | ~para(X2,X3,X4,X5) | ~para(X0,X1,X2,X3))),
  inference(flattening,[],[f175])).
fof(f175,plain,(
  ! [X0,X1,X2,X3,X4,X5] : (para(X0,X1,X4,X5) | (~para(X2,X3,X4,X5) | ~para(X0,X1,X2,X3)))),
  inference(ennf_transformation,[],[f6])).
fof(f6,axiom,(
  ! [X0,X1,X2,X3,X4,X5] : ((para(X2,X3,X4,X5) & para(X0,X1,X2,X3)) => para(X0,X1,X4,X5))),
  file('./7_Book/00EE/04/E051-5.p',ruleD6)).
fof(f2155,plain,(
  spl29_180 | spl29_182),
  inference(avatar_split_clause,[],[f2126,f2153,f2150])).
fof(f2126,plain,(
  ( ! [X2,X0,X3,X1] : (~para(X0,X1,X0,X1) | para(X2,X3,X2,X3)) )),
  inference(resolution,[],[f813,f398])).
fof(f398,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (~eqangle(X0,X1,X4,X5,X2,X3,X4,X5) | para(X0,X1,X2,X3)) )),
  inference(cnf_transformation,[],[f219])).
fof(f219,plain,(
  ! [X0,X1,X2,X3,X4,X5] : (para(X0,X1,X2,X3) | ~eqangle(X0,X1,X4,X5,X2,X3,X4,X5))),
  inference(ennf_transformation,[],[f118])).
fof(f118,plain,(
  ! [X0,X1,X2,X3,X4,X5] : (eqangle(X0,X1,X4,X5,X2,X3,X4,X5) => para(X0,X1,X2,X3))),
  inference(rectify,[],[f39])).
fof(f39,axiom,(
  ! [X0,X1,X2,X3,X8,X9] : (eqangle(X0,X1,X8,X9,X2,X3,X8,X9) => para(X0,X1,X2,X3))),
  file('./7_Book/00EE/04/E051-5.p',ruleD39)).
fof(f813,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (eqangle(X0,X1,X2,X3,X0,X1,X4,X5) | ~para(X2,X3,X4,X5)) )),
  inference(resolution,[],[f378,f399])).
fof(f378,plain,(
  ( ! [X6,X4,X2,X0,X7,X5,X3,X1] : (~eqangle(X0,X1,X2,X3,X4,X5,X6,X7) | eqangle(X2,X3,X0,X1,X6,X7,X4,X5)) )),
  inference(cnf_transformation,[],[f194])).
fof(f194,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6,X7] : (eqangle(X2,X3,X0,X1,X6,X7,X4,X5) | ~eqangle(X0,X1,X2,X3,X4,X5,X6,X7))),
  inference(ennf_transformation,[],[f101])).
fof(f101,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6,X7] : (eqangle(X0,X1,X2,X3,X4,X5,X6,X7) => eqangle(X2,X3,X0,X1,X6,X7,X4,X5))),
  inference(rectify,[],[f19])).
fof(f19,axiom,(
  ! [X0,X1,X2,X3,X8,X9,X10,X11] : (eqangle(X0,X1,X2,X3,X8,X9,X10,X11) => eqangle(X2,X3,X0,X1,X10,X11,X8,X9))),
  file('./7_Book/00EE/04/E051-5.p',ruleD19)).
% SZS output end Proof for E051-5
% ------------------------------
% Version: Vampire 4.4.0 (commit 7916d27 on 2019-08-23 08:50:16 +0100)
% Termination reason: Refutation

% Memory used [KB]: 17654
% Time elapsed: 0.430 s
% ------------------------------
% ------------------------------
