% Refutation found. Thanks to Tanya!
% SZS status Theorem for E051-8
% SZS output start Proof for E051-8
fof(f22561,plain,(
  $false),
  inference(avatar_sat_refutation,[],[f1870,f2046,f2348,f3162,f3167,f3169,f3512,f3825,f3957,f6123,f6904,f6917,f6919,f7379,f7608,f7716,f8507,f8511,f8562,f8567,f12555,f12678,f12889,f13233,f13702,f13805,f13856,f13876,f14577,f15317,f22550,f22560])).
fof(f22560,plain,(
  ~spl26_148 | ~spl26_272 | ~spl26_406 | spl26_553),
  inference(avatar_contradiction_clause,[],[f22559])).
fof(f22559,plain,(
  $false | (~spl26_148 | ~spl26_272 | ~spl26_406 | ~spl26_553)),
  inference(subsumption_resolution,[],[f7607,f15373])).
fof(f15373,plain,(
  ( ! [X4] : (midp(X4,sK22,sK21)) ) | (~spl26_148 | ~spl26_272 | ~spl26_406)),
  inference(subsumption_resolution,[],[f15372,f12431])).
fof(f12431,plain,(
  ( ! [X35,X34] : (midp(X34,X35,X35)) ) | ~spl26_148),
  inference(subsumption_resolution,[],[f12393,f4334])).
fof(f4334,plain,(
  ( ! [X4,X3] : (coll(X3,X4,X4)) ) | ~spl26_148),
  inference(resolution,[],[f1866,f424])).
fof(f424,plain,(
  ( ! [X2,X0,X1] : (~para(X0,X1,X0,X2) | coll(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f266])).
fof(f266,plain,(
  ! [X0,X1,X2] : (coll(X0,X1,X2) | ~para(X0,X1,X0,X2))),
  inference(ennf_transformation,[],[f67])).
fof(f67,axiom,(
  ! [X0,X1,X2] : (para(X0,X1,X0,X2) => coll(X0,X1,X2))),
  file('./7_Book/00EE/04/E051-8.p',ruleD66)).
fof(f1866,plain,(
  ( ! [X2,X3] : (para(X2,X3,X2,X3)) ) | ~spl26_148),
  inference(avatar_component_clause,[],[f1865])).
fof(f1865,plain,(
  spl26_148 <=> ! [X3,X2] : para(X2,X3,X2,X3)),
  introduced(avatar_definition,[new_symbols(naming,[spl26_148])])).
fof(f12393,plain,(
  ( ! [X35,X34] : (~coll(X34,X35,X35) | midp(X34,X35,X35)) ) | ~spl26_148),
  inference(resolution,[],[f6908,f425])).
fof(f425,plain,(
  ( ! [X2,X0,X1] : (~cong(X0,X1,X0,X2) | ~coll(X0,X1,X2) | midp(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f268])).
fof(f268,plain,(
  ! [X0,X1,X2] : (midp(X0,X1,X2) | ~coll(X0,X1,X2) | ~cong(X0,X1,X0,X2))),
  inference(flattening,[],[f267])).
fof(f267,plain,(
  ! [X0,X1,X2] : (midp(X0,X1,X2) | (~coll(X0,X1,X2) | ~cong(X0,X1,X0,X2)))),
  inference(ennf_transformation,[],[f68])).
fof(f68,axiom,(
  ! [X0,X1,X2] : ((coll(X0,X1,X2) & cong(X0,X1,X0,X2)) => midp(X0,X1,X2))),
  file('./7_Book/00EE/04/E051-8.p',ruleD67)).
fof(f6908,plain,(
  ( ! [X35,X36] : (cong(X35,X36,X35,X36)) ) | ~spl26_148),
  inference(subsumption_resolution,[],[f6907,f6791])).
fof(f6791,plain,(
  ( ! [X4,X5,X3] : (cyclic(X3,X4,X3,X5)) ) | ~spl26_148),
  inference(resolution,[],[f6782,f372])).
fof(f372,plain,(
  ( ! [X2,X0,X3,X1] : (~cyclic(X0,X1,X2,X3) | cyclic(X0,X2,X1,X3)) )),
  inference(cnf_transformation,[],[f187])).
fof(f187,plain,(
  ! [X0,X1,X2,X3] : (cyclic(X0,X2,X1,X3) | ~cyclic(X0,X1,X2,X3))),
  inference(ennf_transformation,[],[f15])).
fof(f15,axiom,(
  ! [X0,X1,X2,X3] : (cyclic(X0,X1,X2,X3) => cyclic(X0,X2,X1,X3))),
  file('./7_Book/00EE/04/E051-8.p',ruleD15)).
fof(f6782,plain,(
  ( ! [X6,X8,X7] : (cyclic(X6,X6,X7,X8)) ) | ~spl26_148),
  inference(resolution,[],[f3783,f1866])).
fof(f3783,plain,(
  ( ! [X68,X66,X67] : (~para(X66,X67,X66,X67) | cyclic(X67,X67,X66,X68)) )),
  inference(subsumption_resolution,[],[f2029,f2030])).
fof(f2030,plain,(
  ( ! [X70,X71,X69] : (~para(X69,X70,X69,X70) | coll(X69,X71,X70) | cyclic(X70,X70,X69,X71)) )),
  inference(resolution,[],[f690,f399])).
fof(f399,plain,(
  ( ! [X2,X0,X3,X1] : (~eqangle(X2,X0,X2,X1,X3,X0,X3,X1) | coll(X2,X3,X0) | cyclic(X0,X1,X2,X3)) )),
  inference(cnf_transformation,[],[f221])).
fof(f221,plain,(
  ! [X0,X1,X2,X3] : (cyclic(X0,X1,X2,X3) | coll(X2,X3,X0) | ~eqangle(X2,X0,X2,X1,X3,X0,X3,X1))),
  inference(flattening,[],[f220])).
fof(f220,plain,(
  ! [X0,X1,X2,X3] : (cyclic(X0,X1,X2,X3) | (coll(X2,X3,X0) | ~eqangle(X2,X0,X2,X1,X3,X0,X3,X1)))),
  inference(ennf_transformation,[],[f121])).
fof(f121,plain,(
  ! [X0,X1,X2,X3] : ((~coll(X2,X3,X0) & eqangle(X2,X0,X2,X1,X3,X0,X3,X1)) => cyclic(X0,X1,X2,X3))),
  inference(rectify,[],[f42])).
fof(f42,axiom,(
  ! [X0,X1,X8,X9] : ((~coll(X8,X9,X0) & eqangle(X8,X0,X8,X1,X9,X0,X9,X1)) => cyclic(X0,X1,X8,X9))),
  file('./7_Book/00EE/04/E051-8.p',ruleD42a)).
fof(f690,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (eqangle(X0,X1,X2,X3,X4,X5,X4,X5) | ~para(X0,X1,X2,X3)) )),
  inference(resolution,[],[f378,f397])).
fof(f397,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (eqangle(X0,X1,X4,X5,X2,X3,X4,X5) | ~para(X0,X1,X2,X3)) )),
  inference(cnf_transformation,[],[f218])).
fof(f218,plain,(
  ! [X0,X1,X2,X3,X4,X5] : (eqangle(X0,X1,X4,X5,X2,X3,X4,X5) | ~para(X0,X1,X2,X3))),
  inference(ennf_transformation,[],[f119])).
fof(f119,plain,(
  ! [X0,X1,X2,X3,X4,X5] : (para(X0,X1,X2,X3) => eqangle(X0,X1,X4,X5,X2,X3,X4,X5))),
  inference(rectify,[],[f40])).
fof(f40,axiom,(
  ! [X0,X1,X2,X3,X8,X9] : (para(X0,X1,X2,X3) => eqangle(X0,X1,X8,X9,X2,X3,X8,X9))),
  file('./7_Book/00EE/04/E051-8.p',ruleD40)).
fof(f378,plain,(
  ( ! [X6,X4,X2,X0,X7,X5,X3,X1] : (~eqangle(X0,X1,X2,X3,X4,X5,X6,X7) | eqangle(X0,X1,X4,X5,X2,X3,X6,X7)) )),
  inference(cnf_transformation,[],[f194])).
fof(f194,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6,X7] : (eqangle(X0,X1,X4,X5,X2,X3,X6,X7) | ~eqangle(X0,X1,X2,X3,X4,X5,X6,X7))),
  inference(ennf_transformation,[],[f103])).
fof(f103,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6,X7] : (eqangle(X0,X1,X2,X3,X4,X5,X6,X7) => eqangle(X0,X1,X4,X5,X2,X3,X6,X7))),
  inference(rectify,[],[f21])).
fof(f21,axiom,(
  ! [X0,X1,X2,X3,X8,X9,X10,X11] : (eqangle(X0,X1,X2,X3,X8,X9,X10,X11) => eqangle(X0,X1,X8,X9,X2,X3,X10,X11))),
  file('./7_Book/00EE/04/E051-8.p',ruleD21)).
fof(f2029,plain,(
  ( ! [X68,X66,X67] : (~para(X66,X67,X66,X67) | ~coll(X66,X68,X67) | cyclic(X67,X67,X66,X68)) )),
  inference(resolution,[],[f690,f400])).
fof(f400,plain,(
  ( ! [X2,X0,X3,X1] : (~eqangle(X2,X0,X2,X1,X3,X0,X3,X1) | ~coll(X2,X3,X1) | cyclic(X0,X1,X2,X3)) )),
  inference(cnf_transformation,[],[f223])).
fof(f223,plain,(
  ! [X0,X1,X2,X3] : (cyclic(X0,X1,X2,X3) | ~coll(X2,X3,X1) | ~eqangle(X2,X0,X2,X1,X3,X0,X3,X1))),
  inference(flattening,[],[f222])).
fof(f222,plain,(
  ! [X0,X1,X2,X3] : (cyclic(X0,X1,X2,X3) | (~coll(X2,X3,X1) | ~eqangle(X2,X0,X2,X1,X3,X0,X3,X1)))),
  inference(ennf_transformation,[],[f122])).
fof(f122,plain,(
  ! [X0,X1,X2,X3] : ((coll(X2,X3,X1) & eqangle(X2,X0,X2,X1,X3,X0,X3,X1)) => cyclic(X0,X1,X2,X3))),
  inference(rectify,[],[f43])).
fof(f43,axiom,(
  ! [X0,X1,X8,X9] : ((coll(X8,X9,X1) & eqangle(X8,X0,X8,X1,X9,X0,X9,X1)) => cyclic(X0,X1,X8,X9))),
  file('./7_Book/00EE/04/E051-8.p',ruleD42b)).
fof(f6907,plain,(
  ( ! [X35,X36] : (cong(X35,X36,X35,X36) | ~cyclic(X35,X36,X35,X35)) ) | ~spl26_148),
  inference(subsumption_resolution,[],[f6891,f6791])).
fof(f6891,plain,(
  ( ! [X35,X36] : (cong(X35,X36,X35,X36) | ~cyclic(X35,X36,X35,X36) | ~cyclic(X35,X36,X35,X35)) ) | ~spl26_148),
  inference(resolution,[],[f6791,f1265])).
fof(f1265,plain,(
  ( ! [X10,X8,X7,X9] : (~cyclic(X7,X8,X9,X10) | cong(X7,X8,X7,X8) | ~cyclic(X7,X8,X9,X8) | ~cyclic(X7,X8,X9,X7)) )),
  inference(duplicate_literal_removal,[],[f1259])).
fof(f1259,plain,(
  ( ! [X10,X8,X7,X9] : (cong(X7,X8,X7,X8) | ~cyclic(X7,X8,X9,X10) | ~cyclic(X7,X8,X9,X8) | ~cyclic(X7,X8,X9,X7) | ~cyclic(X7,X8,X9,X10)) )),
  inference(resolution,[],[f401,f398])).
fof(f398,plain,(
  ( ! [X2,X0,X3,X1] : (eqangle(X2,X0,X2,X1,X3,X0,X3,X1) | ~cyclic(X0,X1,X2,X3)) )),
  inference(cnf_transformation,[],[f219])).
fof(f219,plain,(
  ! [X0,X1,X2,X3] : (eqangle(X2,X0,X2,X1,X3,X0,X3,X1) | ~cyclic(X0,X1,X2,X3))),
  inference(ennf_transformation,[],[f120])).
fof(f120,plain,(
  ! [X0,X1,X2,X3] : (cyclic(X0,X1,X2,X3) => eqangle(X2,X0,X2,X1,X3,X0,X3,X1))),
  inference(rectify,[],[f41])).
fof(f41,axiom,(
  ! [X0,X1,X8,X9] : (cyclic(X0,X1,X8,X9) => eqangle(X8,X0,X8,X1,X9,X0,X9,X1))),
  file('./7_Book/00EE/04/E051-8.p',ruleD41)).
fof(f401,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (~eqangle(X2,X0,X2,X1,X5,X3,X5,X4) | cong(X0,X1,X3,X4) | ~cyclic(X0,X1,X2,X5) | ~cyclic(X0,X1,X2,X4) | ~cyclic(X0,X1,X2,X3)) )),
  inference(cnf_transformation,[],[f225])).
fof(f225,plain,(
  ! [X0,X1,X2,X3,X4,X5] : (cong(X0,X1,X3,X4) | ~eqangle(X2,X0,X2,X1,X5,X3,X5,X4) | ~cyclic(X0,X1,X2,X5) | ~cyclic(X0,X1,X2,X4) | ~cyclic(X0,X1,X2,X3))),
  inference(flattening,[],[f224])).
fof(f224,plain,(
  ! [X0,X1,X2,X3,X4,X5] : (cong(X0,X1,X3,X4) | (~eqangle(X2,X0,X2,X1,X5,X3,X5,X4) | ~cyclic(X0,X1,X2,X5) | ~cyclic(X0,X1,X2,X4) | ~cyclic(X0,X1,X2,X3)))),
  inference(ennf_transformation,[],[f123])).
fof(f123,plain,(
  ! [X0,X1,X2,X3,X4,X5] : ((eqangle(X2,X0,X2,X1,X5,X3,X5,X4) & cyclic(X0,X1,X2,X5) & cyclic(X0,X1,X2,X4) & cyclic(X0,X1,X2,X3)) => cong(X0,X1,X3,X4))),
  inference(rectify,[],[f44])).
fof(f44,axiom,(
  ! [X0,X1,X2,X8,X9,X14] : ((eqangle(X2,X0,X2,X1,X14,X8,X14,X9) & cyclic(X0,X1,X2,X14) & cyclic(X0,X1,X2,X9) & cyclic(X0,X1,X2,X8)) => cong(X0,X1,X8,X9))),
  file('./7_Book/00EE/04/E051-8.p',ruleD43)).
fof(f15372,plain,(
  ( ! [X4] : (midp(X4,sK22,sK21) | ~midp(X4,sK23,sK23)) ) | (~spl26_272 | ~spl26_406)),
  inference(subsumption_resolution,[],[f15365,f5543])).
fof(f5543,plain,(
  para(sK23,sK22,sK23,sK21) | ~spl26_406),
  inference(avatar_component_clause,[],[f5542])).
fof(f5542,plain,(
  spl26_406 <=> para(sK23,sK22,sK23,sK21)),
  introduced(avatar_definition,[new_symbols(naming,[spl26_406])])).
fof(f15365,plain,(
  ( ! [X4] : (midp(X4,sK22,sK21) | ~para(sK23,sK22,sK23,sK21) | ~midp(X4,sK23,sK23)) ) | ~spl26_272),
  inference(resolution,[],[f3828,f422])).
fof(f422,plain,(
  ( ! [X4,X2,X0,X3,X1] : (~para(X0,X3,X1,X2) | midp(X4,X2,X3) | ~para(X0,X2,X1,X3) | ~midp(X4,X0,X1)) )),
  inference(cnf_transformation,[],[f263])).
fof(f263,plain,(
  ! [X0,X1,X2,X3,X4] : (midp(X4,X2,X3) | ~para(X0,X3,X1,X2) | ~para(X0,X2,X1,X3) | ~midp(X4,X0,X1))),
  inference(flattening,[],[f262])).
fof(f262,plain,(
  ! [X0,X1,X2,X3,X4] : (midp(X4,X2,X3) | (~para(X0,X3,X1,X2) | ~para(X0,X2,X1,X3) | ~midp(X4,X0,X1)))),
  inference(ennf_transformation,[],[f143])).
fof(f143,plain,(
  ! [X0,X1,X2,X3,X4] : ((para(X0,X3,X1,X2) & para(X0,X2,X1,X3) & midp(X4,X0,X1)) => midp(X4,X2,X3))),
  inference(rectify,[],[f65])).
fof(f65,axiom,(
  ! [X0,X1,X2,X3,X6] : ((para(X0,X3,X1,X2) & para(X0,X2,X1,X3) & midp(X6,X0,X1)) => midp(X6,X2,X3))),
  file('./7_Book/00EE/04/E051-8.p',ruleD64)).
fof(f3828,plain,(
  para(sK23,sK21,sK23,sK22) | ~spl26_272),
  inference(avatar_component_clause,[],[f3827])).
fof(f3827,plain,(
  spl26_272 <=> para(sK23,sK21,sK23,sK22)),
  introduced(avatar_definition,[new_symbols(naming,[spl26_272])])).
fof(f7607,plain,(
  ~midp(sK25,sK22,sK21) | ~spl26_553),
  inference(avatar_component_clause,[],[f7606])).
fof(f7606,plain,(
  spl26_553 <=> ~midp(sK25,sK22,sK21)),
  introduced(avatar_definition,[new_symbols(naming,[spl26_553])])).
fof(f22550,plain,(
  ~spl26_148 | spl26_551),
  inference(avatar_contradiction_clause,[],[f22515])).
fof(f22515,plain,(
  $false | (~spl26_148 | ~spl26_551)),
  inference(resolution,[],[f22437,f7601])).
fof(f7601,plain,(
  ~midp(sK22,sK24,sK25) | ~spl26_551),
  inference(avatar_component_clause,[],[f7600])).
fof(f7600,plain,(
  spl26_551 <=> ~midp(sK22,sK24,sK25)),
  introduced(avatar_definition,[new_symbols(naming,[spl26_551])])).
fof(f22437,plain,(
  ( ! [X61,X62,X63] : (midp(X61,X62,X63)) ) | ~spl26_148),
  inference(subsumption_resolution,[],[f22436,f12431])).
fof(f22436,plain,(
  ( ! [X61,X64,X62,X63] : (midp(X61,X62,X63) | ~midp(X61,X64,X64)) ) | ~spl26_148),
  inference(subsumption_resolution,[],[f22408,f22375])).
fof(f22375,plain,(
  ( ! [X39,X38,X40] : (para(X40,X38,X40,X39)) ) | ~spl26_148),
  inference(subsumption_resolution,[],[f22374,f1866])).
fof(f22374,plain,(
  ( ! [X39,X38,X40] : (~para(X38,X39,X38,X39) | para(X40,X38,X40,X39)) ) | ~spl26_148),
  inference(subsumption_resolution,[],[f22373,f6791])).
fof(f22373,plain,(
  ( ! [X39,X38,X40] : (~cyclic(X38,X39,X38,X40) | ~para(X38,X39,X38,X39) | para(X40,X38,X40,X39)) ) | ~spl26_148),
  inference(subsumption_resolution,[],[f22332,f14881])).
fof(f14881,plain,(
  ( ! [X103,X101,X102,X100] : (cong(X100,X101,X102,X103)) ) | ~spl26_148),
  inference(subsumption_resolution,[],[f14880,f12431])).
fof(f14880,plain,(
  ( ! [X103,X101,X102,X100] : (~midp(X102,X103,X103) | cong(X100,X101,X102,X103)) ) | ~spl26_148),
  inference(subsumption_resolution,[],[f14854,f12431])).
fof(f14854,plain,(
  ( ! [X103,X101,X102,X100] : (~midp(X100,X101,X101) | ~midp(X102,X103,X103) | cong(X100,X101,X102,X103)) ) | ~spl26_148),
  inference(resolution,[],[f2372,f9364])).
fof(f9364,plain,(
  ( ! [X87,X86] : (cong(X87,X87,X86,X86)) ) | ~spl26_148),
  inference(subsumption_resolution,[],[f9363,f6782])).
fof(f9363,plain,(
  ( ! [X88,X87,X86] : (cong(X87,X87,X86,X86) | ~cyclic(X87,X87,X88,X86)) ) | ~spl26_148),
  inference(subsumption_resolution,[],[f9362,f6782])).
fof(f9362,plain,(
  ( ! [X88,X87,X86] : (cong(X87,X87,X86,X86) | ~cyclic(X87,X87,X88,X88) | ~cyclic(X87,X87,X88,X86)) ) | ~spl26_148),
  inference(subsumption_resolution,[],[f9351,f6817])).
fof(f6817,plain,(
  ( ! [X12,X10,X11,X9] : (cyclic(X9,X10,X11,X12)) ) | ~spl26_148),
  inference(subsumption_resolution,[],[f6793,f6782])).
fof(f6793,plain,(
  ( ! [X12,X10,X11,X9] : (cyclic(X9,X10,X11,X12) | ~cyclic(X9,X9,X10,X11)) ) | ~spl26_148),
  inference(resolution,[],[f6782,f374])).
fof(f374,plain,(
  ( ! [X4,X2,X0,X3,X1] : (~cyclic(X0,X1,X2,X4) | cyclic(X1,X2,X3,X4) | ~cyclic(X0,X1,X2,X3)) )),
  inference(cnf_transformation,[],[f190])).
fof(f190,plain,(
  ! [X0,X1,X2,X3,X4] : (cyclic(X1,X2,X3,X4) | ~cyclic(X0,X1,X2,X4) | ~cyclic(X0,X1,X2,X3))),
  inference(flattening,[],[f189])).
fof(f189,plain,(
  ! [X0,X1,X2,X3,X4] : (cyclic(X1,X2,X3,X4) | (~cyclic(X0,X1,X2,X4) | ~cyclic(X0,X1,X2,X3)))),
  inference(ennf_transformation,[],[f17])).
fof(f17,axiom,(
  ! [X0,X1,X2,X3,X4] : ((cyclic(X0,X1,X2,X4) & cyclic(X0,X1,X2,X3)) => cyclic(X1,X2,X3,X4))),
  file('./7_Book/00EE/04/E051-8.p',ruleD17)).
fof(f9351,plain,(
  ( ! [X88,X87,X86] : (~cyclic(X86,X87,X88,X88) | cong(X87,X87,X86,X86) | ~cyclic(X87,X87,X88,X88) | ~cyclic(X87,X87,X88,X86)) )),
  inference(duplicate_literal_removal,[],[f9340])).
fof(f9340,plain,(
  ( ! [X88,X87,X86] : (~cyclic(X86,X87,X88,X88) | cong(X87,X87,X86,X86) | ~cyclic(X87,X87,X88,X88) | ~cyclic(X87,X87,X88,X86) | ~cyclic(X87,X87,X88,X86)) )),
  inference(resolution,[],[f1880,f401])).
fof(f1880,plain,(
  ( ! [X17,X15,X18,X16] : (eqangle(X17,X16,X18,X16,X17,X15,X18,X15) | ~cyclic(X15,X16,X17,X18)) )),
  inference(resolution,[],[f674,f378])).
fof(f674,plain,(
  ( ! [X6,X8,X7,X9] : (eqangle(X6,X7,X6,X8,X9,X7,X9,X8) | ~cyclic(X8,X7,X6,X9)) )),
  inference(resolution,[],[f376,f398])).
fof(f376,plain,(
  ( ! [X6,X4,X2,X0,X7,X5,X3,X1] : (~eqangle(X0,X1,X2,X3,X4,X5,X6,X7) | eqangle(X2,X3,X0,X1,X6,X7,X4,X5)) )),
  inference(cnf_transformation,[],[f192])).
fof(f192,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6,X7] : (eqangle(X2,X3,X0,X1,X6,X7,X4,X5) | ~eqangle(X0,X1,X2,X3,X4,X5,X6,X7))),
  inference(ennf_transformation,[],[f101])).
fof(f101,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6,X7] : (eqangle(X0,X1,X2,X3,X4,X5,X6,X7) => eqangle(X2,X3,X0,X1,X6,X7,X4,X5))),
  inference(rectify,[],[f19])).
fof(f19,axiom,(
  ! [X0,X1,X2,X3,X8,X9,X10,X11] : (eqangle(X0,X1,X2,X3,X8,X9,X10,X11) => eqangle(X2,X3,X0,X1,X10,X11,X8,X9))),
  file('./7_Book/00EE/04/E051-8.p',ruleD19)).
fof(f2372,plain,(
  ( ! [X14,X12,X10,X15,X13,X11] : (~cong(X14,X15,X11,X12) | ~midp(X13,X14,X15) | ~midp(X10,X11,X12) | cong(X13,X14,X10,X11)) )),
  inference(resolution,[],[f714,f433])).
fof(f433,plain,(
  ( ! [X6,X4,X2,X0,X7,X5,X3,X1] : (~eqratio(X0,X1,X2,X3,X4,X5,X6,X7) | ~cong(X4,X5,X6,X7) | cong(X0,X1,X2,X3)) )),
  inference(cnf_transformation,[],[f282])).
fof(f282,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6,X7] : (cong(X0,X1,X2,X3) | ~cong(X4,X5,X6,X7) | ~eqratio(X0,X1,X2,X3,X4,X5,X6,X7))),
  inference(flattening,[],[f281])).
fof(f281,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6,X7] : (cong(X0,X1,X2,X3) | (~cong(X4,X5,X6,X7) | ~eqratio(X0,X1,X2,X3,X4,X5,X6,X7)))),
  inference(ennf_transformation,[],[f148])).
fof(f148,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6,X7] : ((cong(X4,X5,X6,X7) & eqratio(X0,X1,X2,X3,X4,X5,X6,X7)) => cong(X0,X1,X2,X3))),
  inference(rectify,[],[f76])).
fof(f76,axiom,(
  ! [X0,X1,X2,X3,X8,X9,X10,X11] : ((cong(X8,X9,X10,X11) & eqratio(X0,X1,X2,X3,X8,X9,X10,X11)) => cong(X0,X1,X2,X3))),
  file('./7_Book/00EE/04/E051-8.p',ruleD75)).
fof(f714,plain,(
  ( ! [X6,X10,X8,X7,X11,X9] : (eqratio(X6,X7,X8,X9,X7,X10,X9,X11) | ~midp(X8,X9,X11) | ~midp(X6,X7,X10)) )),
  inference(resolution,[],[f386,f428])).
fof(f428,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (eqratio(X4,X0,X0,X1,X5,X2,X2,X3) | ~midp(X5,X2,X3) | ~midp(X4,X0,X1)) )),
  inference(cnf_transformation,[],[f272])).
fof(f272,plain,(
  ! [X0,X1,X2,X3,X4,X5] : (eqratio(X4,X0,X0,X1,X5,X2,X2,X3) | ~midp(X5,X2,X3) | ~midp(X4,X0,X1))),
  inference(flattening,[],[f271])).
fof(f271,plain,(
  ! [X0,X1,X2,X3,X4,X5] : (eqratio(X4,X0,X0,X1,X5,X2,X2,X3) | (~midp(X5,X2,X3) | ~midp(X4,X0,X1)))),
  inference(ennf_transformation,[],[f145])).
fof(f145,plain,(
  ! [X0,X1,X2,X3,X4,X5] : ((midp(X5,X2,X3) & midp(X4,X0,X1)) => eqratio(X4,X0,X0,X1,X5,X2,X2,X3))),
  inference(rectify,[],[f71])).
fof(f71,axiom,(
  ! [X0,X1,X2,X3,X6,X16] : ((midp(X16,X2,X3) & midp(X6,X0,X1)) => eqratio(X6,X0,X0,X1,X16,X2,X2,X3))),
  file('./7_Book/00EE/04/E051-8.p',ruleD70)).
fof(f386,plain,(
  ( ! [X6,X4,X2,X0,X7,X5,X3,X1] : (~eqratio(X0,X1,X2,X3,X4,X5,X6,X7) | eqratio(X0,X1,X4,X5,X2,X3,X6,X7)) )),
  inference(cnf_transformation,[],[f204])).
fof(f204,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6,X7] : (eqratio(X0,X1,X4,X5,X2,X3,X6,X7) | ~eqratio(X0,X1,X2,X3,X4,X5,X6,X7))),
  inference(ennf_transformation,[],[f108])).
fof(f108,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6,X7] : (eqratio(X0,X1,X2,X3,X4,X5,X6,X7) => eqratio(X0,X1,X4,X5,X2,X3,X6,X7))),
  inference(rectify,[],[f29])).
fof(f29,axiom,(
  ! [X0,X1,X2,X3,X8,X9,X10,X11] : (eqratio(X0,X1,X2,X3,X8,X9,X10,X11) => eqratio(X0,X1,X8,X9,X2,X3,X10,X11))),
  file('./7_Book/00EE/04/E051-8.p',ruleD29)).
fof(f22332,plain,(
  ( ! [X39,X38,X40] : (~cong(X38,X38,X38,X39) | ~cyclic(X38,X39,X38,X40) | ~para(X38,X39,X38,X39) | para(X40,X38,X40,X39)) )),
  inference(resolution,[],[f4120,f431])).
fof(f431,plain,(
  ( ! [X6,X4,X2,X0,X7,X5,X3,X1] : (~eqangle(X0,X1,X2,X3,X4,X5,X6,X7) | ~para(X4,X5,X6,X7) | para(X0,X1,X2,X3)) )),
  inference(cnf_transformation,[],[f278])).
fof(f278,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6,X7] : (para(X0,X1,X2,X3) | ~para(X4,X5,X6,X7) | ~eqangle(X0,X1,X2,X3,X4,X5,X6,X7))),
  inference(flattening,[],[f277])).
fof(f277,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6,X7] : (para(X0,X1,X2,X3) | (~para(X4,X5,X6,X7) | ~eqangle(X0,X1,X2,X3,X4,X5,X6,X7)))),
  inference(ennf_transformation,[],[f146])).
fof(f146,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6,X7] : ((para(X4,X5,X6,X7) & eqangle(X0,X1,X2,X3,X4,X5,X6,X7)) => para(X0,X1,X2,X3))),
  inference(rectify,[],[f74])).
fof(f74,axiom,(
  ! [X0,X1,X2,X3,X8,X9,X10,X11] : ((para(X8,X9,X10,X11) & eqangle(X0,X1,X2,X3,X8,X9,X10,X11)) => para(X0,X1,X2,X3))),
  file('./7_Book/00EE/04/E051-8.p',ruleD73)).
fof(f4120,plain,(
  ( ! [X57,X58,X56] : (eqangle(X56,X57,X56,X58,X57,X58,X57,X58) | ~cong(X57,X57,X57,X58) | ~cyclic(X57,X58,X57,X56)) )),
  inference(resolution,[],[f1188,f682])).
fof(f682,plain,(
  ( ! [X6,X8,X7,X9] : (eqangle(X6,X7,X6,X8,X9,X7,X9,X8) | ~cyclic(X7,X8,X9,X6)) )),
  inference(resolution,[],[f377,f398])).
fof(f377,plain,(
  ( ! [X6,X4,X2,X0,X7,X5,X3,X1] : (~eqangle(X0,X1,X2,X3,X4,X5,X6,X7) | eqangle(X4,X5,X6,X7,X0,X1,X2,X3)) )),
  inference(cnf_transformation,[],[f193])).
fof(f193,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6,X7] : (eqangle(X4,X5,X6,X7,X0,X1,X2,X3) | ~eqangle(X0,X1,X2,X3,X4,X5,X6,X7))),
  inference(ennf_transformation,[],[f102])).
fof(f102,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6,X7] : (eqangle(X0,X1,X2,X3,X4,X5,X6,X7) => eqangle(X4,X5,X6,X7,X0,X1,X2,X3))),
  inference(rectify,[],[f20])).
fof(f20,axiom,(
  ! [X0,X1,X2,X3,X8,X9,X10,X11] : (eqangle(X0,X1,X2,X3,X8,X9,X10,X11) => eqangle(X8,X9,X10,X11,X0,X1,X2,X3))),
  file('./7_Book/00EE/04/E051-8.p',ruleD20)).
fof(f1188,plain,(
  ( ! [X39,X45,X43,X41,X44,X42,X40] : (~eqangle(X39,X40,X41,X42,X45,X43,X43,X44) | eqangle(X39,X40,X41,X42,X43,X44,X45,X44) | ~cong(X45,X43,X45,X44)) )),
  inference(resolution,[],[f379,f404])).
fof(f404,plain,(
  ( ! [X2,X0,X1] : (eqangle(X2,X0,X0,X1,X0,X1,X2,X1) | ~cong(X2,X0,X2,X1)) )),
  inference(cnf_transformation,[],[f230])).
fof(f230,plain,(
  ! [X0,X1,X2] : (eqangle(X2,X0,X0,X1,X0,X1,X2,X1) | ~cong(X2,X0,X2,X1))),
  inference(ennf_transformation,[],[f126])).
fof(f126,plain,(
  ! [X0,X1,X2] : (cong(X2,X0,X2,X1) => eqangle(X2,X0,X0,X1,X0,X1,X2,X1))),
  inference(rectify,[],[f47])).
fof(f47,axiom,(
  ! [X0,X1,X7] : (cong(X7,X0,X7,X1) => eqangle(X7,X0,X0,X1,X0,X1,X7,X1))),
  file('./7_Book/00EE/04/E051-8.p',ruleD46)).
fof(f379,plain,(
  ( ! [X6,X4,X2,X0,X10,X8,X7,X5,X3,X1,X11,X9] : (~eqangle(X4,X5,X6,X7,X8,X9,X10,X11) | eqangle(X0,X1,X2,X3,X8,X9,X10,X11) | ~eqangle(X0,X1,X2,X3,X4,X5,X6,X7)) )),
  inference(cnf_transformation,[],[f196])).
fof(f196,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11] : (eqangle(X0,X1,X2,X3,X8,X9,X10,X11) | ~eqangle(X4,X5,X6,X7,X8,X9,X10,X11) | ~eqangle(X0,X1,X2,X3,X4,X5,X6,X7))),
  inference(flattening,[],[f195])).
fof(f195,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11] : (eqangle(X0,X1,X2,X3,X8,X9,X10,X11) | (~eqangle(X4,X5,X6,X7,X8,X9,X10,X11) | ~eqangle(X0,X1,X2,X3,X4,X5,X6,X7)))),
  inference(ennf_transformation,[],[f104])).
fof(f104,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11] : ((eqangle(X4,X5,X6,X7,X8,X9,X10,X11) & eqangle(X0,X1,X2,X3,X4,X5,X6,X7)) => eqangle(X0,X1,X2,X3,X8,X9,X10,X11))),
  inference(rectify,[],[f22])).
fof(f22,axiom,(
  ! [X0,X1,X2,X3,X8,X9,X10,X11,X4,X5,X12,X13] : ((eqangle(X8,X9,X10,X11,X4,X5,X12,X13) & eqangle(X0,X1,X2,X3,X8,X9,X10,X11)) => eqangle(X0,X1,X2,X3,X4,X5,X12,X13))),
  file('./7_Book/00EE/04/E051-8.p',ruleD22)).
fof(f22408,plain,(
  ( ! [X61,X64,X62,X63] : (midp(X61,X62,X63) | ~para(X64,X62,X64,X63) | ~midp(X61,X64,X64)) ) | ~spl26_148),
  inference(resolution,[],[f22375,f422])).
fof(f15317,plain,(
  spl26_272 | ~spl26_64 | ~spl26_148 | ~spl26_328 | ~spl26_372),
  inference(avatar_split_clause,[],[f15311,f5043,f4692,f1865,f986,f3827])).
fof(f986,plain,(
  spl26_64 <=> para(sK23,sK22,sK22,sK21)),
  introduced(avatar_definition,[new_symbols(naming,[spl26_64])])).
fof(f4692,plain,(
  spl26_328 <=> para(sK22,sK22,sK22,sK23)),
  introduced(avatar_definition,[new_symbols(naming,[spl26_328])])).
fof(f5043,plain,(
  spl26_372 <=> para(sK22,sK23,sK22,sK22)),
  introduced(avatar_definition,[new_symbols(naming,[spl26_372])])).
fof(f15311,plain,(
  para(sK23,sK21,sK23,sK22) | (~spl26_64 | ~spl26_148 | ~spl26_328 | ~spl26_372)),
  inference(resolution,[],[f14927,f1758])).
fof(f1758,plain,(
  ( ! [X2,X0,X3,X1] : (~para(X0,X1,X2,X3) | para(X1,X0,X2,X3)) )),
  inference(resolution,[],[f665,f396])).
fof(f396,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (~eqangle(X0,X1,X4,X5,X2,X3,X4,X5) | para(X0,X1,X2,X3)) )),
  inference(cnf_transformation,[],[f217])).
fof(f217,plain,(
  ! [X0,X1,X2,X3,X4,X5] : (para(X0,X1,X2,X3) | ~eqangle(X0,X1,X4,X5,X2,X3,X4,X5))),
  inference(ennf_transformation,[],[f118])).
fof(f118,plain,(
  ! [X0,X1,X2,X3,X4,X5] : (eqangle(X0,X1,X4,X5,X2,X3,X4,X5) => para(X0,X1,X2,X3))),
  inference(rectify,[],[f39])).
fof(f39,axiom,(
  ! [X0,X1,X2,X3,X8,X9] : (eqangle(X0,X1,X8,X9,X2,X3,X8,X9) => para(X0,X1,X2,X3))),
  file('./7_Book/00EE/04/E051-8.p',ruleD39)).
fof(f665,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (eqangle(X0,X1,X2,X3,X4,X5,X2,X3) | ~para(X1,X0,X4,X5)) )),
  inference(resolution,[],[f375,f397])).
fof(f375,plain,(
  ( ! [X6,X4,X2,X0,X7,X5,X3,X1] : (~eqangle(X0,X1,X2,X3,X4,X5,X6,X7) | eqangle(X1,X0,X2,X3,X4,X5,X6,X7)) )),
  inference(cnf_transformation,[],[f191])).
fof(f191,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6,X7] : (eqangle(X1,X0,X2,X3,X4,X5,X6,X7) | ~eqangle(X0,X1,X2,X3,X4,X5,X6,X7))),
  inference(ennf_transformation,[],[f100])).
fof(f100,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6,X7] : (eqangle(X0,X1,X2,X3,X4,X5,X6,X7) => eqangle(X1,X0,X2,X3,X4,X5,X6,X7))),
  inference(rectify,[],[f18])).
fof(f18,axiom,(
  ! [X0,X1,X2,X3,X8,X9,X10,X11] : (eqangle(X0,X1,X2,X3,X8,X9,X10,X11) => eqangle(X1,X0,X2,X3,X8,X9,X10,X11))),
  file('./7_Book/00EE/04/E051-8.p',ruleD18)).
fof(f14927,plain,(
  para(sK21,sK23,sK23,sK22) | (~spl26_64 | ~spl26_148 | ~spl26_328 | ~spl26_372)),
  inference(subsumption_resolution,[],[f9856,f14834])).
fof(f14834,plain,(
  ( ! [X4] : (midp(X4,sK23,sK22)) ) | (~spl26_148 | ~spl26_328 | ~spl26_372)),
  inference(subsumption_resolution,[],[f14833,f12431])).
fof(f14833,plain,(
  ( ! [X4] : (midp(X4,sK23,sK22) | ~midp(X4,sK22,sK22)) ) | (~spl26_328 | ~spl26_372)),
  inference(subsumption_resolution,[],[f14826,f5044])).
fof(f5044,plain,(
  para(sK22,sK23,sK22,sK22) | ~spl26_372),
  inference(avatar_component_clause,[],[f5043])).
fof(f14826,plain,(
  ( ! [X4] : (midp(X4,sK23,sK22) | ~para(sK22,sK23,sK22,sK22) | ~midp(X4,sK22,sK22)) ) | ~spl26_328),
  inference(resolution,[],[f4693,f422])).
fof(f4693,plain,(
  para(sK22,sK22,sK22,sK23) | ~spl26_328),
  inference(avatar_component_clause,[],[f4692])).
fof(f9856,plain,(
  para(sK21,sK23,sK23,sK22) | ~midp(sK21,sK23,sK22) | ~spl26_64),
  inference(resolution,[],[f2131,f8542])).
fof(f8542,plain,(
  para(sK23,sK22,sK21,sK22) | ~spl26_64),
  inference(resolution,[],[f987,f361])).
fof(f361,plain,(
  ( ! [X2,X0,X3,X1] : (~para(X0,X1,X2,X3) | para(X0,X1,X3,X2)) )),
  inference(cnf_transformation,[],[f171])).
fof(f171,plain,(
  ! [X0,X1,X2,X3] : (para(X0,X1,X3,X2) | ~para(X0,X1,X2,X3))),
  inference(ennf_transformation,[],[f4])).
fof(f4,axiom,(
  ! [X0,X1,X2,X3] : (para(X0,X1,X2,X3) => para(X0,X1,X3,X2))),
  file('./7_Book/00EE/04/E051-8.p',ruleD4)).
fof(f987,plain,(
  para(sK23,sK22,sK22,sK21) | ~spl26_64),
  inference(avatar_component_clause,[],[f986])).
fof(f2131,plain,(
  ( ! [X4,X5,X3] : (~para(X4,X5,X3,X5) | para(X3,X4,X4,X5) | ~midp(X3,X4,X5)) )),
  inference(resolution,[],[f980,f426])).
fof(f426,plain,(
  ( ! [X2,X0,X1] : (cong(X0,X1,X0,X2) | ~midp(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f269])).
fof(f269,plain,(
  ! [X0,X1,X2] : (cong(X0,X1,X0,X2) | ~midp(X0,X1,X2))),
  inference(ennf_transformation,[],[f69])).
fof(f69,axiom,(
  ! [X0,X1,X2] : (midp(X0,X1,X2) => cong(X0,X1,X0,X2))),
  file('./7_Book/00EE/04/E051-8.p',ruleD68)).
fof(f980,plain,(
  ( ! [X21,X19,X20] : (~cong(X21,X19,X21,X20) | para(X21,X19,X19,X20) | ~para(X19,X20,X21,X20)) )),
  inference(resolution,[],[f431,f404])).
fof(f14577,plain,(
  ~spl26_66 | ~spl26_148 | spl26_407 | ~spl26_732),
  inference(avatar_contradiction_clause,[],[f14576])).
fof(f14576,plain,(
  $false | (~spl26_66 | ~spl26_148 | ~spl26_407 | ~spl26_732)),
  inference(subsumption_resolution,[],[f14575,f12595])).
fof(f12595,plain,(
  midp(sK23,sK21,sK22) | ~spl26_148),
  inference(subsumption_resolution,[],[f12473,f2417])).
fof(f2417,plain,(
  coll(sK23,sK21,sK22)),
  inference(resolution,[],[f2268,f358])).
fof(f358,plain,(
  ( ! [X2,X0,X1] : (~coll(X0,X1,X2) | coll(X0,X2,X1)) )),
  inference(cnf_transformation,[],[f167])).
fof(f167,plain,(
  ! [X0,X1,X2] : (coll(X0,X2,X1) | ~coll(X0,X1,X2))),
  inference(ennf_transformation,[],[f1])).
fof(f1,axiom,(
  ! [X0,X1,X2] : (coll(X0,X1,X2) => coll(X0,X2,X1))),
  file('./7_Book/00EE/04/E051-8.p',ruleD1)).
fof(f2268,plain,(
  coll(sK23,sK22,sK21)),
  inference(resolution,[],[f2102,f536])).
fof(f536,plain,(
  ( ! [X16] : (~coll(sK21,sK25,X16) | coll(X16,sK22,sK21)) )),
  inference(resolution,[],[f360,f483])).
fof(f483,plain,(
  coll(sK21,sK25,sK22)),
  inference(resolution,[],[f359,f476])).
fof(f476,plain,(
  coll(sK25,sK21,sK22)),
  inference(cnf_transformation,[],[f357])).
fof(f357,plain,(
  ~cong(sK22,sK24,sK25,sK21) & coll(sK25,sK21,sK22) & para(sK20,sK21,sK25,sK24) & coll(sK24,sK20,sK22) & para(sK21,sK22,sK24,sK23) & coll(sK23,sK20,sK21) & eqangle(sK23,sK22,sK22,sK21,sK23,sK22,sK22,sK20)),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK20,sK21,sK22,sK23,sK24,sK25])],[f319,f356])).
fof(f356,plain,(
  ? [X0,X1,X2,X3,X4,X5] : (~cong(X2,X4,X5,X1) & coll(X5,X1,X2) & para(X0,X1,X5,X4) & coll(X4,X0,X2) & para(X1,X2,X4,X3) & coll(X3,X0,X1) & eqangle(X3,X2,X2,X1,X3,X2,X2,X0)) => (~cong(sK22,sK24,sK25,sK21) & coll(sK25,sK21,sK22) & para(sK20,sK21,sK25,sK24) & coll(sK24,sK20,sK22) & para(sK21,sK22,sK24,sK23) & coll(sK23,sK20,sK21) & eqangle(sK23,sK22,sK22,sK21,sK23,sK22,sK22,sK20))),
  introduced(choice_axiom,[])).
fof(f319,plain,(
  ? [X0,X1,X2,X3,X4,X5] : (~cong(X2,X4,X5,X1) & coll(X5,X1,X2) & para(X0,X1,X5,X4) & coll(X4,X0,X2) & para(X1,X2,X4,X3) & coll(X3,X0,X1) & eqangle(X3,X2,X2,X1,X3,X2,X2,X0))),
  inference(flattening,[],[f318])).
fof(f318,plain,(
  ? [X0,X1,X2,X3,X4,X5] : (~cong(X2,X4,X5,X1) & (coll(X5,X1,X2) & para(X0,X1,X5,X4) & coll(X4,X0,X2) & para(X1,X2,X4,X3) & coll(X3,X0,X1) & eqangle(X3,X2,X2,X1,X3,X2,X2,X0)))),
  inference(ennf_transformation,[],[f96])).
fof(f96,negated_conjecture,(
  ~! [X0,X1,X2,X3,X4,X5] : ((coll(X5,X1,X2) & para(X0,X1,X5,X4) & coll(X4,X0,X2) & para(X1,X2,X4,X3) & coll(X3,X0,X1) & eqangle(X3,X2,X2,X1,X3,X2,X2,X0)) => cong(X2,X4,X5,X1))),
  inference(negated_conjecture,[],[f95])).
fof(f95,conjecture,(
  ! [X0,X1,X2,X3,X4,X5] : ((coll(X5,X1,X2) & para(X0,X1,X5,X4) & coll(X4,X0,X2) & para(X1,X2,X4,X3) & coll(X3,X0,X1) & eqangle(X3,X2,X2,X1,X3,X2,X2,X0)) => cong(X2,X4,X5,X1))),
  file('./7_Book/00EE/04/E051-8.p',exemplo7Book00EE04E0518)).
fof(f359,plain,(
  ( ! [X2,X0,X1] : (~coll(X0,X1,X2) | coll(X1,X0,X2)) )),
  inference(cnf_transformation,[],[f168])).
fof(f168,plain,(
  ! [X0,X1,X2] : (coll(X1,X0,X2) | ~coll(X0,X1,X2))),
  inference(ennf_transformation,[],[f2])).
fof(f2,axiom,(
  ! [X0,X1,X2] : (coll(X0,X1,X2) => coll(X1,X0,X2))),
  file('./7_Book/00EE/04/E051-8.p',ruleD2)).
fof(f360,plain,(
  ( ! [X2,X0,X3,X1] : (~coll(X0,X1,X3) | coll(X2,X3,X0) | ~coll(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f170])).
fof(f170,plain,(
  ! [X0,X1,X2,X3] : (coll(X2,X3,X0) | ~coll(X0,X1,X3) | ~coll(X0,X1,X2))),
  inference(flattening,[],[f169])).
fof(f169,plain,(
  ! [X0,X1,X2,X3] : (coll(X2,X3,X0) | (~coll(X0,X1,X3) | ~coll(X0,X1,X2)))),
  inference(ennf_transformation,[],[f3])).
fof(f3,axiom,(
  ! [X0,X1,X2,X3] : ((coll(X0,X1,X3) & coll(X0,X1,X2)) => coll(X2,X3,X0))),
  file('./7_Book/00EE/04/E051-8.p',ruleD3)).
fof(f2102,plain,(
  coll(sK21,sK25,sK23)),
  inference(resolution,[],[f1899,f359])).
fof(f1899,plain,(
  coll(sK25,sK21,sK23)),
  inference(resolution,[],[f1722,f522])).
fof(f522,plain,(
  ( ! [X2] : (~coll(sK23,sK20,X2) | coll(X2,sK21,sK23)) )),
  inference(resolution,[],[f360,f472])).
fof(f472,plain,(
  coll(sK23,sK20,sK21)),
  inference(cnf_transformation,[],[f357])).
fof(f1722,plain,(
  coll(sK23,sK20,sK25)),
  inference(resolution,[],[f1513,f358])).
fof(f1513,plain,(
  coll(sK23,sK25,sK20)),
  inference(resolution,[],[f1488,f359])).
fof(f1488,plain,(
  coll(sK25,sK23,sK20)),
  inference(resolution,[],[f1447,f524])).
fof(f524,plain,(
  ( ! [X4] : (~coll(sK20,sK21,X4) | coll(X4,sK23,sK20)) )),
  inference(resolution,[],[f360,f491])).
fof(f491,plain,(
  coll(sK20,sK21,sK23)),
  inference(resolution,[],[f481,f358])).
fof(f481,plain,(
  coll(sK20,sK23,sK21)),
  inference(resolution,[],[f359,f472])).
fof(f1447,plain,(
  coll(sK20,sK21,sK25)),
  inference(resolution,[],[f1380,f358])).
fof(f1380,plain,(
  coll(sK20,sK25,sK21)),
  inference(resolution,[],[f1378,f529])).
fof(f529,plain,(
  ( ! [X9] : (~coll(sK21,sK22,X9) | coll(X9,sK25,sK21)) )),
  inference(resolution,[],[f360,f495])).
fof(f495,plain,(
  coll(sK21,sK22,sK25)),
  inference(resolution,[],[f483,f358])).
fof(f1378,plain,(
  coll(sK21,sK22,sK20)),
  inference(resolution,[],[f1351,f359])).
fof(f1351,plain,(
  coll(sK22,sK21,sK20)),
  inference(resolution,[],[f1317,f424])).
fof(f1317,plain,(
  para(sK22,sK21,sK22,sK20)),
  inference(resolution,[],[f672,f396])).
fof(f672,plain,(
  eqangle(sK22,sK21,sK23,sK22,sK22,sK20,sK23,sK22)),
  inference(resolution,[],[f376,f471])).
fof(f471,plain,(
  eqangle(sK23,sK22,sK22,sK21,sK23,sK22,sK22,sK20)),
  inference(cnf_transformation,[],[f357])).
fof(f12473,plain,(
  midp(sK23,sK21,sK22) | ~coll(sK23,sK21,sK22) | ~spl26_148),
  inference(resolution,[],[f12431,f641])).
fof(f641,plain,(
  ( ! [X10] : (~midp(sK24,X10,sK21) | midp(sK23,X10,sK22) | ~coll(sK23,X10,sK22)) )),
  inference(resolution,[],[f403,f500])).
fof(f500,plain,(
  para(sK24,sK23,sK21,sK22)),
  inference(resolution,[],[f362,f473])).
fof(f473,plain,(
  para(sK21,sK22,sK24,sK23)),
  inference(cnf_transformation,[],[f357])).
fof(f362,plain,(
  ( ! [X2,X0,X3,X1] : (~para(X0,X1,X2,X3) | para(X2,X3,X0,X1)) )),
  inference(cnf_transformation,[],[f172])).
fof(f172,plain,(
  ! [X0,X1,X2,X3] : (para(X2,X3,X0,X1) | ~para(X0,X1,X2,X3))),
  inference(ennf_transformation,[],[f5])).
fof(f5,axiom,(
  ! [X0,X1,X2,X3] : (para(X0,X1,X2,X3) => para(X2,X3,X0,X1))),
  file('./7_Book/00EE/04/E051-8.p',ruleD5)).
fof(f403,plain,(
  ( ! [X4,X2,X0,X3,X1] : (~para(X3,X4,X1,X2) | ~coll(X4,X0,X2) | midp(X4,X0,X2) | ~midp(X3,X0,X1)) )),
  inference(cnf_transformation,[],[f229])).
fof(f229,plain,(
  ! [X0,X1,X2,X3,X4] : (midp(X4,X0,X2) | ~coll(X4,X0,X2) | ~para(X3,X4,X1,X2) | ~midp(X3,X0,X1))),
  inference(flattening,[],[f228])).
fof(f228,plain,(
  ! [X0,X1,X2,X3,X4] : (midp(X4,X0,X2) | (~coll(X4,X0,X2) | ~para(X3,X4,X1,X2) | ~midp(X3,X0,X1)))),
  inference(ennf_transformation,[],[f125])).
fof(f125,plain,(
  ! [X0,X1,X2,X3,X4] : ((coll(X4,X0,X2) & para(X3,X4,X1,X2) & midp(X3,X0,X1)) => midp(X4,X0,X2))),
  inference(rectify,[],[f46])).
fof(f46,axiom,(
  ! [X0,X1,X2,X4,X5] : ((coll(X5,X0,X2) & para(X4,X5,X1,X2) & midp(X4,X0,X1)) => midp(X5,X0,X2))),
  file('./7_Book/00EE/04/E051-8.p',ruleD45)).
fof(f14575,plain,(
  ~midp(sK23,sK21,sK22) | (~spl26_66 | ~spl26_148 | ~spl26_407 | ~spl26_732)),
  inference(subsumption_resolution,[],[f14571,f12681])).
fof(f12681,plain,(
  midp(sK23,sK22,sK20) | ~spl26_732),
  inference(avatar_component_clause,[],[f12680])).
fof(f12680,plain,(
  spl26_732 <=> midp(sK23,sK22,sK20)),
  introduced(avatar_definition,[new_symbols(naming,[spl26_732])])).
fof(f14571,plain,(
  ~midp(sK23,sK22,sK20) | ~midp(sK23,sK21,sK22) | (~spl26_66 | ~spl26_148 | ~spl26_407)),
  inference(resolution,[],[f13922,f1734])).
fof(f1734,plain,(
  ( ! [X2,X0,X3,X1] : (circle(X0,X1,X2,X3) | ~midp(X0,X1,X3) | ~midp(X0,X2,X1)) )),
  inference(resolution,[],[f612,f517])).
fof(f517,plain,(
  ( ! [X2,X0,X1] : (cong(X0,X1,X0,X2) | ~midp(X0,X2,X1)) )),
  inference(resolution,[],[f381,f426])).
fof(f381,plain,(
  ( ! [X2,X0,X3,X1] : (~cong(X0,X1,X2,X3) | cong(X2,X3,X0,X1)) )),
  inference(cnf_transformation,[],[f198])).
fof(f198,plain,(
  ! [X0,X1,X2,X3] : (cong(X2,X3,X0,X1) | ~cong(X0,X1,X2,X3))),
  inference(ennf_transformation,[],[f24])).
fof(f24,axiom,(
  ! [X0,X1,X2,X3] : (cong(X0,X1,X2,X3) => cong(X2,X3,X0,X1))),
  file('./7_Book/00EE/04/E051-8.p',ruleD24)).
fof(f612,plain,(
  ( ! [X2,X0,X3,X1] : (~cong(X0,X1,X0,X2) | circle(X0,X1,X2,X3) | ~midp(X0,X1,X3)) )),
  inference(resolution,[],[f369,f426])).
fof(f369,plain,(
  ( ! [X2,X0,X3,X1] : (~cong(X3,X0,X3,X2) | circle(X3,X0,X1,X2) | ~cong(X3,X0,X3,X1)) )),
  inference(cnf_transformation,[],[f183])).
fof(f183,plain,(
  ! [X0,X1,X2,X3] : (circle(X3,X0,X1,X2) | ~cong(X3,X0,X3,X2) | ~cong(X3,X0,X3,X1))),
  inference(flattening,[],[f182])).
fof(f182,plain,(
  ! [X0,X1,X2,X3] : (circle(X3,X0,X1,X2) | (~cong(X3,X0,X3,X2) | ~cong(X3,X0,X3,X1)))),
  inference(ennf_transformation,[],[f98])).
fof(f98,plain,(
  ! [X0,X1,X2,X3] : ((cong(X3,X0,X3,X2) & cong(X3,X0,X3,X1)) => circle(X3,X0,X1,X2))),
  inference(rectify,[],[f12])).
fof(f12,axiom,(
  ! [X0,X1,X2,X7] : ((cong(X7,X0,X7,X2) & cong(X7,X0,X7,X1)) => circle(X7,X0,X1,X2))),
  file('./7_Book/00EE/04/E051-8.p',ruleD12)).
fof(f13922,plain,(
  ~circle(sK23,sK22,sK21,sK20) | (~spl26_66 | ~spl26_148 | ~spl26_407)),
  inference(subsumption_resolution,[],[f13921,f5546])).
fof(f5546,plain,(
  ~para(sK23,sK22,sK23,sK21) | ~spl26_407),
  inference(avatar_component_clause,[],[f5545])).
fof(f5545,plain,(
  spl26_407 <=> ~para(sK23,sK22,sK23,sK21)),
  introduced(avatar_definition,[new_symbols(naming,[spl26_407])])).
fof(f13921,plain,(
  ~circle(sK23,sK22,sK21,sK20) | para(sK23,sK22,sK23,sK21) | (~spl26_66 | ~spl26_148)),
  inference(subsumption_resolution,[],[f13904,f13333])).
fof(f13333,plain,(
  midp(sK22,sK21,sK20) | (~spl26_66 | ~spl26_148)),
  inference(resolution,[],[f12595,f8592])).
fof(f8592,plain,(
  ( ! [X2] : (~midp(sK23,X2,sK22) | midp(sK22,X2,sK20)) ) | (~spl26_66 | ~spl26_148)),
  inference(subsumption_resolution,[],[f8588,f4473])).
fof(f4473,plain,(
  ( ! [X2,X0,X1] : (coll(X0,X1,X2)) ) | ~spl26_148),
  inference(subsumption_resolution,[],[f4470,f4440])).
fof(f4440,plain,(
  ( ! [X6,X5] : (coll(X5,X5,X6)) ) | ~spl26_148),
  inference(resolution,[],[f4382,f358])).
fof(f4382,plain,(
  ( ! [X4,X3] : (coll(X3,X4,X3)) ) | ~spl26_148),
  inference(resolution,[],[f4334,f359])).
fof(f4470,plain,(
  ( ! [X2,X0,X1] : (coll(X0,X1,X2) | ~coll(X2,X2,X0)) ) | ~spl26_148),
  inference(resolution,[],[f4440,f360])).
fof(f8588,plain,(
  ( ! [X2] : (~coll(sK22,X2,sK20) | midp(sK22,X2,sK20) | ~midp(sK23,X2,sK22)) ) | ~spl26_66),
  inference(resolution,[],[f990,f403])).
fof(f990,plain,(
  para(sK23,sK22,sK22,sK20) | ~spl26_66),
  inference(avatar_component_clause,[],[f989])).
fof(f989,plain,(
  spl26_66 <=> para(sK23,sK22,sK22,sK20)),
  introduced(avatar_definition,[new_symbols(naming,[spl26_66])])).
fof(f13904,plain,(
  ~midp(sK22,sK21,sK20) | ~circle(sK23,sK22,sK21,sK20) | para(sK23,sK22,sK23,sK21)),
  inference(resolution,[],[f4577,f396])).
fof(f4577,plain,(
  ( ! [X0,X1] : (eqangle(sK23,sK22,sK23,sK22,X0,sK21,X0,X1) | ~midp(X1,sK21,sK20) | ~circle(X0,sK22,sK21,sK20)) )),
  inference(resolution,[],[f1187,f689])).
fof(f689,plain,(
  eqangle(sK23,sK22,sK23,sK22,sK22,sK21,sK22,sK20)),
  inference(resolution,[],[f378,f471])).
fof(f1187,plain,(
  ( ! [X30,X37,X35,X33,X31,X38,X36,X34,X32] : (~eqangle(X30,X31,X32,X33,X37,X35,X37,X38) | eqangle(X30,X31,X32,X33,X34,X35,X34,X36) | ~midp(X36,X35,X38) | ~circle(X34,X37,X35,X38)) )),
  inference(resolution,[],[f379,f408])).
fof(f408,plain,(
  ( ! [X4,X2,X0,X3,X1] : (eqangle(X0,X1,X0,X2,X3,X1,X3,X4) | ~midp(X4,X1,X2) | ~circle(X3,X0,X1,X2)) )),
  inference(cnf_transformation,[],[f238])).
fof(f238,plain,(
  ! [X0,X1,X2,X3,X4] : (eqangle(X0,X1,X0,X2,X3,X1,X3,X4) | ~midp(X4,X1,X2) | ~circle(X3,X0,X1,X2))),
  inference(flattening,[],[f237])).
fof(f237,plain,(
  ! [X0,X1,X2,X3,X4] : (eqangle(X0,X1,X0,X2,X3,X1,X3,X4) | (~midp(X4,X1,X2) | ~circle(X3,X0,X1,X2)))),
  inference(ennf_transformation,[],[f130])).
fof(f130,plain,(
  ! [X0,X1,X2,X3,X4] : ((midp(X4,X1,X2) & circle(X3,X0,X1,X2)) => eqangle(X0,X1,X0,X2,X3,X1,X3,X4))),
  inference(rectify,[],[f51])).
fof(f51,axiom,(
  ! [X0,X1,X2,X7,X6] : ((midp(X6,X1,X2) & circle(X7,X0,X1,X2)) => eqangle(X0,X1,X0,X2,X7,X1,X7,X6))),
  file('./7_Book/00EE/04/E051-8.p',ruleD50)).
fof(f13876,plain,(
  ~spl26_455 | spl26_329),
  inference(avatar_split_clause,[],[f6341,f4689,f5733])).
fof(f5733,plain,(
  spl26_455 <=> ~midp(sK22,sK23,sK22)),
  introduced(avatar_definition,[new_symbols(naming,[spl26_455])])).
fof(f4689,plain,(
  spl26_329 <=> ~para(sK22,sK22,sK22,sK23)),
  introduced(avatar_definition,[new_symbols(naming,[spl26_329])])).
fof(f6341,plain,(
  ~midp(sK22,sK23,sK22) | ~spl26_329),
  inference(resolution,[],[f1496,f4690])).
fof(f4690,plain,(
  ~para(sK22,sK22,sK22,sK23) | ~spl26_329),
  inference(avatar_component_clause,[],[f4689])).
fof(f1496,plain,(
  ( ! [X0,X1] : (para(X0,X0,X0,X1) | ~midp(X0,X1,X0)) )),
  inference(resolution,[],[f569,f517])).
fof(f569,plain,(
  ( ! [X0,X1] : (~cong(X0,X0,X0,X1) | para(X0,X0,X0,X1)) )),
  inference(resolution,[],[f404,f396])).
fof(f13856,plain,(
  spl26_454 | ~spl26_530),
  inference(avatar_split_clause,[],[f13851,f6971,f5736])).
fof(f5736,plain,(
  spl26_454 <=> midp(sK22,sK23,sK22)),
  introduced(avatar_definition,[new_symbols(naming,[spl26_454])])).
fof(f6971,plain,(
  spl26_530 <=> cong(sK22,sK23,sK22,sK22)),
  introduced(avatar_definition,[new_symbols(naming,[spl26_530])])).
fof(f13851,plain,(
  midp(sK22,sK23,sK22) | ~spl26_530),
  inference(subsumption_resolution,[],[f13831,f2752])).
fof(f2752,plain,(
  coll(sK22,sK23,sK22)),
  inference(resolution,[],[f2729,f359])).
fof(f2729,plain,(
  coll(sK23,sK22,sK22)),
  inference(resolution,[],[f1392,f424])).
fof(f1392,plain,(
  para(sK23,sK22,sK23,sK22)),
  inference(subsumption_resolution,[],[f1389,f1317])).
fof(f1389,plain,(
  ~para(sK22,sK21,sK22,sK20) | para(sK23,sK22,sK23,sK22)),
  inference(resolution,[],[f689,f431])).
fof(f13831,plain,(
  ~coll(sK22,sK23,sK22) | midp(sK22,sK23,sK22) | ~spl26_530),
  inference(resolution,[],[f6972,f425])).
fof(f6972,plain,(
  cong(sK22,sK23,sK22,sK22) | ~spl26_530),
  inference(avatar_component_clause,[],[f6971])).
fof(f13805,plain,(
  ~spl26_148 | ~spl26_156 | spl26_531),
  inference(avatar_contradiction_clause,[],[f13804])).
fof(f13804,plain,(
  $false | (~spl26_148 | ~spl26_156 | ~spl26_531)),
  inference(subsumption_resolution,[],[f13797,f12431])).
fof(f13797,plain,(
  ~midp(sK23,sK22,sK22) | (~spl26_156 | ~spl26_531)),
  inference(resolution,[],[f6969,f2343])).
fof(f2343,plain,(
  ( ! [X14,X15,X13] : (cong(X14,X13,X14,X15) | ~midp(X13,X14,X15)) ) | ~spl26_156),
  inference(avatar_component_clause,[],[f2342])).
fof(f2342,plain,(
  spl26_156 <=> ! [X13,X15,X14] : (~midp(X13,X14,X15) | cong(X14,X13,X14,X15))),
  introduced(avatar_definition,[new_symbols(naming,[spl26_156])])).
fof(f6969,plain,(
  ~cong(sK22,sK23,sK22,sK22) | ~spl26_531),
  inference(avatar_component_clause,[],[f6968])).
fof(f6968,plain,(
  spl26_531 <=> ~cong(sK22,sK23,sK22,sK22)),
  introduced(avatar_definition,[new_symbols(naming,[spl26_531])])).
fof(f13702,plain,(
  spl26_346 | ~spl26_78 | ~spl26_712),
  inference(avatar_split_clause,[],[f13695,f12553,f1332,f4854])).
fof(f4854,plain,(
  spl26_346 <=> para(sK20,sK23,sK20,sK22)),
  introduced(avatar_definition,[new_symbols(naming,[spl26_346])])).
fof(f1332,plain,(
  spl26_78 <=> para(sK22,sK20,sK23,sK22)),
  introduced(avatar_definition,[new_symbols(naming,[spl26_78])])).
fof(f12553,plain,(
  spl26_712 <=> cong(sK23,sK20,sK23,sK22)),
  introduced(avatar_definition,[new_symbols(naming,[spl26_712])])).
fof(f13695,plain,(
  para(sK20,sK23,sK20,sK22) | (~spl26_78 | ~spl26_712)),
  inference(subsumption_resolution,[],[f13679,f8620])).
fof(f8620,plain,(
  para(sK20,sK22,sK23,sK22) | ~spl26_78),
  inference(resolution,[],[f1333,f1758])).
fof(f1333,plain,(
  para(sK22,sK20,sK23,sK22) | ~spl26_78),
  inference(avatar_component_clause,[],[f1332])).
fof(f13679,plain,(
  ~para(sK20,sK22,sK23,sK22) | para(sK20,sK23,sK20,sK22) | ~spl26_712),
  inference(resolution,[],[f12554,f1818])).
fof(f1818,plain,(
  ( ! [X23,X21,X22] : (~cong(X21,X22,X21,X23) | ~para(X22,X23,X21,X23) | para(X22,X21,X22,X23)) )),
  inference(resolution,[],[f667,f431])).
fof(f667,plain,(
  ( ! [X12,X10,X11] : (eqangle(X10,X11,X10,X12,X10,X12,X11,X12) | ~cong(X11,X10,X11,X12)) )),
  inference(resolution,[],[f375,f404])).
fof(f12554,plain,(
  cong(sK23,sK20,sK23,sK22) | ~spl26_712),
  inference(avatar_component_clause,[],[f12553])).
fof(f13233,plain,(
  spl26_706 | ~spl26_66 | ~spl26_148),
  inference(avatar_split_clause,[],[f13219,f1865,f989,f12517])).
fof(f12517,plain,(
  spl26_706 <=> midp(sK22,sK20,sK22)),
  introduced(avatar_definition,[new_symbols(naming,[spl26_706])])).
fof(f13219,plain,(
  midp(sK22,sK20,sK22) | (~spl26_66 | ~spl26_148)),
  inference(resolution,[],[f12468,f368])).
fof(f368,plain,(
  ( ! [X2,X0,X1] : (~midp(X2,X1,X0) | midp(X2,X0,X1)) )),
  inference(cnf_transformation,[],[f181])).
fof(f181,plain,(
  ! [X0,X1,X2] : (midp(X2,X0,X1) | ~midp(X2,X1,X0))),
  inference(ennf_transformation,[],[f97])).
fof(f97,plain,(
  ! [X0,X1,X2] : (midp(X2,X1,X0) => midp(X2,X0,X1))),
  inference(rectify,[],[f11])).
fof(f11,axiom,(
  ! [X0,X1,X6] : (midp(X6,X1,X0) => midp(X6,X0,X1))),
  file('./7_Book/00EE/04/E051-8.p',ruleD11)).
fof(f12468,plain,(
  midp(sK22,sK22,sK20) | (~spl26_66 | ~spl26_148)),
  inference(resolution,[],[f12431,f8592])).
fof(f12889,plain,(
  spl26_732 | ~spl26_296),
  inference(avatar_split_clause,[],[f12872,f4138,f12680])).
fof(f4138,plain,(
  spl26_296 <=> cong(sK23,sK22,sK23,sK20)),
  introduced(avatar_definition,[new_symbols(naming,[spl26_296])])).
fof(f12872,plain,(
  midp(sK23,sK22,sK20) | ~spl26_296),
  inference(subsumption_resolution,[],[f12851,f1519])).
fof(f1519,plain,(
  coll(sK23,sK22,sK20)),
  inference(resolution,[],[f1492,f359])).
fof(f1492,plain,(
  coll(sK22,sK23,sK20)),
  inference(resolution,[],[f1460,f524])).
fof(f1460,plain,(
  coll(sK20,sK21,sK22)),
  inference(resolution,[],[f1383,f359])).
fof(f1383,plain,(
  coll(sK21,sK20,sK22)),
  inference(resolution,[],[f1378,f358])).
fof(f12851,plain,(
  ~coll(sK23,sK22,sK20) | midp(sK23,sK22,sK20) | ~spl26_296),
  inference(resolution,[],[f4139,f425])).
fof(f4139,plain,(
  cong(sK23,sK22,sK23,sK20) | ~spl26_296),
  inference(avatar_component_clause,[],[f4138])).
fof(f12678,plain,(
  spl26_296 | ~spl26_66 | ~spl26_70 | ~spl26_148),
  inference(avatar_split_clause,[],[f12677,f1865,f1009,f989,f4138])).
fof(f1009,plain,(
  spl26_70 <=> perp(sK23,sK22,sK22,sK20)),
  introduced(avatar_definition,[new_symbols(naming,[spl26_70])])).
fof(f12677,plain,(
  cong(sK23,sK22,sK23,sK20) | (~spl26_66 | ~spl26_70 | ~spl26_148)),
  inference(subsumption_resolution,[],[f6952,f12468])).
fof(f6952,plain,(
  cong(sK23,sK22,sK23,sK20) | ~midp(sK22,sK22,sK20) | ~spl26_70),
  inference(resolution,[],[f1010,f413])).
fof(f413,plain,(
  ( ! [X2,X0,X3,X1] : (~perp(X3,X2,X0,X1) | cong(X3,X0,X3,X1) | ~midp(X2,X0,X1)) )),
  inference(cnf_transformation,[],[f248])).
fof(f248,plain,(
  ! [X0,X1,X2,X3] : (cong(X3,X0,X3,X1) | ~perp(X3,X2,X0,X1) | ~midp(X2,X0,X1))),
  inference(flattening,[],[f247])).
fof(f247,plain,(
  ! [X0,X1,X2,X3] : (cong(X3,X0,X3,X1) | (~perp(X3,X2,X0,X1) | ~midp(X2,X0,X1)))),
  inference(ennf_transformation,[],[f134])).
fof(f134,plain,(
  ! [X0,X1,X2,X3] : ((perp(X3,X2,X0,X1) & midp(X2,X0,X1)) => cong(X3,X0,X3,X1))),
  inference(rectify,[],[f56])).
fof(f56,axiom,(
  ! [X0,X1,X6,X7] : ((perp(X7,X6,X0,X1) & midp(X6,X0,X1)) => cong(X7,X0,X7,X1))),
  file('./7_Book/00EE/04/E051-8.p',ruleD55)).
fof(f1010,plain,(
  perp(sK23,sK22,sK22,sK20) | ~spl26_70),
  inference(avatar_component_clause,[],[f1009])).
fof(f12555,plain,(
  ~spl26_707 | spl26_712 | ~spl26_70),
  inference(avatar_split_clause,[],[f6996,f1009,f12553,f12520])).
fof(f12520,plain,(
  spl26_707 <=> ~midp(sK22,sK20,sK22)),
  introduced(avatar_definition,[new_symbols(naming,[spl26_707])])).
fof(f6996,plain,(
  cong(sK23,sK20,sK23,sK22) | ~midp(sK22,sK20,sK22) | ~spl26_70),
  inference(resolution,[],[f6954,f413])).
fof(f6954,plain,(
  perp(sK23,sK22,sK20,sK22) | ~spl26_70),
  inference(resolution,[],[f1010,f364])).
fof(f364,plain,(
  ( ! [X2,X0,X3,X1] : (~perp(X0,X1,X2,X3) | perp(X0,X1,X3,X2)) )),
  inference(cnf_transformation,[],[f175])).
fof(f175,plain,(
  ! [X0,X1,X2,X3] : (perp(X0,X1,X3,X2) | ~perp(X0,X1,X2,X3))),
  inference(ennf_transformation,[],[f7])).
fof(f7,axiom,(
  ! [X0,X1,X2,X3] : (perp(X0,X1,X2,X3) => perp(X0,X1,X3,X2))),
  file('./7_Book/00EE/04/E051-8.p',ruleD7)).
fof(f8567,plain,(
  spl26_78 | ~spl26_77),
  inference(avatar_split_clause,[],[f6170,f1326,f1332])).
fof(f1326,plain,(
  spl26_77 <=> ~para(sK22,sK21,sK23,sK22)),
  introduced(avatar_definition,[new_symbols(naming,[spl26_77])])).
fof(f6170,plain,(
  ~para(sK22,sK21,sK23,sK22) | para(sK22,sK20,sK23,sK22)),
  inference(resolution,[],[f1320,f431])).
fof(f1320,plain,(
  eqangle(sK22,sK20,sK23,sK22,sK22,sK21,sK23,sK22)),
  inference(resolution,[],[f672,f377])).
fof(f8562,plain,(
  spl26_76 | ~spl26_64),
  inference(avatar_split_clause,[],[f8543,f986,f1329])).
fof(f1329,plain,(
  spl26_76 <=> para(sK22,sK21,sK23,sK22)),
  introduced(avatar_definition,[new_symbols(naming,[spl26_76])])).
fof(f8543,plain,(
  para(sK22,sK21,sK23,sK22) | ~spl26_64),
  inference(resolution,[],[f987,f362])).
fof(f8511,plain,(
  spl26_66 | ~spl26_70 | ~spl26_148 | ~spl26_154),
  inference(avatar_split_clause,[],[f8499,f2044,f1865,f1009,f989])).
fof(f2044,plain,(
  spl26_154 <=> ! [X44,X46,X45,X47] : (~para(X44,X45,X46,X47) | perp(X44,X45,X46,X47))),
  introduced(avatar_definition,[new_symbols(naming,[spl26_154])])).
fof(f8499,plain,(
  para(sK23,sK22,sK22,sK20) | (~spl26_70 | ~spl26_148 | ~spl26_154)),
  inference(resolution,[],[f7761,f6951])).
fof(f6951,plain,(
  ( ! [X10,X9] : (~perp(X9,X10,sK23,sK22) | para(X9,X10,sK22,sK20)) ) | ~spl26_70),
  inference(resolution,[],[f1010,f366])).
fof(f366,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (~perp(X2,X3,X4,X5) | para(X0,X1,X4,X5) | ~perp(X0,X1,X2,X3)) )),
  inference(cnf_transformation,[],[f178])).
fof(f178,plain,(
  ! [X0,X1,X2,X3,X4,X5] : (para(X0,X1,X4,X5) | ~perp(X2,X3,X4,X5) | ~perp(X0,X1,X2,X3))),
  inference(flattening,[],[f177])).
fof(f177,plain,(
  ! [X0,X1,X2,X3,X4,X5] : (para(X0,X1,X4,X5) | (~perp(X2,X3,X4,X5) | ~perp(X0,X1,X2,X3)))),
  inference(ennf_transformation,[],[f9])).
fof(f9,axiom,(
  ! [X0,X1,X2,X3,X4,X5] : ((perp(X2,X3,X4,X5) & perp(X0,X1,X2,X3)) => para(X0,X1,X4,X5))),
  file('./7_Book/00EE/04/E051-8.p',ruleD9)).
fof(f7761,plain,(
  ( ! [X19,X18] : (perp(X18,X19,X18,X19)) ) | (~spl26_148 | ~spl26_154)),
  inference(resolution,[],[f2045,f1866])).
fof(f2045,plain,(
  ( ! [X47,X45,X46,X44] : (~para(X44,X45,X46,X47) | perp(X44,X45,X46,X47)) ) | ~spl26_154),
  inference(avatar_component_clause,[],[f2044])).
fof(f8507,plain,(
  spl26_64 | ~spl26_68 | ~spl26_148 | ~spl26_154),
  inference(avatar_split_clause,[],[f8500,f2044,f1865,f1006,f986])).
fof(f1006,plain,(
  spl26_68 <=> perp(sK23,sK22,sK22,sK21)),
  introduced(avatar_definition,[new_symbols(naming,[spl26_68])])).
fof(f8500,plain,(
  para(sK23,sK22,sK22,sK21) | (~spl26_68 | ~spl26_148 | ~spl26_154)),
  inference(resolution,[],[f7761,f6929])).
fof(f6929,plain,(
  ( ! [X10,X9] : (~perp(X9,X10,sK23,sK22) | para(X9,X10,sK22,sK21)) ) | ~spl26_68),
  inference(resolution,[],[f1007,f366])).
fof(f1007,plain,(
  perp(sK23,sK22,sK22,sK21) | ~spl26_68),
  inference(avatar_component_clause,[],[f1006])).
fof(f7716,plain,(
  ~spl26_148 | ~spl26_152 | ~spl26_194),
  inference(avatar_contradiction_clause,[],[f7715])).
fof(f7715,plain,(
  $false | (~spl26_148 | ~spl26_152 | ~spl26_194)),
  inference(subsumption_resolution,[],[f7704,f7230])).
fof(f7230,plain,(
  ( ! [X14] : (midp(X14,X14,X14)) ) | ~spl26_148),
  inference(subsumption_resolution,[],[f7209,f4382])).
fof(f7209,plain,(
  ( ! [X14] : (~coll(X14,X14,X14) | midp(X14,X14,X14)) ) | ~spl26_148),
  inference(resolution,[],[f6819,f425])).
fof(f6819,plain,(
  ( ! [X35] : (cong(X35,X35,X35,X35)) ) | ~spl26_148),
  inference(subsumption_resolution,[],[f6816,f6782])).
fof(f6816,plain,(
  ( ! [X35,X36] : (cong(X35,X35,X35,X35) | ~cyclic(X35,X35,X36,X35)) ) | ~spl26_148),
  inference(duplicate_literal_removal,[],[f6801])).
fof(f6801,plain,(
  ( ! [X35,X36] : (cong(X35,X35,X35,X35) | ~cyclic(X35,X35,X36,X35) | ~cyclic(X35,X35,X36,X35)) ) | ~spl26_148),
  inference(resolution,[],[f6782,f1265])).
fof(f7704,plain,(
  ~midp(sK21,sK21,sK21) | (~spl26_148 | ~spl26_152 | ~spl26_194)),
  inference(resolution,[],[f7408,f1496])).
fof(f7408,plain,(
  ( ! [X11] : (~para(sK21,X11,sK21,sK21)) ) | (~spl26_148 | ~spl26_152 | ~spl26_194)),
  inference(subsumption_resolution,[],[f7399,f6905])).
fof(f6905,plain,(
  ( ! [X24,X23,X22] : (~para(X22,X24,X22,X23) | ~perp(X22,X23,X22,X24)) ) | (~spl26_148 | ~spl26_152)),
  inference(subsumption_resolution,[],[f6887,f2042])).
fof(f2042,plain,(
  ( ! [X48,X49] : (~perp(X48,X49,X48,X49)) ) | ~spl26_152),
  inference(avatar_component_clause,[],[f2041])).
fof(f2041,plain,(
  spl26_152 <=> ! [X49,X48] : ~perp(X48,X49,X48,X49)),
  introduced(avatar_definition,[new_symbols(naming,[spl26_152])])).
fof(f6887,plain,(
  ( ! [X24,X23,X22] : (perp(X22,X23,X22,X23) | ~para(X22,X24,X22,X23) | ~perp(X22,X23,X22,X24)) ) | ~spl26_148),
  inference(resolution,[],[f6791,f997])).
fof(f997,plain,(
  ( ! [X6,X8,X7,X9] : (~cyclic(X9,X8,X6,X7) | perp(X9,X7,X6,X7) | ~para(X9,X8,X6,X7) | ~perp(X6,X7,X6,X8)) )),
  inference(resolution,[],[f432,f412])).
fof(f412,plain,(
  ( ! [X2,X0,X3,X1] : (eqangle(X0,X3,X2,X3,X2,X3,X2,X1) | ~para(X0,X1,X2,X3) | ~cyclic(X0,X1,X2,X3)) )),
  inference(cnf_transformation,[],[f246])).
fof(f246,plain,(
  ! [X0,X1,X2,X3] : (eqangle(X0,X3,X2,X3,X2,X3,X2,X1) | ~para(X0,X1,X2,X3) | ~cyclic(X0,X1,X2,X3))),
  inference(flattening,[],[f245])).
fof(f245,plain,(
  ! [X0,X1,X2,X3] : (eqangle(X0,X3,X2,X3,X2,X3,X2,X1) | (~para(X0,X1,X2,X3) | ~cyclic(X0,X1,X2,X3)))),
  inference(ennf_transformation,[],[f55])).
fof(f55,axiom,(
  ! [X0,X1,X2,X3] : ((para(X0,X1,X2,X3) & cyclic(X0,X1,X2,X3)) => eqangle(X0,X3,X2,X3,X2,X3,X2,X1))),
  file('./7_Book/00EE/04/E051-8.p',ruleD54)).
fof(f432,plain,(
  ( ! [X6,X4,X2,X0,X7,X5,X3,X1] : (~eqangle(X0,X1,X2,X3,X4,X5,X6,X7) | ~perp(X4,X5,X6,X7) | perp(X0,X1,X2,X3)) )),
  inference(cnf_transformation,[],[f280])).
fof(f280,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6,X7] : (perp(X0,X1,X2,X3) | ~perp(X4,X5,X6,X7) | ~eqangle(X0,X1,X2,X3,X4,X5,X6,X7))),
  inference(flattening,[],[f279])).
fof(f279,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6,X7] : (perp(X0,X1,X2,X3) | (~perp(X4,X5,X6,X7) | ~eqangle(X0,X1,X2,X3,X4,X5,X6,X7)))),
  inference(ennf_transformation,[],[f147])).
fof(f147,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6,X7] : ((perp(X4,X5,X6,X7) & eqangle(X0,X1,X2,X3,X4,X5,X6,X7)) => perp(X0,X1,X2,X3))),
  inference(rectify,[],[f75])).
fof(f75,axiom,(
  ! [X0,X1,X2,X3,X8,X9,X10,X11] : ((perp(X8,X9,X10,X11) & eqangle(X0,X1,X2,X3,X8,X9,X10,X11)) => perp(X0,X1,X2,X3))),
  file('./7_Book/00EE/04/E051-8.p',ruleD74)).
fof(f7399,plain,(
  ( ! [X11] : (perp(sK21,sK21,sK21,X11) | ~para(sK21,X11,sK21,sK21)) ) | ~spl26_194),
  inference(resolution,[],[f3549,f955])).
fof(f955,plain,(
  ( ! [X2,X0,X3,X1] : (~circle(X0,X1,X3,X1) | perp(X0,X1,X1,X2) | ~para(X1,X2,X1,X1)) )),
  inference(resolution,[],[f407,f397])).
fof(f407,plain,(
  ( ! [X4,X2,X0,X3,X1] : (~eqangle(X0,X4,X0,X1,X2,X0,X2,X1) | perp(X3,X0,X0,X4) | ~circle(X3,X0,X1,X2)) )),
  inference(cnf_transformation,[],[f236])).
fof(f236,plain,(
  ! [X0,X1,X2,X3,X4] : (perp(X3,X0,X0,X4) | ~eqangle(X0,X4,X0,X1,X2,X0,X2,X1) | ~circle(X3,X0,X1,X2))),
  inference(flattening,[],[f235])).
fof(f235,plain,(
  ! [X0,X1,X2,X3,X4] : (perp(X3,X0,X0,X4) | (~eqangle(X0,X4,X0,X1,X2,X0,X2,X1) | ~circle(X3,X0,X1,X2)))),
  inference(ennf_transformation,[],[f129])).
fof(f129,plain,(
  ! [X0,X1,X2,X3,X4] : ((eqangle(X0,X4,X0,X1,X2,X0,X2,X1) & circle(X3,X0,X1,X2)) => perp(X3,X0,X0,X4))),
  inference(rectify,[],[f50])).
fof(f50,axiom,(
  ! [X0,X1,X2,X7,X15] : ((eqangle(X0,X15,X0,X1,X2,X0,X2,X1) & circle(X7,X0,X1,X2)) => perp(X7,X0,X0,X15))),
  file('./7_Book/00EE/04/E051-8.p',ruleD49)).
fof(f3549,plain,(
  circle(sK21,sK21,sK21,sK21) | ~spl26_194),
  inference(avatar_component_clause,[],[f3548])).
fof(f3548,plain,(
  spl26_194 <=> circle(sK21,sK21,sK21,sK21)),
  introduced(avatar_definition,[new_symbols(naming,[spl26_194])])).
fof(f7608,plain,(
  ~spl26_551 | ~spl26_553),
  inference(avatar_split_clause,[],[f7577,f7606,f7600])).
fof(f7577,plain,(
  ~midp(sK25,sK22,sK21) | ~midp(sK22,sK24,sK25)),
  inference(resolution,[],[f1745,f477])).
fof(f477,plain,(
  ~cong(sK22,sK24,sK25,sK21)),
  inference(cnf_transformation,[],[f357])).
fof(f1745,plain,(
  ( ! [X14,X12,X15,X13] : (cong(X12,X13,X14,X15) | ~midp(X14,X12,X15) | ~midp(X12,X13,X14)) )),
  inference(resolution,[],[f617,f514])).
fof(f514,plain,(
  ( ! [X2,X0,X1] : (cong(X0,X1,X2,X0) | ~midp(X0,X1,X2)) )),
  inference(resolution,[],[f380,f426])).
fof(f380,plain,(
  ( ! [X2,X0,X3,X1] : (~cong(X0,X1,X2,X3) | cong(X0,X1,X3,X2)) )),
  inference(cnf_transformation,[],[f197])).
fof(f197,plain,(
  ! [X0,X1,X2,X3] : (cong(X0,X1,X3,X2) | ~cong(X0,X1,X2,X3))),
  inference(ennf_transformation,[],[f23])).
fof(f23,axiom,(
  ! [X0,X1,X2,X3] : (cong(X0,X1,X2,X3) => cong(X0,X1,X3,X2))),
  file('./7_Book/00EE/04/E051-8.p',ruleD23)).
fof(f617,plain,(
  ( ! [X4,X2,X0,X3,X1] : (~cong(X0,X1,X2,X4) | cong(X0,X1,X2,X3) | ~midp(X2,X4,X3)) )),
  inference(resolution,[],[f382,f426])).
fof(f382,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (~cong(X2,X3,X4,X5) | cong(X0,X1,X4,X5) | ~cong(X0,X1,X2,X3)) )),
  inference(cnf_transformation,[],[f200])).
fof(f200,plain,(
  ! [X0,X1,X2,X3,X4,X5] : (cong(X0,X1,X4,X5) | ~cong(X2,X3,X4,X5) | ~cong(X0,X1,X2,X3))),
  inference(flattening,[],[f199])).
fof(f199,plain,(
  ! [X0,X1,X2,X3,X4,X5] : (cong(X0,X1,X4,X5) | (~cong(X2,X3,X4,X5) | ~cong(X0,X1,X2,X3)))),
  inference(ennf_transformation,[],[f25])).
fof(f25,axiom,(
  ! [X0,X1,X2,X3,X4,X5] : ((cong(X2,X3,X4,X5) & cong(X0,X1,X2,X3)) => cong(X0,X1,X4,X5))),
  file('./7_Book/00EE/04/E051-8.p',ruleD25)).
fof(f7379,plain,(
  ~spl26_148 | spl26_195),
  inference(avatar_contradiction_clause,[],[f7378])).
fof(f7378,plain,(
  $false | (~spl26_148 | ~spl26_195)),
  inference(subsumption_resolution,[],[f7371,f7230])).
fof(f7371,plain,(
  ~midp(sK21,sK21,sK21) | ~spl26_195),
  inference(duplicate_literal_removal,[],[f7354])).
fof(f7354,plain,(
  ~midp(sK21,sK21,sK21) | ~midp(sK21,sK21,sK21) | ~spl26_195),
  inference(resolution,[],[f1734,f3552])).
fof(f3552,plain,(
  ~circle(sK21,sK21,sK21,sK21) | ~spl26_195),
  inference(avatar_component_clause,[],[f3551])).
fof(f3551,plain,(
  spl26_195 <=> ~circle(sK21,sK21,sK21,sK21)),
  introduced(avatar_definition,[new_symbols(naming,[spl26_195])])).
fof(f6919,plain,(
  spl26_70 | ~spl26_69),
  inference(avatar_split_clause,[],[f1374,f1003,f1009])).
fof(f1003,plain,(
  spl26_69 <=> ~perp(sK23,sK22,sK22,sK21)),
  introduced(avatar_definition,[new_symbols(naming,[spl26_69])])).
fof(f1374,plain,(
  ~perp(sK23,sK22,sK22,sK21) | perp(sK23,sK22,sK22,sK20)),
  inference(resolution,[],[f680,f432])).
fof(f680,plain,(
  eqangle(sK23,sK22,sK22,sK20,sK23,sK22,sK22,sK21)),
  inference(resolution,[],[f377,f471])).
fof(f6917,plain,(
  ~spl26_148 | spl26_271),
  inference(avatar_contradiction_clause,[],[f6916])).
fof(f6916,plain,(
  $false | (~spl26_148 | ~spl26_271)),
  inference(subsumption_resolution,[],[f3824,f6908])).
fof(f3824,plain,(
  ~cong(sK22,sK21,sK22,sK21) | ~spl26_271),
  inference(avatar_component_clause,[],[f3823])).
fof(f3823,plain,(
  spl26_271 <=> ~cong(sK22,sK21,sK22,sK21)),
  introduced(avatar_definition,[new_symbols(naming,[spl26_271])])).
fof(f6904,plain,(
  ~spl26_92 | ~spl26_96 | ~spl26_148 | spl26_269),
  inference(avatar_contradiction_clause,[],[f6903])).
fof(f6903,plain,(
  $false | (~spl26_92 | ~spl26_96 | ~spl26_148 | ~spl26_269)),
  inference(resolution,[],[f6791,f4737])).
fof(f4737,plain,(
  ~cyclic(sK22,sK23,sK22,sK23) | (~spl26_92 | ~spl26_96 | ~spl26_269)),
  inference(subsumption_resolution,[],[f4736,f3918])).
fof(f3918,plain,(
  cyclic(sK22,sK23,sK22,sK22) | ~spl26_96),
  inference(resolution,[],[f1420,f372])).
fof(f1420,plain,(
  cyclic(sK22,sK22,sK23,sK22) | ~spl26_96),
  inference(avatar_component_clause,[],[f1419])).
fof(f1419,plain,(
  spl26_96 <=> cyclic(sK22,sK22,sK23,sK22)),
  introduced(avatar_definition,[new_symbols(naming,[spl26_96])])).
fof(f4736,plain,(
  ~cyclic(sK22,sK23,sK22,sK23) | ~cyclic(sK22,sK23,sK22,sK22) | (~spl26_92 | ~spl26_269)),
  inference(subsumption_resolution,[],[f4666,f3818])).
fof(f3818,plain,(
  ~cong(sK22,sK23,sK22,sK23) | ~spl26_269),
  inference(avatar_component_clause,[],[f3817])).
fof(f3817,plain,(
  spl26_269 <=> ~cong(sK22,sK23,sK22,sK23)),
  introduced(avatar_definition,[new_symbols(naming,[spl26_269])])).
fof(f4666,plain,(
  cong(sK22,sK23,sK22,sK23) | ~cyclic(sK22,sK23,sK22,sK23) | ~cyclic(sK22,sK23,sK22,sK22) | ~spl26_92),
  inference(resolution,[],[f3801,f1265])).
fof(f3801,plain,(
  cyclic(sK22,sK23,sK22,sK21) | ~spl26_92),
  inference(resolution,[],[f1408,f372])).
fof(f1408,plain,(
  cyclic(sK22,sK22,sK23,sK21) | ~spl26_92),
  inference(avatar_component_clause,[],[f1407])).
fof(f1407,plain,(
  spl26_92 <=> cyclic(sK22,sK22,sK23,sK21)),
  introduced(avatar_definition,[new_symbols(naming,[spl26_92])])).
fof(f6123,plain,(
  ~spl26_347 | ~spl26_94 | spl26_373),
  inference(avatar_split_clause,[],[f6122,f5046,f1413,f4857])).
fof(f4857,plain,(
  spl26_347 <=> ~para(sK20,sK23,sK20,sK22)),
  introduced(avatar_definition,[new_symbols(naming,[spl26_347])])).
fof(f1413,plain,(
  spl26_94 <=> cyclic(sK22,sK22,sK23,sK20)),
  introduced(avatar_definition,[new_symbols(naming,[spl26_94])])).
fof(f5046,plain,(
  spl26_373 <=> ~para(sK22,sK23,sK22,sK22)),
  introduced(avatar_definition,[new_symbols(naming,[spl26_373])])).
fof(f6122,plain,(
  ~para(sK20,sK23,sK20,sK22) | (~spl26_94 | ~spl26_373)),
  inference(subsumption_resolution,[],[f6116,f5047])).
fof(f5047,plain,(
  ~para(sK22,sK23,sK22,sK22) | ~spl26_373),
  inference(avatar_component_clause,[],[f5046])).
fof(f6116,plain,(
  para(sK22,sK23,sK22,sK22) | ~para(sK20,sK23,sK20,sK22) | ~spl26_94),
  inference(resolution,[],[f4897,f978])).
fof(f978,plain,(
  ( ! [X12,X10,X13,X11] : (~cyclic(X11,X12,X13,X10) | para(X13,X11,X13,X12) | ~para(X10,X11,X10,X12)) )),
  inference(resolution,[],[f431,f398])).
fof(f4897,plain,(
  cyclic(sK23,sK22,sK22,sK20) | ~spl26_94),
  inference(resolution,[],[f3874,f373])).
fof(f373,plain,(
  ( ! [X2,X0,X3,X1] : (~cyclic(X0,X1,X2,X3) | cyclic(X1,X0,X2,X3)) )),
  inference(cnf_transformation,[],[f188])).
fof(f188,plain,(
  ! [X0,X1,X2,X3] : (cyclic(X1,X0,X2,X3) | ~cyclic(X0,X1,X2,X3))),
  inference(ennf_transformation,[],[f16])).
fof(f16,axiom,(
  ! [X0,X1,X2,X3] : (cyclic(X0,X1,X2,X3) => cyclic(X1,X0,X2,X3))),
  file('./7_Book/00EE/04/E051-8.p',ruleD16)).
fof(f3874,plain,(
  cyclic(sK22,sK23,sK22,sK20) | ~spl26_94),
  inference(resolution,[],[f1414,f372])).
fof(f1414,plain,(
  cyclic(sK22,sK22,sK23,sK20) | ~spl26_94),
  inference(avatar_component_clause,[],[f1413])).
fof(f3957,plain,(
  ~spl26_150),
  inference(avatar_contradiction_clause,[],[f3952])).
fof(f3952,plain,(
  $false | ~spl26_150),
  inference(resolution,[],[f1869,f1392])).
fof(f1869,plain,(
  ( ! [X0,X1] : (~para(X0,X1,X0,X1)) ) | ~spl26_150),
  inference(avatar_component_clause,[],[f1868])).
fof(f1868,plain,(
  spl26_150 <=> ! [X1,X0] : ~para(X0,X1,X0,X1)),
  introduced(avatar_definition,[new_symbols(naming,[spl26_150])])).
fof(f3825,plain,(
  ~spl26_269 | ~spl26_271 | spl26_69 | ~spl26_92),
  inference(avatar_split_clause,[],[f3812,f1407,f1003,f3823,f3817])).
fof(f3812,plain,(
  ~cong(sK22,sK21,sK22,sK21) | ~cong(sK22,sK23,sK22,sK23) | (~spl26_69 | ~spl26_92)),
  inference(subsumption_resolution,[],[f3804,f1004])).
fof(f1004,plain,(
  ~perp(sK23,sK22,sK22,sK21) | ~spl26_69),
  inference(avatar_component_clause,[],[f1003])).
fof(f3804,plain,(
  perp(sK23,sK22,sK22,sK21) | ~cong(sK22,sK21,sK22,sK21) | ~cong(sK22,sK23,sK22,sK23) | ~spl26_92),
  inference(resolution,[],[f1408,f415])).
fof(f415,plain,(
  ( ! [X2,X0,X3,X1] : (~cyclic(X0,X1,X2,X3) | perp(X2,X0,X0,X3) | ~cong(X0,X3,X1,X3) | ~cong(X0,X2,X1,X2)) )),
  inference(cnf_transformation,[],[f252])).
fof(f252,plain,(
  ! [X0,X1,X2,X3] : (perp(X2,X0,X0,X3) | ~cyclic(X0,X1,X2,X3) | ~cong(X0,X3,X1,X3) | ~cong(X0,X2,X1,X2))),
  inference(flattening,[],[f251])).
fof(f251,plain,(
  ! [X0,X1,X2,X3] : (perp(X2,X0,X0,X3) | (~cyclic(X0,X1,X2,X3) | ~cong(X0,X3,X1,X3) | ~cong(X0,X2,X1,X2)))),
  inference(ennf_transformation,[],[f136])).
fof(f136,plain,(
  ! [X0,X1,X2,X3] : ((cyclic(X0,X1,X2,X3) & cong(X0,X3,X1,X3) & cong(X0,X2,X1,X2)) => perp(X2,X0,X0,X3))),
  inference(rectify,[],[f58])).
fof(f58,axiom,(
  ! [X0,X1,X8,X9] : ((cyclic(X0,X1,X8,X9) & cong(X0,X9,X1,X9) & cong(X0,X8,X1,X8)) => perp(X8,X0,X0,X9))),
  file('./7_Book/00EE/04/E051-8.p',ruleD57)).
fof(f3512,plain,(
  ~spl26_74 | ~spl26_148),
  inference(avatar_contradiction_clause,[],[f3511])).
fof(f3511,plain,(
  $false | (~spl26_74 | ~spl26_148)),
  inference(subsumption_resolution,[],[f3510,f2357])).
fof(f2357,plain,(
  ( ! [X0,X1] : (~midp(X0,X1,X0)) ) | ~spl26_74),
  inference(subsumption_resolution,[],[f2352,f368])).
fof(f2352,plain,(
  ( ! [X0,X1] : (~midp(X0,X0,X1) | ~midp(X0,X1,X0)) ) | ~spl26_74),
  inference(resolution,[],[f1022,f517])).
fof(f1022,plain,(
  ( ! [X6,X8,X7] : (~cong(X6,X7,X7,X8) | ~midp(X6,X7,X8)) ) | ~spl26_74),
  inference(avatar_component_clause,[],[f1021])).
fof(f1021,plain,(
  spl26_74 <=> ! [X7,X8,X6] : (~cong(X6,X7,X7,X8) | ~midp(X6,X7,X8))),
  introduced(avatar_definition,[new_symbols(naming,[spl26_74])])).
fof(f3510,plain,(
  ( ! [X14] : (midp(X14,X14,X14)) ) | ~spl26_148),
  inference(subsumption_resolution,[],[f3491,f3093])).
fof(f3093,plain,(
  ( ! [X4,X3] : (coll(X3,X4,X3)) ) | ~spl26_148),
  inference(resolution,[],[f3077,f359])).
fof(f3077,plain,(
  ( ! [X4,X3] : (coll(X3,X4,X4)) ) | ~spl26_148),
  inference(resolution,[],[f1866,f424])).
fof(f3491,plain,(
  ( ! [X14] : (~coll(X14,X14,X14) | midp(X14,X14,X14)) ) | ~spl26_148),
  inference(resolution,[],[f3180,f425])).
fof(f3180,plain,(
  ( ! [X3] : (cong(X3,X3,X3,X3)) ) | ~spl26_148),
  inference(subsumption_resolution,[],[f3178,f2050])).
fof(f2050,plain,(
  ( ! [X70,X71,X69] : (cyclic(X70,X70,X69,X71)) ) | ~spl26_148),
  inference(subsumption_resolution,[],[f2049,f2048])).
fof(f2048,plain,(
  ( ! [X68,X66,X67] : (~coll(X66,X68,X67) | cyclic(X67,X67,X66,X68)) ) | ~spl26_148),
  inference(subsumption_resolution,[],[f2029,f1866])).
fof(f2049,plain,(
  ( ! [X70,X71,X69] : (coll(X69,X71,X70) | cyclic(X70,X70,X69,X71)) ) | ~spl26_148),
  inference(subsumption_resolution,[],[f2030,f1866])).
fof(f3178,plain,(
  ( ! [X4,X3] : (cong(X3,X3,X3,X3) | ~cyclic(X3,X3,X4,X3)) ) | ~spl26_148),
  inference(duplicate_literal_removal,[],[f3177])).
fof(f3177,plain,(
  ( ! [X4,X3] : (cong(X3,X3,X3,X3) | ~cyclic(X3,X3,X4,X3) | ~cyclic(X3,X3,X4,X3)) ) | ~spl26_148),
  inference(resolution,[],[f1265,f2050])).
fof(f3169,plain,(
  spl26_95 | ~spl26_148),
  inference(avatar_contradiction_clause,[],[f3168])).
fof(f3168,plain,(
  $false | (~spl26_95 | ~spl26_148)),
  inference(subsumption_resolution,[],[f1417,f2050])).
fof(f1417,plain,(
  ~cyclic(sK22,sK22,sK23,sK20) | ~spl26_95),
  inference(avatar_component_clause,[],[f1416])).
fof(f1416,plain,(
  spl26_95 <=> ~cyclic(sK22,sK22,sK23,sK20)),
  introduced(avatar_definition,[new_symbols(naming,[spl26_95])])).
fof(f3167,plain,(
  spl26_93 | ~spl26_148),
  inference(avatar_contradiction_clause,[],[f3166])).
fof(f3166,plain,(
  $false | (~spl26_93 | ~spl26_148)),
  inference(subsumption_resolution,[],[f1411,f2050])).
fof(f1411,plain,(
  ~cyclic(sK22,sK22,sK23,sK21) | ~spl26_93),
  inference(avatar_component_clause,[],[f1410])).
fof(f1410,plain,(
  spl26_93 <=> ~cyclic(sK22,sK22,sK23,sK21)),
  introduced(avatar_definition,[new_symbols(naming,[spl26_93])])).
fof(f3162,plain,(
  spl26_97 | ~spl26_148),
  inference(avatar_contradiction_clause,[],[f3161])).
fof(f3161,plain,(
  $false | (~spl26_97 | ~spl26_148)),
  inference(resolution,[],[f2050,f1423])).
fof(f1423,plain,(
  ~cyclic(sK22,sK22,sK23,sK22) | ~spl26_97),
  inference(avatar_component_clause,[],[f1422])).
fof(f1422,plain,(
  spl26_97 <=> ~cyclic(sK22,sK22,sK23,sK22)),
  introduced(avatar_definition,[new_symbols(naming,[spl26_97])])).
fof(f2348,plain,(
  spl26_156 | spl26_74),
  inference(avatar_split_clause,[],[f2335,f1021,f2342])).
fof(f2335,plain,(
  ( ! [X14,X12,X10,X15,X13,X11] : (~midp(X10,X11,X12) | ~midp(X13,X14,X15) | ~cong(X10,X11,X11,X12) | cong(X14,X13,X14,X15)) )),
  inference(resolution,[],[f698,f433])).
fof(f698,plain,(
  ( ! [X6,X10,X8,X7,X11,X9] : (eqratio(X6,X7,X6,X8,X9,X10,X10,X11) | ~midp(X9,X10,X11) | ~midp(X7,X6,X8)) )),
  inference(resolution,[],[f383,f428])).
fof(f383,plain,(
  ( ! [X6,X4,X2,X0,X7,X5,X3,X1] : (~eqratio(X0,X1,X2,X3,X4,X5,X6,X7) | eqratio(X1,X0,X2,X3,X4,X5,X6,X7)) )),
  inference(cnf_transformation,[],[f201])).
fof(f201,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6,X7] : (eqratio(X1,X0,X2,X3,X4,X5,X6,X7) | ~eqratio(X0,X1,X2,X3,X4,X5,X6,X7))),
  inference(ennf_transformation,[],[f105])).
fof(f105,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6,X7] : (eqratio(X0,X1,X2,X3,X4,X5,X6,X7) => eqratio(X1,X0,X2,X3,X4,X5,X6,X7))),
  inference(rectify,[],[f26])).
fof(f26,axiom,(
  ! [X0,X1,X2,X3,X8,X9,X10,X11] : (eqratio(X0,X1,X2,X3,X8,X9,X10,X11) => eqratio(X1,X0,X2,X3,X8,X9,X10,X11))),
  file('./7_Book/00EE/04/E051-8.p',ruleD26)).
fof(f2046,plain,(
  spl26_152 | spl26_154),
  inference(avatar_split_clause,[],[f2023,f2044,f2041])).
fof(f2023,plain,(
  ( ! [X47,X45,X48,X46,X44,X49] : (~para(X44,X45,X46,X47) | ~perp(X48,X49,X48,X49) | perp(X44,X45,X46,X47)) )),
  inference(resolution,[],[f690,f432])).
fof(f1870,plain,(
  spl26_148 | spl26_150),
  inference(avatar_split_clause,[],[f1841,f1868,f1865])).
fof(f1841,plain,(
  ( ! [X2,X0,X3,X1] : (~para(X0,X1,X0,X1) | para(X2,X3,X2,X3)) )),
  inference(resolution,[],[f673,f396])).
fof(f673,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (eqangle(X0,X1,X2,X3,X0,X1,X4,X5) | ~para(X2,X3,X4,X5)) )),
  inference(resolution,[],[f376,f397])).
% SZS output end Proof for E051-8
% ------------------------------
% Version: Vampire 4.4.0 (commit 7916d27 on 2019-08-23 08:50:16 +0100)
% Termination reason: Refutation

% Memory used [KB]: 17014
% Time elapsed: 0.521 s
% ------------------------------
% ------------------------------
