% Refutation found. Thanks to Tanya!
% SZS status Theorem for 26
% SZS output start Proof for 26
fof(f60360,plain,(
  $false),
  inference(avatar_sat_refutation,[],[f586,f3325,f5736,f60349,f60351,f60353,f60357])).
fof(f60357,plain,(
  spl42_13 | ~spl42_314),
  inference(avatar_contradiction_clause,[],[f60356])).
fof(f60356,plain,(
  $false | (~spl42_13 | ~spl42_314)),
  inference(subsumption_resolution,[],[f571,f57105])).
fof(f57105,plain,(
  ( ! [X251,X249,X246,X244,X250,X248,X247,X245] : (eqangle(X244,X245,X246,X247,X248,X249,X250,X251)) ) | ~spl42_314),
  inference(subsumption_resolution,[],[f57104,f45683])).
fof(f45683,plain,(
  ( ! [X6,X8,X7,X9] : (para(X6,X7,X8,X9)) ) | ~spl42_314),
  inference(subsumption_resolution,[],[f44905,f44851])).
fof(f44851,plain,(
  ( ! [X61,X62,X63] : (midp(X61,X62,X63)) ) | ~spl42_314),
  inference(subsumption_resolution,[],[f44850,f20071])).
fof(f20071,plain,(
  ( ! [X30,X29] : (midp(X29,X30,X30)) ) | ~spl42_314),
  inference(subsumption_resolution,[],[f20041,f5898])).
fof(f5898,plain,(
  ( ! [X4,X3] : (coll(X3,X4,X4)) ) | ~spl42_314),
  inference(resolution,[],[f3321,f433])).
fof(f433,plain,(
  ( ! [X2,X0,X1] : (~para(X0,X1,X0,X2) | coll(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f268])).
fof(f268,plain,(
  ! [X0,X1,X2] : (coll(X0,X1,X2) | ~para(X0,X1,X0,X2))),
  inference(ennf_transformation,[],[f67])).
fof(f67,axiom,(
  ! [X0,X1,X2] : (para(X0,X1,X0,X2) => coll(X0,X1,X2))),
  file('./6_GDD_FULL/21-40/26.p',ruleD66)).
fof(f3321,plain,(
  ( ! [X2,X3] : (para(X2,X3,X2,X3)) ) | ~spl42_314),
  inference(avatar_component_clause,[],[f3320])).
fof(f3320,plain,(
  spl42_314 <=> ! [X3,X2] : para(X2,X3,X2,X3)),
  introduced(avatar_definition,[new_symbols(naming,[spl42_314])])).
fof(f20041,plain,(
  ( ! [X30,X29] : (~coll(X29,X30,X30) | midp(X29,X30,X30)) ) | ~spl42_314),
  inference(resolution,[],[f20030,f434])).
fof(f434,plain,(
  ( ! [X2,X0,X1] : (~cong(X0,X1,X0,X2) | ~coll(X0,X1,X2) | midp(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f270])).
fof(f270,plain,(
  ! [X0,X1,X2] : (midp(X0,X1,X2) | ~coll(X0,X1,X2) | ~cong(X0,X1,X0,X2))),
  inference(flattening,[],[f269])).
fof(f269,plain,(
  ! [X0,X1,X2] : (midp(X0,X1,X2) | (~coll(X0,X1,X2) | ~cong(X0,X1,X0,X2)))),
  inference(ennf_transformation,[],[f68])).
fof(f68,axiom,(
  ! [X0,X1,X2] : ((coll(X0,X1,X2) & cong(X0,X1,X0,X2)) => midp(X0,X1,X2))),
  file('./6_GDD_FULL/21-40/26.p',ruleD67)).
fof(f20030,plain,(
  ( ! [X103,X104] : (cong(X103,X104,X103,X104)) ) | ~spl42_314),
  inference(subsumption_resolution,[],[f20029,f18577])).
fof(f18577,plain,(
  ( ! [X12,X10,X11,X9] : (cyclic(X9,X10,X11,X12)) ) | ~spl42_314),
  inference(subsumption_resolution,[],[f18555,f16194])).
fof(f16194,plain,(
  ( ! [X2,X0,X1] : (cyclic(X0,X0,X1,X2)) ) | ~spl42_314),
  inference(resolution,[],[f5474,f3321])).
fof(f5474,plain,(
  ( ! [X97,X98,X96] : (~para(X96,X97,X96,X97) | cyclic(X97,X97,X96,X98)) )),
  inference(subsumption_resolution,[],[f3537,f3538])).
fof(f3538,plain,(
  ( ! [X101,X99,X100] : (~para(X99,X100,X99,X100) | coll(X99,X101,X100) | cyclic(X100,X100,X99,X101)) )),
  inference(resolution,[],[f1076,f408])).
fof(f408,plain,(
  ( ! [X2,X0,X3,X1] : (~eqangle(X2,X0,X2,X1,X3,X0,X3,X1) | coll(X2,X3,X0) | cyclic(X0,X1,X2,X3)) )),
  inference(cnf_transformation,[],[f223])).
fof(f223,plain,(
  ! [X0,X1,X2,X3] : (cyclic(X0,X1,X2,X3) | coll(X2,X3,X0) | ~eqangle(X2,X0,X2,X1,X3,X0,X3,X1))),
  inference(flattening,[],[f222])).
fof(f222,plain,(
  ! [X0,X1,X2,X3] : (cyclic(X0,X1,X2,X3) | (coll(X2,X3,X0) | ~eqangle(X2,X0,X2,X1,X3,X0,X3,X1)))),
  inference(ennf_transformation,[],[f121])).
fof(f121,plain,(
  ! [X0,X1,X2,X3] : ((~coll(X2,X3,X0) & eqangle(X2,X0,X2,X1,X3,X0,X3,X1)) => cyclic(X0,X1,X2,X3))),
  inference(rectify,[],[f42])).
fof(f42,axiom,(
  ! [X0,X1,X8,X9] : ((~coll(X8,X9,X0) & eqangle(X8,X0,X8,X1,X9,X0,X9,X1)) => cyclic(X0,X1,X8,X9))),
  file('./6_GDD_FULL/21-40/26.p',ruleD42a)).
fof(f1076,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (eqangle(X0,X1,X2,X3,X4,X5,X4,X5) | ~para(X0,X1,X2,X3)) )),
  inference(resolution,[],[f387,f406])).
fof(f406,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (eqangle(X0,X1,X4,X5,X2,X3,X4,X5) | ~para(X0,X1,X2,X3)) )),
  inference(cnf_transformation,[],[f220])).
fof(f220,plain,(
  ! [X0,X1,X2,X3,X4,X5] : (eqangle(X0,X1,X4,X5,X2,X3,X4,X5) | ~para(X0,X1,X2,X3))),
  inference(ennf_transformation,[],[f119])).
fof(f119,plain,(
  ! [X0,X1,X2,X3,X4,X5] : (para(X0,X1,X2,X3) => eqangle(X0,X1,X4,X5,X2,X3,X4,X5))),
  inference(rectify,[],[f40])).
fof(f40,axiom,(
  ! [X0,X1,X2,X3,X8,X9] : (para(X0,X1,X2,X3) => eqangle(X0,X1,X8,X9,X2,X3,X8,X9))),
  file('./6_GDD_FULL/21-40/26.p',ruleD40)).
fof(f387,plain,(
  ( ! [X6,X4,X2,X0,X7,X5,X3,X1] : (~eqangle(X0,X1,X2,X3,X4,X5,X6,X7) | eqangle(X0,X1,X4,X5,X2,X3,X6,X7)) )),
  inference(cnf_transformation,[],[f196])).
fof(f196,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6,X7] : (eqangle(X0,X1,X4,X5,X2,X3,X6,X7) | ~eqangle(X0,X1,X2,X3,X4,X5,X6,X7))),
  inference(ennf_transformation,[],[f103])).
fof(f103,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6,X7] : (eqangle(X0,X1,X2,X3,X4,X5,X6,X7) => eqangle(X0,X1,X4,X5,X2,X3,X6,X7))),
  inference(rectify,[],[f21])).
fof(f21,axiom,(
  ! [X0,X1,X2,X3,X8,X9,X10,X11] : (eqangle(X0,X1,X2,X3,X8,X9,X10,X11) => eqangle(X0,X1,X8,X9,X2,X3,X10,X11))),
  file('./6_GDD_FULL/21-40/26.p',ruleD21)).
fof(f3537,plain,(
  ( ! [X97,X98,X96] : (~para(X96,X97,X96,X97) | ~coll(X96,X98,X97) | cyclic(X97,X97,X96,X98)) )),
  inference(resolution,[],[f1076,f409])).
fof(f409,plain,(
  ( ! [X2,X0,X3,X1] : (~eqangle(X2,X0,X2,X1,X3,X0,X3,X1) | ~coll(X2,X3,X1) | cyclic(X0,X1,X2,X3)) )),
  inference(cnf_transformation,[],[f225])).
fof(f225,plain,(
  ! [X0,X1,X2,X3] : (cyclic(X0,X1,X2,X3) | ~coll(X2,X3,X1) | ~eqangle(X2,X0,X2,X1,X3,X0,X3,X1))),
  inference(flattening,[],[f224])).
fof(f224,plain,(
  ! [X0,X1,X2,X3] : (cyclic(X0,X1,X2,X3) | (~coll(X2,X3,X1) | ~eqangle(X2,X0,X2,X1,X3,X0,X3,X1)))),
  inference(ennf_transformation,[],[f122])).
fof(f122,plain,(
  ! [X0,X1,X2,X3] : ((coll(X2,X3,X1) & eqangle(X2,X0,X2,X1,X3,X0,X3,X1)) => cyclic(X0,X1,X2,X3))),
  inference(rectify,[],[f43])).
fof(f43,axiom,(
  ! [X0,X1,X8,X9] : ((coll(X8,X9,X1) & eqangle(X8,X0,X8,X1,X9,X0,X9,X1)) => cyclic(X0,X1,X8,X9))),
  file('./6_GDD_FULL/21-40/26.p',ruleD42b)).
fof(f18555,plain,(
  ( ! [X12,X10,X11,X9] : (cyclic(X9,X10,X11,X12) | ~cyclic(X9,X9,X10,X11)) ) | ~spl42_314),
  inference(resolution,[],[f16194,f383])).
fof(f383,plain,(
  ( ! [X4,X2,X0,X3,X1] : (~cyclic(X0,X1,X2,X4) | cyclic(X1,X2,X3,X4) | ~cyclic(X0,X1,X2,X3)) )),
  inference(cnf_transformation,[],[f192])).
fof(f192,plain,(
  ! [X0,X1,X2,X3,X4] : (cyclic(X1,X2,X3,X4) | ~cyclic(X0,X1,X2,X4) | ~cyclic(X0,X1,X2,X3))),
  inference(flattening,[],[f191])).
fof(f191,plain,(
  ! [X0,X1,X2,X3,X4] : (cyclic(X1,X2,X3,X4) | (~cyclic(X0,X1,X2,X4) | ~cyclic(X0,X1,X2,X3)))),
  inference(ennf_transformation,[],[f17])).
fof(f17,axiom,(
  ! [X0,X1,X2,X3,X4] : ((cyclic(X0,X1,X2,X4) & cyclic(X0,X1,X2,X3)) => cyclic(X1,X2,X3,X4))),
  file('./6_GDD_FULL/21-40/26.p',ruleD17)).
fof(f20029,plain,(
  ( ! [X103,X105,X104] : (cong(X103,X104,X103,X104) | ~cyclic(X103,X104,X105,X104)) ) | ~spl42_314),
  inference(subsumption_resolution,[],[f20028,f18577])).
fof(f20028,plain,(
  ( ! [X103,X105,X104] : (cong(X103,X104,X103,X104) | ~cyclic(X103,X104,X105,X103) | ~cyclic(X103,X104,X105,X104)) ) | ~spl42_314),
  inference(subsumption_resolution,[],[f20013,f18553])).
fof(f18553,plain,(
  ( ! [X4,X5,X3] : (cyclic(X3,X4,X3,X5)) ) | ~spl42_314),
  inference(resolution,[],[f16194,f381])).
fof(f381,plain,(
  ( ! [X2,X0,X3,X1] : (~cyclic(X0,X1,X2,X3) | cyclic(X0,X2,X1,X3)) )),
  inference(cnf_transformation,[],[f189])).
fof(f189,plain,(
  ! [X0,X1,X2,X3] : (cyclic(X0,X2,X1,X3) | ~cyclic(X0,X1,X2,X3))),
  inference(ennf_transformation,[],[f15])).
fof(f15,axiom,(
  ! [X0,X1,X2,X3] : (cyclic(X0,X1,X2,X3) => cyclic(X0,X2,X1,X3))),
  file('./6_GDD_FULL/21-40/26.p',ruleD15)).
fof(f20013,plain,(
  ( ! [X103,X105,X104] : (~cyclic(X103,X104,X103,X105) | cong(X103,X104,X103,X104) | ~cyclic(X103,X104,X105,X103) | ~cyclic(X103,X104,X105,X104)) )),
  inference(duplicate_literal_removal,[],[f20010])).
fof(f20010,plain,(
  ( ! [X103,X105,X104] : (~cyclic(X103,X104,X103,X105) | cong(X103,X104,X103,X104) | ~cyclic(X103,X104,X105,X103) | ~cyclic(X103,X104,X105,X104) | ~cyclic(X103,X104,X105,X103)) )),
  inference(resolution,[],[f3227,f410])).
fof(f410,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (~eqangle(X2,X0,X2,X1,X5,X3,X5,X4) | cong(X0,X1,X3,X4) | ~cyclic(X0,X1,X2,X5) | ~cyclic(X0,X1,X2,X4) | ~cyclic(X0,X1,X2,X3)) )),
  inference(cnf_transformation,[],[f227])).
fof(f227,plain,(
  ! [X0,X1,X2,X3,X4,X5] : (cong(X0,X1,X3,X4) | ~eqangle(X2,X0,X2,X1,X5,X3,X5,X4) | ~cyclic(X0,X1,X2,X5) | ~cyclic(X0,X1,X2,X4) | ~cyclic(X0,X1,X2,X3))),
  inference(flattening,[],[f226])).
fof(f226,plain,(
  ! [X0,X1,X2,X3,X4,X5] : (cong(X0,X1,X3,X4) | (~eqangle(X2,X0,X2,X1,X5,X3,X5,X4) | ~cyclic(X0,X1,X2,X5) | ~cyclic(X0,X1,X2,X4) | ~cyclic(X0,X1,X2,X3)))),
  inference(ennf_transformation,[],[f123])).
fof(f123,plain,(
  ! [X0,X1,X2,X3,X4,X5] : ((eqangle(X2,X0,X2,X1,X5,X3,X5,X4) & cyclic(X0,X1,X2,X5) & cyclic(X0,X1,X2,X4) & cyclic(X0,X1,X2,X3)) => cong(X0,X1,X3,X4))),
  inference(rectify,[],[f44])).
fof(f44,axiom,(
  ! [X0,X1,X2,X8,X9,X14] : ((eqangle(X2,X0,X2,X1,X14,X8,X14,X9) & cyclic(X0,X1,X2,X14) & cyclic(X0,X1,X2,X9) & cyclic(X0,X1,X2,X8)) => cong(X0,X1,X8,X9))),
  file('./6_GDD_FULL/21-40/26.p',ruleD43)).
fof(f3227,plain,(
  ( ! [X14,X12,X13,X11] : (eqangle(X14,X11,X14,X12,X11,X13,X13,X12) | ~cyclic(X11,X12,X13,X14)) )),
  inference(resolution,[],[f1056,f386])).
fof(f386,plain,(
  ( ! [X6,X4,X2,X0,X7,X5,X3,X1] : (~eqangle(X0,X1,X2,X3,X4,X5,X6,X7) | eqangle(X4,X5,X6,X7,X0,X1,X2,X3)) )),
  inference(cnf_transformation,[],[f195])).
fof(f195,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6,X7] : (eqangle(X4,X5,X6,X7,X0,X1,X2,X3) | ~eqangle(X0,X1,X2,X3,X4,X5,X6,X7))),
  inference(ennf_transformation,[],[f102])).
fof(f102,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6,X7] : (eqangle(X0,X1,X2,X3,X4,X5,X6,X7) => eqangle(X4,X5,X6,X7,X0,X1,X2,X3))),
  inference(rectify,[],[f20])).
fof(f20,axiom,(
  ! [X0,X1,X2,X3,X8,X9,X10,X11] : (eqangle(X0,X1,X2,X3,X8,X9,X10,X11) => eqangle(X8,X9,X10,X11,X0,X1,X2,X3))),
  file('./6_GDD_FULL/21-40/26.p',ruleD20)).
fof(f1056,plain,(
  ( ! [X6,X8,X7,X9] : (eqangle(X6,X7,X7,X8,X9,X6,X9,X8) | ~cyclic(X6,X8,X7,X9)) )),
  inference(resolution,[],[f384,f407])).
fof(f407,plain,(
  ( ! [X2,X0,X3,X1] : (eqangle(X2,X0,X2,X1,X3,X0,X3,X1) | ~cyclic(X0,X1,X2,X3)) )),
  inference(cnf_transformation,[],[f221])).
fof(f221,plain,(
  ! [X0,X1,X2,X3] : (eqangle(X2,X0,X2,X1,X3,X0,X3,X1) | ~cyclic(X0,X1,X2,X3))),
  inference(ennf_transformation,[],[f120])).
fof(f120,plain,(
  ! [X0,X1,X2,X3] : (cyclic(X0,X1,X2,X3) => eqangle(X2,X0,X2,X1,X3,X0,X3,X1))),
  inference(rectify,[],[f41])).
fof(f41,axiom,(
  ! [X0,X1,X8,X9] : (cyclic(X0,X1,X8,X9) => eqangle(X8,X0,X8,X1,X9,X0,X9,X1))),
  file('./6_GDD_FULL/21-40/26.p',ruleD41)).
fof(f384,plain,(
  ( ! [X6,X4,X2,X0,X7,X5,X3,X1] : (~eqangle(X0,X1,X2,X3,X4,X5,X6,X7) | eqangle(X1,X0,X2,X3,X4,X5,X6,X7)) )),
  inference(cnf_transformation,[],[f193])).
fof(f193,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6,X7] : (eqangle(X1,X0,X2,X3,X4,X5,X6,X7) | ~eqangle(X0,X1,X2,X3,X4,X5,X6,X7))),
  inference(ennf_transformation,[],[f100])).
fof(f100,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6,X7] : (eqangle(X0,X1,X2,X3,X4,X5,X6,X7) => eqangle(X1,X0,X2,X3,X4,X5,X6,X7))),
  inference(rectify,[],[f18])).
fof(f18,axiom,(
  ! [X0,X1,X2,X3,X8,X9,X10,X11] : (eqangle(X0,X1,X2,X3,X8,X9,X10,X11) => eqangle(X1,X0,X2,X3,X8,X9,X10,X11))),
  file('./6_GDD_FULL/21-40/26.p',ruleD18)).
fof(f44850,plain,(
  ( ! [X61,X64,X62,X63] : (midp(X61,X62,X63) | ~midp(X61,X64,X64)) ) | ~spl42_314),
  inference(subsumption_resolution,[],[f44750,f44717])).
fof(f44717,plain,(
  ( ! [X39,X37,X38] : (para(X39,X37,X39,X38)) ) | ~spl42_314),
  inference(subsumption_resolution,[],[f44716,f3321])).
fof(f44716,plain,(
  ( ! [X39,X37,X38] : (~para(X37,X38,X37,X38) | para(X39,X37,X39,X38)) ) | ~spl42_314),
  inference(subsumption_resolution,[],[f44715,f18553])).
fof(f44715,plain,(
  ( ! [X39,X37,X38] : (~cyclic(X37,X38,X37,X39) | ~para(X37,X38,X37,X38) | para(X39,X37,X39,X38)) ) | ~spl42_314),
  inference(subsumption_resolution,[],[f44670,f36612])).
fof(f36612,plain,(
  ( ! [X76,X74,X77,X75] : (cong(X74,X75,X76,X77)) ) | ~spl42_314),
  inference(subsumption_resolution,[],[f36611,f20071])).
fof(f36611,plain,(
  ( ! [X76,X74,X77,X75] : (~midp(X76,X77,X77) | cong(X74,X75,X76,X77)) ) | ~spl42_314),
  inference(subsumption_resolution,[],[f36601,f20071])).
fof(f36601,plain,(
  ( ! [X76,X74,X77,X75] : (~midp(X74,X75,X75) | ~midp(X76,X77,X77) | cong(X74,X75,X76,X77)) ) | ~spl42_314),
  inference(resolution,[],[f22823,f4023])).
fof(f4023,plain,(
  ( ! [X14,X12,X10,X15,X13,X11] : (~cong(X14,X15,X11,X12) | ~midp(X13,X14,X15) | ~midp(X10,X11,X12) | cong(X13,X14,X10,X11)) )),
  inference(resolution,[],[f1099,f442])).
fof(f442,plain,(
  ( ! [X6,X4,X2,X0,X7,X5,X3,X1] : (~eqratio(X0,X1,X2,X3,X4,X5,X6,X7) | ~cong(X4,X5,X6,X7) | cong(X0,X1,X2,X3)) )),
  inference(cnf_transformation,[],[f284])).
fof(f284,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6,X7] : (cong(X0,X1,X2,X3) | ~cong(X4,X5,X6,X7) | ~eqratio(X0,X1,X2,X3,X4,X5,X6,X7))),
  inference(flattening,[],[f283])).
fof(f283,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6,X7] : (cong(X0,X1,X2,X3) | (~cong(X4,X5,X6,X7) | ~eqratio(X0,X1,X2,X3,X4,X5,X6,X7)))),
  inference(ennf_transformation,[],[f148])).
fof(f148,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6,X7] : ((cong(X4,X5,X6,X7) & eqratio(X0,X1,X2,X3,X4,X5,X6,X7)) => cong(X0,X1,X2,X3))),
  inference(rectify,[],[f76])).
fof(f76,axiom,(
  ! [X0,X1,X2,X3,X8,X9,X10,X11] : ((cong(X8,X9,X10,X11) & eqratio(X0,X1,X2,X3,X8,X9,X10,X11)) => cong(X0,X1,X2,X3))),
  file('./6_GDD_FULL/21-40/26.p',ruleD75)).
fof(f1099,plain,(
  ( ! [X6,X10,X8,X7,X11,X9] : (eqratio(X6,X7,X8,X9,X7,X10,X9,X11) | ~midp(X8,X9,X11) | ~midp(X6,X7,X10)) )),
  inference(resolution,[],[f395,f437])).
fof(f437,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (eqratio(X4,X0,X0,X1,X5,X2,X2,X3) | ~midp(X5,X2,X3) | ~midp(X4,X0,X1)) )),
  inference(cnf_transformation,[],[f274])).
fof(f274,plain,(
  ! [X0,X1,X2,X3,X4,X5] : (eqratio(X4,X0,X0,X1,X5,X2,X2,X3) | ~midp(X5,X2,X3) | ~midp(X4,X0,X1))),
  inference(flattening,[],[f273])).
fof(f273,plain,(
  ! [X0,X1,X2,X3,X4,X5] : (eqratio(X4,X0,X0,X1,X5,X2,X2,X3) | (~midp(X5,X2,X3) | ~midp(X4,X0,X1)))),
  inference(ennf_transformation,[],[f145])).
fof(f145,plain,(
  ! [X0,X1,X2,X3,X4,X5] : ((midp(X5,X2,X3) & midp(X4,X0,X1)) => eqratio(X4,X0,X0,X1,X5,X2,X2,X3))),
  inference(rectify,[],[f71])).
fof(f71,axiom,(
  ! [X0,X1,X2,X3,X6,X16] : ((midp(X16,X2,X3) & midp(X6,X0,X1)) => eqratio(X6,X0,X0,X1,X16,X2,X2,X3))),
  file('./6_GDD_FULL/21-40/26.p',ruleD70)).
fof(f395,plain,(
  ( ! [X6,X4,X2,X0,X7,X5,X3,X1] : (~eqratio(X0,X1,X2,X3,X4,X5,X6,X7) | eqratio(X0,X1,X4,X5,X2,X3,X6,X7)) )),
  inference(cnf_transformation,[],[f206])).
fof(f206,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6,X7] : (eqratio(X0,X1,X4,X5,X2,X3,X6,X7) | ~eqratio(X0,X1,X2,X3,X4,X5,X6,X7))),
  inference(ennf_transformation,[],[f108])).
fof(f108,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6,X7] : (eqratio(X0,X1,X2,X3,X4,X5,X6,X7) => eqratio(X0,X1,X4,X5,X2,X3,X6,X7))),
  inference(rectify,[],[f29])).
fof(f29,axiom,(
  ! [X0,X1,X2,X3,X8,X9,X10,X11] : (eqratio(X0,X1,X2,X3,X8,X9,X10,X11) => eqratio(X0,X1,X8,X9,X2,X3,X10,X11))),
  file('./6_GDD_FULL/21-40/26.p',ruleD29)).
fof(f22823,plain,(
  ( ! [X103,X102] : (cong(X103,X103,X102,X102)) ) | ~spl42_314),
  inference(subsumption_resolution,[],[f22822,f16194])).
fof(f22822,plain,(
  ( ! [X103,X102,X104] : (cong(X103,X103,X102,X102) | ~cyclic(X103,X103,X104,X102)) ) | ~spl42_314),
  inference(subsumption_resolution,[],[f22821,f16194])).
fof(f22821,plain,(
  ( ! [X103,X102,X104] : (cong(X103,X103,X102,X102) | ~cyclic(X103,X103,X104,X104) | ~cyclic(X103,X103,X104,X102)) ) | ~spl42_314),
  inference(subsumption_resolution,[],[f22811,f18577])).
fof(f22811,plain,(
  ( ! [X103,X102,X104] : (~cyclic(X102,X103,X104,X104) | cong(X103,X103,X102,X102) | ~cyclic(X103,X103,X104,X104) | ~cyclic(X103,X103,X104,X102)) )),
  inference(duplicate_literal_removal,[],[f22807])).
fof(f22807,plain,(
  ( ! [X103,X102,X104] : (~cyclic(X102,X103,X104,X104) | cong(X103,X103,X102,X102) | ~cyclic(X103,X103,X104,X104) | ~cyclic(X103,X103,X104,X102) | ~cyclic(X103,X103,X104,X102)) )),
  inference(resolution,[],[f3335,f410])).
fof(f3335,plain,(
  ( ! [X17,X15,X18,X16] : (eqangle(X17,X16,X18,X16,X17,X15,X18,X15) | ~cyclic(X15,X16,X17,X18)) )),
  inference(resolution,[],[f1063,f387])).
fof(f1063,plain,(
  ( ! [X6,X8,X7,X9] : (eqangle(X6,X7,X6,X8,X9,X7,X9,X8) | ~cyclic(X8,X7,X6,X9)) )),
  inference(resolution,[],[f385,f407])).
fof(f385,plain,(
  ( ! [X6,X4,X2,X0,X7,X5,X3,X1] : (~eqangle(X0,X1,X2,X3,X4,X5,X6,X7) | eqangle(X2,X3,X0,X1,X6,X7,X4,X5)) )),
  inference(cnf_transformation,[],[f194])).
fof(f194,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6,X7] : (eqangle(X2,X3,X0,X1,X6,X7,X4,X5) | ~eqangle(X0,X1,X2,X3,X4,X5,X6,X7))),
  inference(ennf_transformation,[],[f101])).
fof(f101,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6,X7] : (eqangle(X0,X1,X2,X3,X4,X5,X6,X7) => eqangle(X2,X3,X0,X1,X6,X7,X4,X5))),
  inference(rectify,[],[f19])).
fof(f19,axiom,(
  ! [X0,X1,X2,X3,X8,X9,X10,X11] : (eqangle(X0,X1,X2,X3,X8,X9,X10,X11) => eqangle(X2,X3,X0,X1,X10,X11,X8,X9))),
  file('./6_GDD_FULL/21-40/26.p',ruleD19)).
fof(f44670,plain,(
  ( ! [X39,X37,X38] : (~cong(X37,X37,X37,X38) | ~cyclic(X37,X38,X37,X39) | ~para(X37,X38,X37,X38) | para(X39,X37,X39,X38)) )),
  inference(resolution,[],[f6089,f440])).
fof(f440,plain,(
  ( ! [X6,X4,X2,X0,X7,X5,X3,X1] : (~eqangle(X0,X1,X2,X3,X4,X5,X6,X7) | ~para(X4,X5,X6,X7) | para(X0,X1,X2,X3)) )),
  inference(cnf_transformation,[],[f280])).
fof(f280,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6,X7] : (para(X0,X1,X2,X3) | ~para(X4,X5,X6,X7) | ~eqangle(X0,X1,X2,X3,X4,X5,X6,X7))),
  inference(flattening,[],[f279])).
fof(f279,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6,X7] : (para(X0,X1,X2,X3) | (~para(X4,X5,X6,X7) | ~eqangle(X0,X1,X2,X3,X4,X5,X6,X7)))),
  inference(ennf_transformation,[],[f146])).
fof(f146,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6,X7] : ((para(X4,X5,X6,X7) & eqangle(X0,X1,X2,X3,X4,X5,X6,X7)) => para(X0,X1,X2,X3))),
  inference(rectify,[],[f74])).
fof(f74,axiom,(
  ! [X0,X1,X2,X3,X8,X9,X10,X11] : ((para(X8,X9,X10,X11) & eqangle(X0,X1,X2,X3,X8,X9,X10,X11)) => para(X0,X1,X2,X3))),
  file('./6_GDD_FULL/21-40/26.p',ruleD73)).
fof(f6089,plain,(
  ( ! [X57,X58,X56] : (eqangle(X56,X57,X56,X58,X57,X58,X57,X58) | ~cong(X57,X57,X57,X58) | ~cyclic(X57,X58,X57,X56)) )),
  inference(resolution,[],[f1438,f1070])).
fof(f1070,plain,(
  ( ! [X6,X8,X7,X9] : (eqangle(X6,X7,X6,X8,X9,X7,X9,X8) | ~cyclic(X7,X8,X9,X6)) )),
  inference(resolution,[],[f386,f407])).
fof(f1438,plain,(
  ( ! [X39,X37,X35,X41,X38,X36,X40] : (~eqangle(X35,X36,X37,X38,X41,X39,X39,X40) | eqangle(X35,X36,X37,X38,X39,X40,X41,X40) | ~cong(X41,X39,X41,X40)) )),
  inference(resolution,[],[f388,f413])).
fof(f413,plain,(
  ( ! [X2,X0,X1] : (eqangle(X2,X0,X0,X1,X0,X1,X2,X1) | ~cong(X2,X0,X2,X1)) )),
  inference(cnf_transformation,[],[f232])).
fof(f232,plain,(
  ! [X0,X1,X2] : (eqangle(X2,X0,X0,X1,X0,X1,X2,X1) | ~cong(X2,X0,X2,X1))),
  inference(ennf_transformation,[],[f126])).
fof(f126,plain,(
  ! [X0,X1,X2] : (cong(X2,X0,X2,X1) => eqangle(X2,X0,X0,X1,X0,X1,X2,X1))),
  inference(rectify,[],[f47])).
fof(f47,axiom,(
  ! [X0,X1,X7] : (cong(X7,X0,X7,X1) => eqangle(X7,X0,X0,X1,X0,X1,X7,X1))),
  file('./6_GDD_FULL/21-40/26.p',ruleD46)).
fof(f388,plain,(
  ( ! [X6,X4,X2,X0,X10,X8,X7,X5,X3,X1,X11,X9] : (~eqangle(X4,X5,X6,X7,X8,X9,X10,X11) | eqangle(X0,X1,X2,X3,X8,X9,X10,X11) | ~eqangle(X0,X1,X2,X3,X4,X5,X6,X7)) )),
  inference(cnf_transformation,[],[f198])).
fof(f198,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11] : (eqangle(X0,X1,X2,X3,X8,X9,X10,X11) | ~eqangle(X4,X5,X6,X7,X8,X9,X10,X11) | ~eqangle(X0,X1,X2,X3,X4,X5,X6,X7))),
  inference(flattening,[],[f197])).
fof(f197,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11] : (eqangle(X0,X1,X2,X3,X8,X9,X10,X11) | (~eqangle(X4,X5,X6,X7,X8,X9,X10,X11) | ~eqangle(X0,X1,X2,X3,X4,X5,X6,X7)))),
  inference(ennf_transformation,[],[f104])).
fof(f104,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11] : ((eqangle(X4,X5,X6,X7,X8,X9,X10,X11) & eqangle(X0,X1,X2,X3,X4,X5,X6,X7)) => eqangle(X0,X1,X2,X3,X8,X9,X10,X11))),
  inference(rectify,[],[f22])).
fof(f22,axiom,(
  ! [X0,X1,X2,X3,X8,X9,X10,X11,X4,X5,X12,X13] : ((eqangle(X8,X9,X10,X11,X4,X5,X12,X13) & eqangle(X0,X1,X2,X3,X8,X9,X10,X11)) => eqangle(X0,X1,X2,X3,X4,X5,X12,X13))),
  file('./6_GDD_FULL/21-40/26.p',ruleD22)).
fof(f44750,plain,(
  ( ! [X61,X64,X62,X63] : (midp(X61,X62,X63) | ~para(X64,X62,X64,X63) | ~midp(X61,X64,X64)) ) | ~spl42_314),
  inference(resolution,[],[f44717,f431])).
fof(f431,plain,(
  ( ! [X4,X2,X0,X3,X1] : (~para(X0,X3,X1,X2) | midp(X4,X2,X3) | ~para(X0,X2,X1,X3) | ~midp(X4,X0,X1)) )),
  inference(cnf_transformation,[],[f265])).
fof(f265,plain,(
  ! [X0,X1,X2,X3,X4] : (midp(X4,X2,X3) | ~para(X0,X3,X1,X2) | ~para(X0,X2,X1,X3) | ~midp(X4,X0,X1))),
  inference(flattening,[],[f264])).
fof(f264,plain,(
  ! [X0,X1,X2,X3,X4] : (midp(X4,X2,X3) | (~para(X0,X3,X1,X2) | ~para(X0,X2,X1,X3) | ~midp(X4,X0,X1)))),
  inference(ennf_transformation,[],[f143])).
fof(f143,plain,(
  ! [X0,X1,X2,X3,X4] : ((para(X0,X3,X1,X2) & para(X0,X2,X1,X3) & midp(X4,X0,X1)) => midp(X4,X2,X3))),
  inference(rectify,[],[f65])).
fof(f65,axiom,(
  ! [X0,X1,X2,X3,X6] : ((para(X0,X3,X1,X2) & para(X0,X2,X1,X3) & midp(X6,X0,X1)) => midp(X6,X2,X3))),
  file('./6_GDD_FULL/21-40/26.p',ruleD64)).
fof(f44905,plain,(
  ( ! [X6,X10,X8,X7,X9] : (para(X6,X7,X8,X9) | ~midp(X6,X10,X8)) ) | ~spl42_314),
  inference(resolution,[],[f44851,f411])).
fof(f411,plain,(
  ( ! [X4,X2,X0,X3,X1] : (~midp(X4,X0,X2) | para(X3,X4,X1,X2) | ~midp(X3,X0,X1)) )),
  inference(cnf_transformation,[],[f229])).
fof(f229,plain,(
  ! [X0,X1,X2,X3,X4] : (para(X3,X4,X1,X2) | ~midp(X4,X0,X2) | ~midp(X3,X0,X1))),
  inference(flattening,[],[f228])).
fof(f228,plain,(
  ! [X0,X1,X2,X3,X4] : (para(X3,X4,X1,X2) | (~midp(X4,X0,X2) | ~midp(X3,X0,X1)))),
  inference(ennf_transformation,[],[f124])).
fof(f124,plain,(
  ! [X0,X1,X2,X3,X4] : ((midp(X4,X0,X2) & midp(X3,X0,X1)) => para(X3,X4,X1,X2))),
  inference(rectify,[],[f45])).
fof(f45,axiom,(
  ! [X0,X1,X2,X4,X5] : ((midp(X5,X0,X2) & midp(X4,X0,X1)) => para(X4,X5,X1,X2))),
  file('./6_GDD_FULL/21-40/26.p',ruleD44)).
fof(f57104,plain,(
  ( ! [X251,X249,X246,X244,X250,X248,X247,X245] : (eqangle(X244,X245,X246,X247,X248,X249,X250,X251) | ~para(X250,X251,X246,X247)) ) | ~spl42_314),
  inference(subsumption_resolution,[],[f57008,f45683])).
fof(f57008,plain,(
  ( ! [X251,X249,X246,X244,X250,X248,X247,X245] : (eqangle(X244,X245,X246,X247,X248,X249,X250,X251) | ~para(X245,X244,X248,X249) | ~para(X250,X251,X246,X247)) )),
  inference(resolution,[],[f3187,f3295])).
fof(f3295,plain,(
  ( ! [X21,X19,X17,X20,X18,X16] : (eqangle(X20,X21,X18,X19,X20,X21,X16,X17) | ~para(X16,X17,X18,X19)) )),
  inference(resolution,[],[f1062,f386])).
fof(f1062,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (eqangle(X0,X1,X2,X3,X0,X1,X4,X5) | ~para(X2,X3,X4,X5)) )),
  inference(resolution,[],[f385,f406])).
fof(f3187,plain,(
  ( ! [X30,X28,X37,X35,X33,X31,X29,X36,X34,X32] : (~eqangle(X32,X33,X34,X35,X29,X28,X36,X37) | eqangle(X32,X33,X34,X35,X30,X31,X36,X37) | ~para(X28,X29,X30,X31)) )),
  inference(resolution,[],[f1055,f388])).
fof(f1055,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (eqangle(X0,X1,X2,X3,X4,X5,X2,X3) | ~para(X1,X0,X4,X5)) )),
  inference(resolution,[],[f384,f406])).
fof(f571,plain,(
  ~eqangle(sK34,sK35,sK35,sK36,sK40,sK39,sK39,sK41) | ~spl42_13),
  inference(avatar_component_clause,[],[f570])).
fof(f570,plain,(
  spl42_13 <=> ~eqangle(sK34,sK35,sK35,sK36,sK40,sK39,sK39,sK41)),
  introduced(avatar_definition,[new_symbols(naming,[spl42_13])])).
fof(f60353,plain,(
  spl42_11 | ~spl42_314),
  inference(avatar_contradiction_clause,[],[f60352])).
fof(f60352,plain,(
  $false | (~spl42_11 | ~spl42_314)),
  inference(subsumption_resolution,[],[f562,f57105])).
fof(f562,plain,(
  ~eqangle(sK34,sK35,sK35,sK36,sK41,sK40,sK40,sK39) | ~spl42_11),
  inference(avatar_component_clause,[],[f561])).
fof(f561,plain,(
  spl42_11 <=> ~eqangle(sK34,sK35,sK35,sK36,sK41,sK40,sK40,sK39)),
  introduced(avatar_definition,[new_symbols(naming,[spl42_11])])).
fof(f60351,plain,(
  spl42_7 | ~spl42_314),
  inference(avatar_contradiction_clause,[],[f60350])).
fof(f60350,plain,(
  $false | (~spl42_7 | ~spl42_314)),
  inference(subsumption_resolution,[],[f548,f57105])).
fof(f548,plain,(
  ~eqangle(sK34,sK35,sK35,sK36,sK40,sK41,sK41,sK39) | ~spl42_7),
  inference(avatar_component_clause,[],[f547])).
fof(f547,plain,(
  spl42_7 <=> ~eqangle(sK34,sK35,sK35,sK36,sK40,sK41,sK41,sK39)),
  introduced(avatar_definition,[new_symbols(naming,[spl42_7])])).
fof(f60349,plain,(
  spl42_1 | ~spl42_314),
  inference(avatar_contradiction_clause,[],[f60348])).
fof(f60348,plain,(
  $false | (~spl42_1 | ~spl42_314)),
  inference(subsumption_resolution,[],[f60347,f20229])).
fof(f20229,plain,(
  ( ! [X100] : (midp(X100,sK35,sK34)) ) | ~spl42_314),
  inference(subsumption_resolution,[],[f20180,f20071])).
fof(f20180,plain,(
  ( ! [X101,X100] : (midp(X100,sK35,sK34) | ~midp(X100,X101,X101)) ) | ~spl42_314),
  inference(resolution,[],[f20071,f2191])).
fof(f2191,plain,(
  ( ! [X2,X3,X1] : (~midp(sK29,X1,X2) | midp(X3,sK35,sK34) | ~midp(X3,X1,X2)) )),
  inference(subsumption_resolution,[],[f2186,f715])).
fof(f715,plain,(
  ( ! [X19,X18] : (para(X18,sK35,X19,sK34) | ~midp(sK29,X18,X19)) )),
  inference(resolution,[],[f430,f587])).
fof(f587,plain,(
  midp(sK29,sK35,sK34)),
  inference(resolution,[],[f377,f495])).
fof(f495,plain,(
  midp(sK29,sK34,sK35)),
  inference(cnf_transformation,[],[f366])).
fof(f366,plain,(
  (~eqangle(sK34,sK35,sK35,sK36,sK40,sK41,sK41,sK39) | (~eqangle(sK35,sK34,sK34,sK36,sK40,sK39,sK39,sK41) & ~eqangle(sK34,sK35,sK35,sK36,sK41,sK40,sK40,sK39)) | (~eqangle(sK35,sK34,sK34,sK36,sK40,sK39,sK39,sK41) & ~eqangle(sK34,sK35,sK35,sK36,sK40,sK39,sK39,sK41)) | (~eqangle(sK35,sK34,sK34,sK36,sK40,sK41,sK41,sK39) & ~eqangle(sK34,sK35,sK35,sK36,sK41,sK40,sK40,sK39)) | sP1(sK39,sK41,sK40,sK36,sK34,sK35) | sP0(sK39,sK40,sK41,sK36,sK34,sK35) | ~eqangle(sK35,sK34,sK34,sK36,sK41,sK40,sK40,sK39)) & perp(sK37,sK34,sK22,sK41) & midp(sK22,sK37,sK34) & perp(sK36,sK34,sK23,sK41) & midp(sK23,sK36,sK34) & perp(sK36,sK37,sK24,sK41) & midp(sK24,sK36,sK37) & perp(sK37,sK36,sK25,sK40) & midp(sK25,sK37,sK36) & perp(sK35,sK36,sK26,sK40) & midp(sK26,sK35,sK36) & perp(sK35,sK37,sK27,sK40) & midp(sK27,sK35,sK37) & perp(sK37,sK35,sK28,sK39) & midp(sK28,sK37,sK35) & perp(sK34,sK35,sK29,sK39) & midp(sK29,sK34,sK35) & perp(sK34,sK37,sK30,sK39) & midp(sK30,sK34,sK37) & perp(sK35,sK36,sK31,sK38) & midp(sK31,sK35,sK36) & perp(sK34,sK36,sK32,sK38) & midp(sK32,sK34,sK36) & perp(sK34,sK35,sK33,sK38) & midp(sK33,sK34,sK35) & perp(sK35,sK36,sK34,sK37) & perp(sK34,sK36,sK35,sK37) & perp(sK34,sK35,sK36,sK37)),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK22,sK23,sK24,sK25,sK26,sK27,sK28,sK29,sK30,sK31,sK32,sK33,sK34,sK35,sK36,sK37,sK38,sK39,sK40,sK41])],[f324,f365])).
fof(f365,plain,(
  ? [X0,X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19] : ((~eqangle(X12,X13,X13,X14,X18,X19,X19,X17) | (~eqangle(X13,X12,X12,X14,X18,X17,X17,X19) & ~eqangle(X12,X13,X13,X14,X19,X18,X18,X17)) | (~eqangle(X13,X12,X12,X14,X18,X17,X17,X19) & ~eqangle(X12,X13,X13,X14,X18,X17,X17,X19)) | (~eqangle(X13,X12,X12,X14,X18,X19,X19,X17) & ~eqangle(X12,X13,X13,X14,X19,X18,X18,X17)) | sP1(X17,X19,X18,X14,X12,X13) | sP0(X17,X18,X19,X14,X12,X13) | ~eqangle(X13,X12,X12,X14,X19,X18,X18,X17)) & perp(X15,X12,X0,X19) & midp(X0,X15,X12) & perp(X14,X12,X1,X19) & midp(X1,X14,X12) & perp(X14,X15,X2,X19) & midp(X2,X14,X15) & perp(X15,X14,X3,X18) & midp(X3,X15,X14) & perp(X13,X14,X4,X18) & midp(X4,X13,X14) & perp(X13,X15,X5,X18) & midp(X5,X13,X15) & perp(X15,X13,X6,X17) & midp(X6,X15,X13) & perp(X12,X13,X7,X17) & midp(X7,X12,X13) & perp(X12,X15,X8,X17) & midp(X8,X12,X15) & perp(X13,X14,X9,X16) & midp(X9,X13,X14) & perp(X12,X14,X10,X16) & midp(X10,X12,X14) & perp(X12,X13,X11,X16) & midp(X11,X12,X13) & perp(X13,X14,X12,X15) & perp(X12,X14,X13,X15) & perp(X12,X13,X14,X15)) => ((~eqangle(sK34,sK35,sK35,sK36,sK40,sK41,sK41,sK39) | (~eqangle(sK35,sK34,sK34,sK36,sK40,sK39,sK39,sK41) & ~eqangle(sK34,sK35,sK35,sK36,sK41,sK40,sK40,sK39)) | (~eqangle(sK35,sK34,sK34,sK36,sK40,sK39,sK39,sK41) & ~eqangle(sK34,sK35,sK35,sK36,sK40,sK39,sK39,sK41)) | (~eqangle(sK35,sK34,sK34,sK36,sK40,sK41,sK41,sK39) & ~eqangle(sK34,sK35,sK35,sK36,sK41,sK40,sK40,sK39)) | sP1(sK39,sK41,sK40,sK36,sK34,sK35) | sP0(sK39,sK40,sK41,sK36,sK34,sK35) | ~eqangle(sK35,sK34,sK34,sK36,sK41,sK40,sK40,sK39)) & perp(sK37,sK34,sK22,sK41) & midp(sK22,sK37,sK34) & perp(sK36,sK34,sK23,sK41) & midp(sK23,sK36,sK34) & perp(sK36,sK37,sK24,sK41) & midp(sK24,sK36,sK37) & perp(sK37,sK36,sK25,sK40) & midp(sK25,sK37,sK36) & perp(sK35,sK36,sK26,sK40) & midp(sK26,sK35,sK36) & perp(sK35,sK37,sK27,sK40) & midp(sK27,sK35,sK37) & perp(sK37,sK35,sK28,sK39) & midp(sK28,sK37,sK35) & perp(sK34,sK35,sK29,sK39) & midp(sK29,sK34,sK35) & perp(sK34,sK37,sK30,sK39) & midp(sK30,sK34,sK37) & perp(sK35,sK36,sK31,sK38) & midp(sK31,sK35,sK36) & perp(sK34,sK36,sK32,sK38) & midp(sK32,sK34,sK36) & perp(sK34,sK35,sK33,sK38) & midp(sK33,sK34,sK35) & perp(sK35,sK36,sK34,sK37) & perp(sK34,sK36,sK35,sK37) & perp(sK34,sK35,sK36,sK37))),
  introduced(choice_axiom,[])).
fof(f324,plain,(
  ? [X0,X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19] : ((~eqangle(X12,X13,X13,X14,X18,X19,X19,X17) | (~eqangle(X13,X12,X12,X14,X18,X17,X17,X19) & ~eqangle(X12,X13,X13,X14,X19,X18,X18,X17)) | (~eqangle(X13,X12,X12,X14,X18,X17,X17,X19) & ~eqangle(X12,X13,X13,X14,X18,X17,X17,X19)) | (~eqangle(X13,X12,X12,X14,X18,X19,X19,X17) & ~eqangle(X12,X13,X13,X14,X19,X18,X18,X17)) | sP1(X17,X19,X18,X14,X12,X13) | sP0(X17,X18,X19,X14,X12,X13) | ~eqangle(X13,X12,X12,X14,X19,X18,X18,X17)) & perp(X15,X12,X0,X19) & midp(X0,X15,X12) & perp(X14,X12,X1,X19) & midp(X1,X14,X12) & perp(X14,X15,X2,X19) & midp(X2,X14,X15) & perp(X15,X14,X3,X18) & midp(X3,X15,X14) & perp(X13,X14,X4,X18) & midp(X4,X13,X14) & perp(X13,X15,X5,X18) & midp(X5,X13,X15) & perp(X15,X13,X6,X17) & midp(X6,X15,X13) & perp(X12,X13,X7,X17) & midp(X7,X12,X13) & perp(X12,X15,X8,X17) & midp(X8,X12,X15) & perp(X13,X14,X9,X16) & midp(X9,X13,X14) & perp(X12,X14,X10,X16) & midp(X10,X12,X14) & perp(X12,X13,X11,X16) & midp(X11,X12,X13) & perp(X13,X14,X12,X15) & perp(X12,X14,X13,X15) & perp(X12,X13,X14,X15))),
  inference(definition_folding,[],[f321,f323,f322])).
fof(f322,plain,(
  ! [X17,X18,X19,X14,X12,X13] : ((~eqangle(X13,X12,X12,X14,X19,X18,X18,X17) & ~eqangle(X12,X13,X13,X14,X18,X19,X19,X17)) | ~sP0(X17,X18,X19,X14,X12,X13))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP0])])).
fof(f323,plain,(
  ! [X17,X19,X18,X14,X12,X13] : ((~eqangle(X13,X12,X12,X14,X18,X19,X19,X17) & ~eqangle(X12,X13,X13,X14,X18,X17,X17,X19)) | ~sP1(X17,X19,X18,X14,X12,X13))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP1])])).
fof(f321,plain,(
  ? [X0,X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19] : ((~eqangle(X12,X13,X13,X14,X18,X19,X19,X17) | (~eqangle(X13,X12,X12,X14,X18,X17,X17,X19) & ~eqangle(X12,X13,X13,X14,X19,X18,X18,X17)) | (~eqangle(X13,X12,X12,X14,X18,X17,X17,X19) & ~eqangle(X12,X13,X13,X14,X18,X17,X17,X19)) | (~eqangle(X13,X12,X12,X14,X18,X19,X19,X17) & ~eqangle(X12,X13,X13,X14,X19,X18,X18,X17)) | (~eqangle(X13,X12,X12,X14,X18,X19,X19,X17) & ~eqangle(X12,X13,X13,X14,X18,X17,X17,X19)) | (~eqangle(X13,X12,X12,X14,X19,X18,X18,X17) & ~eqangle(X12,X13,X13,X14,X18,X19,X19,X17)) | ~eqangle(X13,X12,X12,X14,X19,X18,X18,X17)) & perp(X15,X12,X0,X19) & midp(X0,X15,X12) & perp(X14,X12,X1,X19) & midp(X1,X14,X12) & perp(X14,X15,X2,X19) & midp(X2,X14,X15) & perp(X15,X14,X3,X18) & midp(X3,X15,X14) & perp(X13,X14,X4,X18) & midp(X4,X13,X14) & perp(X13,X15,X5,X18) & midp(X5,X13,X15) & perp(X15,X13,X6,X17) & midp(X6,X15,X13) & perp(X12,X13,X7,X17) & midp(X7,X12,X13) & perp(X12,X15,X8,X17) & midp(X8,X12,X15) & perp(X13,X14,X9,X16) & midp(X9,X13,X14) & perp(X12,X14,X10,X16) & midp(X10,X12,X14) & perp(X12,X13,X11,X16) & midp(X11,X12,X13) & perp(X13,X14,X12,X15) & perp(X12,X14,X13,X15) & perp(X12,X13,X14,X15))),
  inference(flattening,[],[f320])).
fof(f320,plain,(
  ? [X0,X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19] : ((~eqangle(X12,X13,X13,X14,X18,X19,X19,X17) | (~eqangle(X13,X12,X12,X14,X18,X17,X17,X19) & ~eqangle(X12,X13,X13,X14,X19,X18,X18,X17)) | (~eqangle(X13,X12,X12,X14,X18,X17,X17,X19) & ~eqangle(X12,X13,X13,X14,X18,X17,X17,X19)) | (~eqangle(X13,X12,X12,X14,X18,X19,X19,X17) & ~eqangle(X12,X13,X13,X14,X19,X18,X18,X17)) | (~eqangle(X13,X12,X12,X14,X18,X19,X19,X17) & ~eqangle(X12,X13,X13,X14,X18,X17,X17,X19)) | (~eqangle(X13,X12,X12,X14,X19,X18,X18,X17) & ~eqangle(X12,X13,X13,X14,X18,X19,X19,X17)) | ~eqangle(X13,X12,X12,X14,X19,X18,X18,X17)) & (perp(X15,X12,X0,X19) & midp(X0,X15,X12) & perp(X14,X12,X1,X19) & midp(X1,X14,X12) & perp(X14,X15,X2,X19) & midp(X2,X14,X15) & perp(X15,X14,X3,X18) & midp(X3,X15,X14) & perp(X13,X14,X4,X18) & midp(X4,X13,X14) & perp(X13,X15,X5,X18) & midp(X5,X13,X15) & perp(X15,X13,X6,X17) & midp(X6,X15,X13) & perp(X12,X13,X7,X17) & midp(X7,X12,X13) & perp(X12,X15,X8,X17) & midp(X8,X12,X15) & perp(X13,X14,X9,X16) & midp(X9,X13,X14) & perp(X12,X14,X10,X16) & midp(X10,X12,X14) & perp(X12,X13,X11,X16) & midp(X11,X12,X13) & perp(X13,X14,X12,X15) & perp(X12,X14,X13,X15) & perp(X12,X13,X14,X15)))),
  inference(ennf_transformation,[],[f168])).
fof(f168,plain,(
  ~! [X0,X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19] : ((perp(X15,X12,X0,X19) & midp(X0,X15,X12) & perp(X14,X12,X1,X19) & midp(X1,X14,X12) & perp(X14,X15,X2,X19) & midp(X2,X14,X15) & perp(X15,X14,X3,X18) & midp(X3,X15,X14) & perp(X13,X14,X4,X18) & midp(X4,X13,X14) & perp(X13,X15,X5,X18) & midp(X5,X13,X15) & perp(X15,X13,X6,X17) & midp(X6,X15,X13) & perp(X12,X13,X7,X17) & midp(X7,X12,X13) & perp(X12,X15,X8,X17) & midp(X8,X12,X15) & perp(X13,X14,X9,X16) & midp(X9,X13,X14) & perp(X12,X14,X10,X16) & midp(X10,X12,X14) & perp(X12,X13,X11,X16) & midp(X11,X12,X13) & perp(X13,X14,X12,X15) & perp(X12,X14,X13,X15) & perp(X12,X13,X14,X15)) => (eqangle(X12,X13,X13,X14,X18,X19,X19,X17) & (eqangle(X13,X12,X12,X14,X18,X17,X17,X19) | eqangle(X12,X13,X13,X14,X19,X18,X18,X17)) & (eqangle(X13,X12,X12,X14,X18,X17,X17,X19) | eqangle(X12,X13,X13,X14,X18,X17,X17,X19)) & (eqangle(X13,X12,X12,X14,X18,X19,X19,X17) | eqangle(X12,X13,X13,X14,X19,X18,X18,X17)) & (eqangle(X13,X12,X12,X14,X18,X19,X19,X17) | eqangle(X12,X13,X13,X14,X18,X17,X17,X19)) & (eqangle(X13,X12,X12,X14,X19,X18,X18,X17) | eqangle(X12,X13,X13,X14,X18,X19,X19,X17)) & eqangle(X13,X12,X12,X14,X19,X18,X18,X17)))),
  inference(flattening,[],[f167])).
fof(f167,plain,(
  ~! [X0,X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11] : ! [X12,X13,X14,X15,X16,X17,X18,X19] : ((perp(X15,X12,X0,X19) & midp(X0,X15,X12) & perp(X14,X12,X1,X19) & midp(X1,X14,X12) & perp(X14,X15,X2,X19) & midp(X2,X14,X15) & perp(X15,X14,X3,X18) & midp(X3,X15,X14) & perp(X13,X14,X4,X18) & midp(X4,X13,X14) & perp(X13,X15,X5,X18) & midp(X5,X13,X15) & perp(X15,X13,X6,X17) & midp(X6,X15,X13) & perp(X12,X13,X7,X17) & midp(X7,X12,X13) & perp(X12,X15,X8,X17) & midp(X8,X12,X15) & perp(X13,X14,X9,X16) & midp(X9,X13,X14) & perp(X12,X14,X10,X16) & midp(X10,X12,X14) & perp(X12,X13,X11,X16) & midp(X11,X12,X13) & perp(X13,X14,X12,X15) & perp(X12,X14,X13,X15) & perp(X12,X13,X14,X15)) => (eqangle(X12,X13,X13,X14,X18,X19,X19,X17) & (eqangle(X13,X12,X12,X14,X18,X17,X17,X19) | eqangle(X12,X13,X13,X14,X19,X18,X18,X17)) & (eqangle(X13,X12,X12,X14,X18,X17,X17,X19) | eqangle(X12,X13,X13,X14,X18,X17,X17,X19)) & (eqangle(X13,X12,X12,X14,X18,X19,X19,X17) | eqangle(X12,X13,X13,X14,X19,X18,X18,X17)) & (eqangle(X13,X12,X12,X14,X18,X19,X19,X17) | eqangle(X12,X13,X13,X14,X18,X17,X17,X19)) & (eqangle(X13,X12,X12,X14,X19,X18,X18,X17) | eqangle(X12,X13,X13,X14,X18,X19,X19,X17)) & eqangle(X13,X12,X12,X14,X19,X18,X18,X17)))),
  inference(rectify,[],[f96])).
fof(f96,negated_conjecture,(
  ~! [X32,X31,X30,X29,X28,X27,X26,X25,X24,X23,X22,X21] : ! [X0,X1,X2,X13,X7,X18,X19,X20] : ((perp(X13,X0,X32,X20) & midp(X32,X13,X0) & perp(X2,X0,X31,X20) & midp(X31,X2,X0) & perp(X2,X13,X30,X20) & midp(X30,X2,X13) & perp(X13,X2,X29,X19) & midp(X29,X13,X2) & perp(X1,X2,X28,X19) & midp(X28,X1,X2) & perp(X1,X13,X27,X19) & midp(X27,X1,X13) & perp(X13,X1,X26,X18) & midp(X26,X13,X1) & perp(X0,X1,X25,X18) & midp(X25,X0,X1) & perp(X0,X13,X24,X18) & midp(X24,X0,X13) & perp(X1,X2,X23,X7) & midp(X23,X1,X2) & perp(X0,X2,X22,X7) & midp(X22,X0,X2) & perp(X0,X1,X21,X7) & midp(X21,X0,X1) & perp(X1,X2,X0,X13) & perp(X0,X2,X1,X13) & perp(X0,X1,X2,X13)) => (eqangle(X0,X1,X1,X2,X19,X20,X20,X18) & (eqangle(X1,X0,X0,X2,X19,X18,X18,X20) | eqangle(X0,X1,X1,X2,X20,X19,X19,X18)) & (eqangle(X1,X0,X0,X2,X19,X18,X18,X20) | eqangle(X0,X1,X1,X2,X19,X18,X18,X20)) & (eqangle(X1,X0,X0,X2,X19,X20,X20,X18) | eqangle(X0,X1,X1,X2,X20,X19,X19,X18)) & (eqangle(X1,X0,X0,X2,X19,X20,X20,X18) | eqangle(X0,X1,X1,X2,X19,X18,X18,X20)) & (eqangle(X1,X0,X0,X2,X20,X19,X19,X18) | eqangle(X0,X1,X1,X2,X19,X20,X20,X18)) & eqangle(X1,X0,X0,X2,X20,X19,X19,X18)))),
  inference(negated_conjecture,[],[f95])).
fof(f95,conjecture,(
  ( ! [X30,X28,X26,X24,X23,X21,X31,X29,X27,X25,X32,X22] : (! [X0,X1,X2,X13,X7,X18,X19,X20] : ((perp(X13,X0,X32,X20) & midp(X32,X13,X0) & perp(X2,X0,X31,X20) & midp(X31,X2,X0) & perp(X2,X13,X30,X20) & midp(X30,X2,X13) & perp(X13,X2,X29,X19) & midp(X29,X13,X2) & perp(X1,X2,X28,X19) & midp(X28,X1,X2) & perp(X1,X13,X27,X19) & midp(X27,X1,X13) & perp(X13,X1,X26,X18) & midp(X26,X13,X1) & perp(X0,X1,X25,X18) & midp(X25,X0,X1) & perp(X0,X13,X24,X18) & midp(X24,X0,X13) & perp(X1,X2,X23,X7) & midp(X23,X1,X2) & perp(X0,X2,X22,X7) & midp(X22,X0,X2) & perp(X0,X1,X21,X7) & midp(X21,X0,X1) & perp(X1,X2,X0,X13) & perp(X0,X2,X1,X13) & perp(X0,X1,X2,X13)) => (eqangle(X0,X1,X1,X2,X19,X20,X20,X18) & (eqangle(X1,X0,X0,X2,X19,X18,X18,X20) | eqangle(X0,X1,X1,X2,X20,X19,X19,X18)) & (eqangle(X1,X0,X0,X2,X19,X18,X18,X20) | eqangle(X0,X1,X1,X2,X19,X18,X18,X20)) & (eqangle(X1,X0,X0,X2,X19,X20,X20,X18) | eqangle(X0,X1,X1,X2,X20,X19,X19,X18)) & (eqangle(X1,X0,X0,X2,X19,X20,X20,X18) | eqangle(X0,X1,X1,X2,X19,X18,X18,X20)) & (eqangle(X1,X0,X0,X2,X20,X19,X19,X18) | eqangle(X0,X1,X1,X2,X19,X20,X20,X18)) & eqangle(X1,X0,X0,X2,X20,X19,X19,X18)))) )),
  file('./6_GDD_FULL/21-40/26.p',exemplo6GDDFULL214026)).
fof(f377,plain,(
  ( ! [X2,X0,X1] : (~midp(X2,X1,X0) | midp(X2,X0,X1)) )),
  inference(cnf_transformation,[],[f183])).
fof(f183,plain,(
  ! [X0,X1,X2] : (midp(X2,X0,X1) | ~midp(X2,X1,X0))),
  inference(ennf_transformation,[],[f97])).
fof(f97,plain,(
  ! [X0,X1,X2] : (midp(X2,X1,X0) => midp(X2,X0,X1))),
  inference(rectify,[],[f11])).
fof(f11,axiom,(
  ! [X0,X1,X6] : (midp(X6,X1,X0) => midp(X6,X0,X1))),
  file('./6_GDD_FULL/21-40/26.p',ruleD11)).
fof(f430,plain,(
  ( ! [X4,X2,X0,X3,X1] : (~midp(X4,X2,X3) | para(X0,X2,X1,X3) | ~midp(X4,X0,X1)) )),
  inference(cnf_transformation,[],[f263])).
fof(f263,plain,(
  ! [X0,X1,X2,X3,X4] : (para(X0,X2,X1,X3) | ~midp(X4,X2,X3) | ~midp(X4,X0,X1))),
  inference(flattening,[],[f262])).
fof(f262,plain,(
  ! [X0,X1,X2,X3,X4] : (para(X0,X2,X1,X3) | (~midp(X4,X2,X3) | ~midp(X4,X0,X1)))),
  inference(ennf_transformation,[],[f142])).
fof(f142,plain,(
  ! [X0,X1,X2,X3,X4] : ((midp(X4,X2,X3) & midp(X4,X0,X1)) => para(X0,X2,X1,X3))),
  inference(rectify,[],[f64])).
fof(f64,axiom,(
  ! [X0,X1,X2,X3,X6] : ((midp(X6,X2,X3) & midp(X6,X0,X1)) => para(X0,X2,X1,X3))),
  file('./6_GDD_FULL/21-40/26.p',ruleD63)).
fof(f2186,plain,(
  ( ! [X2,X3,X1] : (~midp(sK29,X1,X2) | midp(X3,sK35,sK34) | ~para(X1,sK35,X2,sK34) | ~midp(X3,X1,X2)) )),
  inference(resolution,[],[f708,f431])).
fof(f708,plain,(
  ( ! [X4,X5] : (para(X4,sK34,X5,sK35) | ~midp(sK29,X4,X5)) )),
  inference(resolution,[],[f430,f495])).
fof(f60347,plain,(
  ~midp(sK40,sK35,sK34) | (~spl42_1 | ~spl42_314)),
  inference(subsumption_resolution,[],[f60343,f20223])).
fof(f20223,plain,(
  ( ! [X82] : (midp(X82,sK34,sK36)) ) | ~spl42_314),
  inference(subsumption_resolution,[],[f20162,f20071])).
fof(f20162,plain,(
  ( ! [X83,X82] : (midp(X82,sK34,sK36) | ~midp(X82,X83,X83)) ) | ~spl42_314),
  inference(resolution,[],[f20071,f2288])).
fof(f2288,plain,(
  ( ! [X2,X3,X1] : (~midp(sK23,X1,X2) | midp(X3,sK34,sK36) | ~midp(X3,X1,X2)) )),
  inference(subsumption_resolution,[],[f2283,f707])).
fof(f707,plain,(
  ( ! [X2,X3] : (para(X2,sK34,X3,sK36) | ~midp(sK23,X2,X3)) )),
  inference(resolution,[],[f430,f595])).
fof(f595,plain,(
  midp(sK23,sK34,sK36)),
  inference(resolution,[],[f377,f507])).
fof(f507,plain,(
  midp(sK23,sK36,sK34)),
  inference(cnf_transformation,[],[f366])).
fof(f2283,plain,(
  ( ! [X2,X3,X1] : (~midp(sK23,X1,X2) | midp(X3,sK34,sK36) | ~para(X1,sK34,X2,sK36) | ~midp(X3,X1,X2)) )),
  inference(resolution,[],[f718,f431])).
fof(f718,plain,(
  ( ! [X24,X25] : (para(X24,sK36,X25,sK34) | ~midp(sK23,X24,X25)) )),
  inference(resolution,[],[f430,f507])).
fof(f60343,plain,(
  ~midp(sK40,sK34,sK36) | ~midp(sK40,sK35,sK34) | (~spl42_1 | ~spl42_314)),
  inference(resolution,[],[f60333,f3131])).
fof(f3131,plain,(
  ( ! [X2,X0,X3,X1] : (circle(X0,X1,X2,X3) | ~midp(X0,X1,X3) | ~midp(X0,X2,X1)) )),
  inference(resolution,[],[f1022,f668])).
fof(f668,plain,(
  ( ! [X2,X0,X1] : (cong(X0,X1,X0,X2) | ~midp(X0,X2,X1)) )),
  inference(resolution,[],[f390,f435])).
fof(f435,plain,(
  ( ! [X2,X0,X1] : (cong(X0,X1,X0,X2) | ~midp(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f271])).
fof(f271,plain,(
  ! [X0,X1,X2] : (cong(X0,X1,X0,X2) | ~midp(X0,X1,X2))),
  inference(ennf_transformation,[],[f69])).
fof(f69,axiom,(
  ! [X0,X1,X2] : (midp(X0,X1,X2) => cong(X0,X1,X0,X2))),
  file('./6_GDD_FULL/21-40/26.p',ruleD68)).
fof(f390,plain,(
  ( ! [X2,X0,X3,X1] : (~cong(X0,X1,X2,X3) | cong(X2,X3,X0,X1)) )),
  inference(cnf_transformation,[],[f200])).
fof(f200,plain,(
  ! [X0,X1,X2,X3] : (cong(X2,X3,X0,X1) | ~cong(X0,X1,X2,X3))),
  inference(ennf_transformation,[],[f24])).
fof(f24,axiom,(
  ! [X0,X1,X2,X3] : (cong(X0,X1,X2,X3) => cong(X2,X3,X0,X1))),
  file('./6_GDD_FULL/21-40/26.p',ruleD24)).
fof(f1022,plain,(
  ( ! [X2,X0,X3,X1] : (~cong(X0,X1,X0,X2) | circle(X0,X1,X2,X3) | ~midp(X0,X1,X3)) )),
  inference(resolution,[],[f378,f435])).
fof(f378,plain,(
  ( ! [X2,X0,X3,X1] : (~cong(X3,X0,X3,X2) | circle(X3,X0,X1,X2) | ~cong(X3,X0,X3,X1)) )),
  inference(cnf_transformation,[],[f185])).
fof(f185,plain,(
  ! [X0,X1,X2,X3] : (circle(X3,X0,X1,X2) | ~cong(X3,X0,X3,X2) | ~cong(X3,X0,X3,X1))),
  inference(flattening,[],[f184])).
fof(f184,plain,(
  ! [X0,X1,X2,X3] : (circle(X3,X0,X1,X2) | (~cong(X3,X0,X3,X2) | ~cong(X3,X0,X3,X1)))),
  inference(ennf_transformation,[],[f98])).
fof(f98,plain,(
  ! [X0,X1,X2,X3] : ((cong(X3,X0,X3,X2) & cong(X3,X0,X3,X1)) => circle(X3,X0,X1,X2))),
  inference(rectify,[],[f12])).
fof(f12,axiom,(
  ! [X0,X1,X2,X7] : ((cong(X7,X0,X7,X2) & cong(X7,X0,X7,X1)) => circle(X7,X0,X1,X2))),
  file('./6_GDD_FULL/21-40/26.p',ruleD12)).
fof(f60333,plain,(
  ~circle(sK40,sK34,sK35,sK36) | (~spl42_1 | ~spl42_314)),
  inference(subsumption_resolution,[],[f60332,f20754])).
fof(f20754,plain,(
  ( ! [X0] : (midp(X0,sK35,sK36)) ) | ~spl42_314),
  inference(resolution,[],[f20226,f377])).
fof(f20226,plain,(
  ( ! [X91] : (midp(X91,sK36,sK35)) ) | ~spl42_314),
  inference(subsumption_resolution,[],[f20171,f20071])).
fof(f20171,plain,(
  ( ! [X92,X91] : (midp(X91,sK36,sK35) | ~midp(X91,X92,X92)) ) | ~spl42_314),
  inference(resolution,[],[f20071,f2243])).
fof(f2243,plain,(
  ( ! [X2,X3,X1] : (~midp(sK26,X1,X2) | midp(X3,sK36,sK35) | ~midp(X3,X1,X2)) )),
  inference(subsumption_resolution,[],[f2238,f721])).
fof(f721,plain,(
  ( ! [X30,X31] : (para(X30,sK36,X31,sK35) | ~midp(sK26,X30,X31)) )),
  inference(resolution,[],[f430,f593])).
fof(f593,plain,(
  midp(sK26,sK36,sK35)),
  inference(resolution,[],[f377,f501])).
fof(f501,plain,(
  midp(sK26,sK35,sK36)),
  inference(cnf_transformation,[],[f366])).
fof(f2238,plain,(
  ( ! [X2,X3,X1] : (~midp(sK26,X1,X2) | midp(X3,sK36,sK35) | ~para(X1,sK36,X2,sK35) | ~midp(X3,X1,X2)) )),
  inference(resolution,[],[f712,f431])).
fof(f712,plain,(
  ( ! [X12,X13] : (para(X12,sK35,X13,sK36) | ~midp(sK26,X12,X13)) )),
  inference(resolution,[],[f430,f501])).
fof(f60332,plain,(
  ~circle(sK40,sK34,sK35,sK36) | ~midp(sK39,sK35,sK36) | (~spl42_1 | ~spl42_314)),
  inference(subsumption_resolution,[],[f60281,f44746])).
fof(f44746,plain,(
  ( ! [X47,X48,X46] : (para(X46,X47,X48,X46)) ) | ~spl42_314),
  inference(resolution,[],[f44717,f370])).
fof(f370,plain,(
  ( ! [X2,X0,X3,X1] : (~para(X0,X1,X2,X3) | para(X0,X1,X3,X2)) )),
  inference(cnf_transformation,[],[f173])).
fof(f173,plain,(
  ! [X0,X1,X2,X3] : (para(X0,X1,X3,X2) | ~para(X0,X1,X2,X3))),
  inference(ennf_transformation,[],[f4])).
fof(f4,axiom,(
  ! [X0,X1,X2,X3] : (para(X0,X1,X2,X3) => para(X0,X1,X3,X2))),
  file('./6_GDD_FULL/21-40/26.p',ruleD4)).
fof(f60281,plain,(
  ~para(sK40,sK35,sK41,sK40) | ~circle(sK40,sK34,sK35,sK36) | ~midp(sK39,sK35,sK36) | ~spl42_1),
  inference(resolution,[],[f5992,f530])).
fof(f530,plain,(
  ~eqangle(sK35,sK34,sK34,sK36,sK41,sK40,sK40,sK39) | ~spl42_1),
  inference(avatar_component_clause,[],[f529])).
fof(f529,plain,(
  spl42_1 <=> ~eqangle(sK35,sK34,sK34,sK36,sK41,sK40,sK40,sK39)),
  introduced(avatar_definition,[new_symbols(naming,[spl42_1])])).
fof(f5992,plain,(
  ( ! [X191,X189,X194,X192,X190,X195,X193] : (eqangle(X189,X190,X190,X191,X192,X193,X194,X195) | ~para(X194,X189,X192,X193) | ~circle(X194,X190,X189,X191) | ~midp(X195,X189,X191)) )),
  inference(resolution,[],[f1434,f1129])).
fof(f1129,plain,(
  ( ! [X21,X19,X22,X20,X18] : (eqangle(X19,X22,X22,X20,X21,X19,X21,X18) | ~circle(X21,X22,X19,X20) | ~midp(X18,X19,X20)) )),
  inference(resolution,[],[f417,f384])).
fof(f417,plain,(
  ( ! [X4,X2,X0,X3,X1] : (eqangle(X0,X1,X0,X2,X3,X1,X3,X4) | ~midp(X4,X1,X2) | ~circle(X3,X0,X1,X2)) )),
  inference(cnf_transformation,[],[f240])).
fof(f240,plain,(
  ! [X0,X1,X2,X3,X4] : (eqangle(X0,X1,X0,X2,X3,X1,X3,X4) | ~midp(X4,X1,X2) | ~circle(X3,X0,X1,X2))),
  inference(flattening,[],[f239])).
fof(f239,plain,(
  ! [X0,X1,X2,X3,X4] : (eqangle(X0,X1,X0,X2,X3,X1,X3,X4) | (~midp(X4,X1,X2) | ~circle(X3,X0,X1,X2)))),
  inference(ennf_transformation,[],[f130])).
fof(f130,plain,(
  ! [X0,X1,X2,X3,X4] : ((midp(X4,X1,X2) & circle(X3,X0,X1,X2)) => eqangle(X0,X1,X0,X2,X3,X1,X3,X4))),
  inference(rectify,[],[f51])).
fof(f51,axiom,(
  ! [X0,X1,X2,X7,X6] : ((midp(X6,X1,X2) & circle(X7,X0,X1,X2)) => eqangle(X0,X1,X0,X2,X7,X1,X7,X6))),
  file('./6_GDD_FULL/21-40/26.p',ruleD50)).
fof(f1434,plain,(
  ( ! [X6,X4,X2,X0,X8,X7,X5,X3,X1,X9] : (~eqangle(X0,X1,X2,X3,X8,X9,X6,X7) | eqangle(X0,X1,X2,X3,X4,X5,X6,X7) | ~para(X8,X9,X4,X5)) )),
  inference(resolution,[],[f388,f406])).
fof(f5736,plain,(
  ~spl42_316),
  inference(avatar_contradiction_clause,[],[f5707])).
fof(f5707,plain,(
  $false | ~spl42_316),
  inference(resolution,[],[f3324,f2656])).
fof(f2656,plain,(
  para(sK22,sK41,sK22,sK41)),
  inference(resolution,[],[f1005,f656])).
fof(f656,plain,(
  perp(sK22,sK41,sK37,sK34)),
  inference(resolution,[],[f374,f510])).
fof(f510,plain,(
  perp(sK37,sK34,sK22,sK41)),
  inference(cnf_transformation,[],[f366])).
fof(f374,plain,(
  ( ! [X2,X0,X3,X1] : (~perp(X0,X1,X2,X3) | perp(X2,X3,X0,X1)) )),
  inference(cnf_transformation,[],[f178])).
fof(f178,plain,(
  ! [X0,X1,X2,X3] : (perp(X2,X3,X0,X1) | ~perp(X0,X1,X2,X3))),
  inference(ennf_transformation,[],[f8])).
fof(f8,axiom,(
  ! [X0,X1,X2,X3] : (perp(X0,X1,X2,X3) => perp(X2,X3,X0,X1))),
  file('./6_GDD_FULL/21-40/26.p',ruleD8)).
fof(f1005,plain,(
  ( ! [X35,X34] : (~perp(X34,X35,sK37,sK34) | para(X34,X35,sK22,sK41)) )),
  inference(resolution,[],[f375,f510])).
fof(f375,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (~perp(X2,X3,X4,X5) | para(X0,X1,X4,X5) | ~perp(X0,X1,X2,X3)) )),
  inference(cnf_transformation,[],[f180])).
fof(f180,plain,(
  ! [X0,X1,X2,X3,X4,X5] : (para(X0,X1,X4,X5) | ~perp(X2,X3,X4,X5) | ~perp(X0,X1,X2,X3))),
  inference(flattening,[],[f179])).
fof(f179,plain,(
  ! [X0,X1,X2,X3,X4,X5] : (para(X0,X1,X4,X5) | (~perp(X2,X3,X4,X5) | ~perp(X0,X1,X2,X3)))),
  inference(ennf_transformation,[],[f9])).
fof(f9,axiom,(
  ! [X0,X1,X2,X3,X4,X5] : ((perp(X2,X3,X4,X5) & perp(X0,X1,X2,X3)) => para(X0,X1,X4,X5))),
  file('./6_GDD_FULL/21-40/26.p',ruleD9)).
fof(f3324,plain,(
  ( ! [X0,X1] : (~para(X0,X1,X0,X1)) ) | ~spl42_316),
  inference(avatar_component_clause,[],[f3323])).
fof(f3323,plain,(
  spl42_316 <=> ! [X1,X0] : ~para(X0,X1,X0,X1)),
  introduced(avatar_definition,[new_symbols(naming,[spl42_316])])).
fof(f3325,plain,(
  spl42_314 | spl42_316),
  inference(avatar_split_clause,[],[f3292,f3323,f3320])).
fof(f3292,plain,(
  ( ! [X2,X0,X3,X1] : (~para(X0,X1,X0,X1) | para(X2,X3,X2,X3)) )),
  inference(resolution,[],[f1062,f405])).
fof(f405,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (~eqangle(X0,X1,X4,X5,X2,X3,X4,X5) | para(X0,X1,X2,X3)) )),
  inference(cnf_transformation,[],[f219])).
fof(f219,plain,(
  ! [X0,X1,X2,X3,X4,X5] : (para(X0,X1,X2,X3) | ~eqangle(X0,X1,X4,X5,X2,X3,X4,X5))),
  inference(ennf_transformation,[],[f118])).
fof(f118,plain,(
  ! [X0,X1,X2,X3,X4,X5] : (eqangle(X0,X1,X4,X5,X2,X3,X4,X5) => para(X0,X1,X2,X3))),
  inference(rectify,[],[f39])).
fof(f39,axiom,(
  ! [X0,X1,X2,X3,X8,X9] : (eqangle(X0,X1,X8,X9,X2,X3,X8,X9) => para(X0,X1,X2,X3))),
  file('./6_GDD_FULL/21-40/26.p',ruleD39)).
fof(f586,plain,(
  ~spl42_1 | ~spl42_13 | ~spl42_11 | ~spl42_7),
  inference(avatar_split_clause,[],[f585,f547,f561,f570,f529])).
fof(f585,plain,(
  ~eqangle(sK34,sK35,sK35,sK36,sK40,sK41,sK41,sK39) | ~eqangle(sK34,sK35,sK35,sK36,sK41,sK40,sK40,sK39) | ~eqangle(sK34,sK35,sK35,sK36,sK40,sK39,sK39,sK41) | ~eqangle(sK35,sK34,sK34,sK36,sK41,sK40,sK40,sK39)),
  inference(subsumption_resolution,[],[f584,f483])).
fof(f483,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (~eqangle(X5,X4,X4,X3,X2,X1,X1,X0) | ~sP0(X0,X1,X2,X3,X4,X5)) )),
  inference(cnf_transformation,[],[f364])).
fof(f364,plain,(
  ! [X0,X1,X2,X3,X4,X5] : ((~eqangle(X5,X4,X4,X3,X2,X1,X1,X0) & ~eqangle(X4,X5,X5,X3,X1,X2,X2,X0)) | ~sP0(X0,X1,X2,X3,X4,X5))),
  inference(rectify,[],[f363])).
fof(f363,plain,(
  ! [X17,X18,X19,X14,X12,X13] : ((~eqangle(X13,X12,X12,X14,X19,X18,X18,X17) & ~eqangle(X12,X13,X13,X14,X18,X19,X19,X17)) | ~sP0(X17,X18,X19,X14,X12,X13))),
  inference(nnf_transformation,[],[f322])).
fof(f584,plain,(
  ~eqangle(sK34,sK35,sK35,sK36,sK40,sK41,sK41,sK39) | ~eqangle(sK34,sK35,sK35,sK36,sK41,sK40,sK40,sK39) | ~eqangle(sK34,sK35,sK35,sK36,sK40,sK39,sK39,sK41) | sP0(sK39,sK40,sK41,sK36,sK34,sK35) | ~eqangle(sK35,sK34,sK34,sK36,sK41,sK40,sK40,sK39)),
  inference(subsumption_resolution,[],[f519,f480])).
fof(f480,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (~eqangle(X4,X5,X5,X3,X2,X0,X0,X1) | ~sP1(X0,X1,X2,X3,X4,X5)) )),
  inference(cnf_transformation,[],[f362])).
fof(f362,plain,(
  ! [X0,X1,X2,X3,X4,X5] : ((~eqangle(X5,X4,X4,X3,X2,X1,X1,X0) & ~eqangle(X4,X5,X5,X3,X2,X0,X0,X1)) | ~sP1(X0,X1,X2,X3,X4,X5))),
  inference(rectify,[],[f361])).
fof(f361,plain,(
  ! [X17,X19,X18,X14,X12,X13] : ((~eqangle(X13,X12,X12,X14,X18,X19,X19,X17) & ~eqangle(X12,X13,X13,X14,X18,X17,X17,X19)) | ~sP1(X17,X19,X18,X14,X12,X13))),
  inference(nnf_transformation,[],[f323])).
fof(f519,plain,(
  ~eqangle(sK34,sK35,sK35,sK36,sK40,sK41,sK41,sK39) | ~eqangle(sK34,sK35,sK35,sK36,sK41,sK40,sK40,sK39) | ~eqangle(sK34,sK35,sK35,sK36,sK40,sK39,sK39,sK41) | sP1(sK39,sK41,sK40,sK36,sK34,sK35) | sP0(sK39,sK40,sK41,sK36,sK34,sK35) | ~eqangle(sK35,sK34,sK34,sK36,sK41,sK40,sK40,sK39)),
  inference(duplicate_literal_removal,[],[f511])).
fof(f511,plain,(
  ~eqangle(sK34,sK35,sK35,sK36,sK40,sK41,sK41,sK39) | ~eqangle(sK34,sK35,sK35,sK36,sK41,sK40,sK40,sK39) | ~eqangle(sK34,sK35,sK35,sK36,sK40,sK39,sK39,sK41) | ~eqangle(sK34,sK35,sK35,sK36,sK41,sK40,sK40,sK39) | sP1(sK39,sK41,sK40,sK36,sK34,sK35) | sP0(sK39,sK40,sK41,sK36,sK34,sK35) | ~eqangle(sK35,sK34,sK34,sK36,sK41,sK40,sK40,sK39)),
  inference(cnf_transformation,[],[f366])).
% SZS output end Proof for 26
% ------------------------------
% Version: Vampire 4.4.0 (commit 7916d27 on 2019-08-23 08:50:16 +0100)
% Termination reason: Refutation

% Memory used [KB]: 37483
% Time elapsed: 1.254 s
% ------------------------------
% ------------------------------
