% Refutation found. Thanks to Tanya!
% SZS status Theorem for 37
% SZS output start Proof for 37
9. ! [X0,X1,X2,X3,X4,X5] : ((perp(X2,X3,X4,X5) & perp(X0,X1,X2,X3)) => para(X0,X1,X4,X5)) [input]
18. ! [X0,X1,X2,X3,X8,X9,X10,X11] : (eqangle(X0,X1,X2,X3,X8,X9,X10,X11) => eqangle(X1,X0,X2,X3,X8,X9,X10,X11)) [input]
19. ! [X0,X1,X2,X3,X8,X9,X10,X11] : (eqangle(X0,X1,X2,X3,X8,X9,X10,X11) => eqangle(X2,X3,X0,X1,X10,X11,X8,X9)) [input]
23. ! [X0,X1,X2,X3] : (cong(X0,X1,X2,X3) => cong(X0,X1,X3,X2)) [input]
25. ! [X0,X1,X2,X3,X4,X5] : ((cong(X2,X3,X4,X5) & cong(X0,X1,X2,X3)) => cong(X0,X1,X4,X5)) [input]
40. ! [X0,X1,X2,X3,X8,X9] : (para(X0,X1,X2,X3) => eqangle(X0,X1,X8,X9,X2,X3,X8,X9)) [input]
51. ! [X0,X1,X2,X7,X6] : ((eqangle(X0,X1,X0,X2,X7,X1,X7,X6) & coll(X6,X1,X2) & circle(X7,X0,X1,X2)) => midp(X6,X1,X2)) [input]
68. ! [X0,X1,X2] : (midp(X0,X1,X2) => cong(X0,X1,X0,X2)) [input]
94. ! [X0,X1,X2,X7,X13,X18,X19,X20,X21] : (coll(X19,X0,X13) & circle(X7,X0,X19,X21) & coll(X18,X2,X13) & circle(X7,X2,X18,X20) & perp(X1,X2,X0,X13) & perp(X0,X2,X1,X13) & perp(X0,X1,X2,X13) & circle(X7,X0,X1,X2)) => cong(X1,X19,X1,X18) [input]
95. ~! [X18,X1,X19] : (! [X0,X1,X2,X7,X13,X18,X19,X20,X21] : (coll(X19,X0,X13) & circle(X7,X0,X19,X21) & coll(X18,X2,X13) & circle(X7,X2,X18,X20) & perp(X1,X2,X0,X13) & perp(X0,X2,X1,X13) & perp(X0,X1,X2,X13) & circle(X7,X0,X1,X2)) => cong(X1,X19,X1,X18)) [negated conjecture 94]
99. ! [X0,X1,X2,X3,X4,X5,X6,X7] : (eqangle(X0,X1,X2,X3,X4,X5,X6,X7) => eqangle(X1,X0,X2,X3,X4,X5,X6,X7)) [rectify 18]
100. ! [X0,X1,X2,X3,X4,X5,X6,X7] : (eqangle(X0,X1,X2,X3,X4,X5,X6,X7) => eqangle(X2,X3,X0,X1,X6,X7,X4,X5)) [rectify 19]
118. ! [X0,X1,X2,X3,X4,X5] : (para(X0,X1,X2,X3) => eqangle(X0,X1,X4,X5,X2,X3,X4,X5)) [rectify 40]
129. ! [X0,X1,X2,X3,X4] : ((eqangle(X0,X1,X0,X2,X3,X1,X3,X4) & coll(X4,X1,X2) & circle(X3,X0,X1,X2)) => midp(X4,X1,X2)) [rectify 51]
165. ~! [X0,X1,X2] : (! [X3,X4,X5,X6,X7,X8,X9,X10,X11] : (coll(X9,X3,X7) & circle(X6,X3,X9,X11) & coll(X8,X5,X7) & circle(X6,X5,X8,X10) & perp(X4,X5,X3,X7) & perp(X3,X5,X4,X7) & perp(X3,X4,X5,X7) & circle(X6,X3,X4,X5)) => cong(X1,X2,X1,X0)) [rectify 95]
176. ! [X0,X1,X2,X3,X4,X5] : (para(X0,X1,X4,X5) | (~perp(X2,X3,X4,X5) | ~perp(X0,X1,X2,X3))) [ennf transformation 9]
177. ! [X0,X1,X2,X3,X4,X5] : (para(X0,X1,X4,X5) | ~perp(X2,X3,X4,X5) | ~perp(X0,X1,X2,X3)) [flattening 176]
190. ! [X0,X1,X2,X3,X4,X5,X6,X7] : (eqangle(X1,X0,X2,X3,X4,X5,X6,X7) | ~eqangle(X0,X1,X2,X3,X4,X5,X6,X7)) [ennf transformation 99]
191. ! [X0,X1,X2,X3,X4,X5,X6,X7] : (eqangle(X2,X3,X0,X1,X6,X7,X4,X5) | ~eqangle(X0,X1,X2,X3,X4,X5,X6,X7)) [ennf transformation 100]
196. ! [X0,X1,X2,X3] : (cong(X0,X1,X3,X2) | ~cong(X0,X1,X2,X3)) [ennf transformation 23]
198. ! [X0,X1,X2,X3,X4,X5] : (cong(X0,X1,X4,X5) | (~cong(X2,X3,X4,X5) | ~cong(X0,X1,X2,X3))) [ennf transformation 25]
199. ! [X0,X1,X2,X3,X4,X5] : (cong(X0,X1,X4,X5) | ~cong(X2,X3,X4,X5) | ~cong(X0,X1,X2,X3)) [flattening 198]
217. ! [X0,X1,X2,X3,X4,X5] : (eqangle(X0,X1,X4,X5,X2,X3,X4,X5) | ~para(X0,X1,X2,X3)) [ennf transformation 118]
236. ! [X0,X1,X2,X3,X4] : (midp(X4,X1,X2) | (~eqangle(X0,X1,X0,X2,X3,X1,X3,X4) | ~coll(X4,X1,X2) | ~circle(X3,X0,X1,X2))) [ennf transformation 129]
237. ! [X0,X1,X2,X3,X4] : (midp(X4,X1,X2) | ~eqangle(X0,X1,X0,X2,X3,X1,X3,X4) | ~coll(X4,X1,X2) | ~circle(X3,X0,X1,X2)) [flattening 236]
266. ! [X0,X1,X2] : (cong(X0,X1,X0,X2) | ~midp(X0,X1,X2)) [ennf transformation 68]
315. ? [X0,X1,X2] : (~cong(X1,X2,X1,X0) & ! [X3,X4,X5,X6,X7,X8,X9,X10,X11] : (coll(X9,X3,X7) & circle(X6,X3,X9,X11) & coll(X8,X5,X7) & circle(X6,X5,X8,X10) & perp(X4,X5,X3,X7) & perp(X3,X5,X4,X7) & perp(X3,X4,X5,X7) & circle(X6,X3,X4,X5))) [ennf transformation 165]
352. ? [X0,X1,X2] : (~cong(X1,X2,X1,X0) & ! [X3,X4,X5,X6,X7,X8,X9,X10,X11] : (coll(X9,X3,X7) & circle(X6,X3,X9,X11) & coll(X8,X5,X7) & circle(X6,X5,X8,X10) & perp(X4,X5,X3,X7) & perp(X3,X5,X4,X7) & perp(X3,X4,X5,X7) & circle(X6,X3,X4,X5))) => (~cong(sK21,sK22,sK21,sK20) & ! [X3,X4,X5,X6,X7,X8,X9,X10,X11] : (coll(X9,X3,X7) & circle(X6,X3,X9,X11) & coll(X8,X5,X7) & circle(X6,X5,X8,X10) & perp(X4,X5,X3,X7) & perp(X3,X5,X4,X7) & perp(X3,X4,X5,X7) & circle(X6,X3,X4,X5))) [choice axiom]
353. ~cong(sK21,sK22,sK21,sK20) & ! [X3,X4,X5,X6,X7,X8,X9,X10,X11] : (coll(X9,X3,X7) & circle(X6,X3,X9,X11) & coll(X8,X5,X7) & circle(X6,X5,X8,X10) & perp(X4,X5,X3,X7) & perp(X3,X5,X4,X7) & perp(X3,X4,X5,X7) & circle(X6,X3,X4,X5)) [skolemisation 315,352]
362. ~perp(X2,X3,X4,X5) | para(X0,X1,X4,X5) | ~perp(X0,X1,X2,X3) [cnf transformation 177]
371. ~eqangle(X0,X1,X2,X3,X4,X5,X6,X7) | eqangle(X1,X0,X2,X3,X4,X5,X6,X7) [cnf transformation 190]
372. ~eqangle(X0,X1,X2,X3,X4,X5,X6,X7) | eqangle(X2,X3,X0,X1,X6,X7,X4,X5) [cnf transformation 191]
376. ~cong(X0,X1,X2,X3) | cong(X0,X1,X3,X2) [cnf transformation 196]
378. ~cong(X2,X3,X4,X5) | cong(X0,X1,X4,X5) | ~cong(X0,X1,X2,X3) [cnf transformation 199]
393. eqangle(X0,X1,X4,X5,X2,X3,X4,X5) | ~para(X0,X1,X2,X3) [cnf transformation 217]
405. ~eqangle(X0,X1,X0,X2,X3,X1,X3,X4) | midp(X4,X1,X2) | ~coll(X4,X1,X2) | ~circle(X3,X0,X1,X2) [cnf transformation 237]
422. cong(X0,X1,X0,X2) | ~midp(X0,X1,X2) [cnf transformation 266]
470. perp(X4,X5,X3,X7) [cnf transformation 353]
473. circle(X6,X3,X9,X11) [cnf transformation 353]
474. coll(X9,X3,X7) [cnf transformation 353]
475. ~cong(sK21,sK22,sK21,sK20) [cnf transformation 353]
481. cong(X0,X1,X2,X0) | ~midp(X0,X1,X2) [resolution 376,422]
534. para(X0,X1,X2,X3) | ~perp(X0,X1,X4,X5) [resolution 362,470]
536. para(X0,X1,X2,X3) [subsumption resolution 534,470]
564. ~cong(X5,X6,X7,X9) | cong(X5,X6,X7,X8) | ~midp(X7,X9,X8) [resolution 378,422]
586. cong(X12,X13,X14,X15) | ~midp(X14,X12,X15) | ~midp(X12,X13,X14) [resolution 564,481]
607. eqangle(X0,X1,X2,X3,X4,X5,X2,X3) | ~para(X1,X0,X4,X5) [resolution 371,393]
611. eqangle(X0,X1,X2,X3,X4,X5,X2,X3) [subsumption resolution 607,536]
614. eqangle(X0,X1,X2,X3,X0,X1,X4,X5) [resolution 372,611]
621. eqangle(X8,X9,X10,X11,X9,X8,X12,X13) [resolution 614,371]
630. eqangle(X8,X9,X10,X11,X12,X13,X11,X10) [resolution 621,372]
708. ~midp(sK21,sK21,sK20) | ~midp(sK21,sK22,sK21) [resolution 586,475]
721. 1 <=> ~midp(sK21,sK22,sK21) [avatar definition]
722. ~midp(sK21,sK22,sK21) <- (1) [avatar component clause 721]
727. 3 <=> ~midp(sK21,sK21,sK20) [avatar definition]
728. ~midp(sK21,sK21,sK20) <- (3) [avatar component clause 727]
729. ~1 | ~3 [avatar split clause 708,727,721]
1240. midp(X0,X1,X0) | ~coll(X0,X1,X0) | ~circle(X2,X2,X1,X0) [resolution 405,611]
1242. midp(X6,X7,X8) | ~coll(X6,X7,X8) | ~circle(X8,X6,X7,X8) [resolution 405,630]
1271. midp(X0,X1,X0) | ~circle(X2,X2,X1,X0) [subsumption resolution 1240,474]
1272. midp(X0,X1,X0) [subsumption resolution 1271,473]
1273. midp(X6,X7,X8) | ~circle(X8,X6,X7,X8) [subsumption resolution 1242,474]
1274. midp(X6,X7,X8) [subsumption resolution 1273,473]
1296. $false <- (1) [resolution 1272,722]
1297. 1 [avatar contradiction clause 1296]
1298. $false <- (3) [subsumption resolution 728,1274]
1299. 3 [avatar contradiction clause 1298]
1300. $false [avatar sat refutation 729,1297,1299]
% SZS output end Proof for 37
% ------------------------------
% Version: Vampire 4.4.0 (commit 7916d27 on 2019-08-23 08:50:16 +0100)
% Termination reason: Refutation

% Memory used [KB]: 5373
% Time elapsed: 0.036 s
% ------------------------------
% ------------------------------
