% Refutation found. Thanks to Tanya!
% SZS status Theorem for E022-8
% SZS output start Proof for E022-8
fof(f52426,plain,(
  $false),
  inference(avatar_sat_refutation,[],[f695,f773,f1027,f2306,f2482,f2807,f4117,f5159,f9215,f9655,f11055,f16003,f16025,f16186,f17122,f52417])).
fof(f52417,plain,(
  ~spl37_30 | ~spl37_54 | ~spl37_60 | ~spl37_232 | ~spl37_238 | ~spl37_244 | ~spl37_650 | ~spl37_742),
  inference(avatar_contradiction_clause,[],[f52394])).
fof(f52394,plain,(
  $false | (~spl37_30 | ~spl37_54 | ~spl37_60 | ~spl37_232 | ~spl37_238 | ~spl37_244 | ~spl37_650 | ~spl37_742)),
  inference(resolution,[],[f35202,f482])).
fof(f482,plain,(
  ~para(sK27,sK24,sK26,sK25)),
  inference(cnf_transformation,[],[f358])).
fof(f358,plain,(
  ~para(sK27,sK24,sK26,sK25) & perp(sK22,sK24,sK24,sK27) & circle(sK20,sK21,sK26,sK36) & coll(sK21,sK24,sK26) & circle(sK20,sK23,sK25,sK35) & coll(sK23,sK24,sK25) & circle(sK22,sK21,sK24,sK34) & circle(sK22,sK21,sK23,sK33) & circle(sK20,sK21,sK23,sK32) & circle(sK22,sK21,sK30,sK31) & circle(sK20,sK21,sK28,sK29)),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK20,sK21,sK22,sK23,sK24,sK25,sK26,sK27,sK28,sK29,sK30,sK31,sK32,sK33,sK34,sK35,sK36])],[f320,f357])).
fof(f357,plain,(
  ? [X0,X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16] : (~para(X7,X4,X6,X5) & perp(X2,X4,X4,X7) & circle(X0,X1,X6,X16) & coll(X1,X4,X6) & circle(X0,X3,X5,X15) & coll(X3,X4,X5) & circle(X2,X1,X4,X14) & circle(X2,X1,X3,X13) & circle(X0,X1,X3,X12) & circle(X2,X1,X10,X11) & circle(X0,X1,X8,X9)) => (~para(sK27,sK24,sK26,sK25) & perp(sK22,sK24,sK24,sK27) & circle(sK20,sK21,sK26,sK36) & coll(sK21,sK24,sK26) & circle(sK20,sK23,sK25,sK35) & coll(sK23,sK24,sK25) & circle(sK22,sK21,sK24,sK34) & circle(sK22,sK21,sK23,sK33) & circle(sK20,sK21,sK23,sK32) & circle(sK22,sK21,sK30,sK31) & circle(sK20,sK21,sK28,sK29))),
  introduced(choice_axiom,[])).
fof(f320,plain,(
  ? [X0,X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16] : (~para(X7,X4,X6,X5) & perp(X2,X4,X4,X7) & circle(X0,X1,X6,X16) & coll(X1,X4,X6) & circle(X0,X3,X5,X15) & coll(X3,X4,X5) & circle(X2,X1,X4,X14) & circle(X2,X1,X3,X13) & circle(X0,X1,X3,X12) & circle(X2,X1,X10,X11) & circle(X0,X1,X8,X9))),
  inference(flattening,[],[f319])).
fof(f319,plain,(
  ? [X0,X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16] : (~para(X7,X4,X6,X5) & (perp(X2,X4,X4,X7) & circle(X0,X1,X6,X16) & coll(X1,X4,X6) & circle(X0,X3,X5,X15) & coll(X3,X4,X5) & circle(X2,X1,X4,X14) & circle(X2,X1,X3,X13) & circle(X0,X1,X3,X12) & circle(X2,X1,X10,X11) & circle(X0,X1,X8,X9)))),
  inference(ennf_transformation,[],[f167])).
fof(f167,plain,(
  ~! [X0,X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16] : ((perp(X2,X4,X4,X7) & circle(X0,X1,X6,X16) & coll(X1,X4,X6) & circle(X0,X3,X5,X15) & coll(X3,X4,X5) & circle(X2,X1,X4,X14) & circle(X2,X1,X3,X13) & circle(X0,X1,X3,X12) & circle(X2,X1,X10,X11) & circle(X0,X1,X8,X9)) => para(X7,X4,X6,X5))),
  inference(rectify,[],[f96])).
fof(f96,negated_conjecture,(
  ~! [X0,X1,X2,X3,X4,X5,X12,X13,X18,X19,X20,X21,X22,X23,X24,X25,X26] : ((perp(X2,X4,X4,X13) & circle(X0,X1,X12,X26) & coll(X1,X4,X12) & circle(X0,X3,X5,X25) & coll(X3,X4,X5) & circle(X2,X1,X4,X24) & circle(X2,X1,X3,X23) & circle(X0,X1,X3,X22) & circle(X2,X1,X20,X21) & circle(X0,X1,X18,X19)) => para(X13,X4,X12,X5))),
  inference(negated_conjecture,[],[f95])).
fof(f95,conjecture,(
  ! [X0,X1,X2,X3,X4,X5,X12,X13,X18,X19,X20,X21,X22,X23,X24,X25,X26] : ((perp(X2,X4,X4,X13) & circle(X0,X1,X12,X26) & coll(X1,X4,X12) & circle(X0,X3,X5,X25) & coll(X3,X4,X5) & circle(X2,X1,X4,X24) & circle(X2,X1,X3,X23) & circle(X0,X1,X3,X22) & circle(X2,X1,X20,X21) & circle(X0,X1,X18,X19)) => para(X13,X4,X12,X5))),
  file('./7_Book/00EE/01/E022-8.p',exemplo7Book00EE01E0228)).
fof(f35202,plain,(
  ( ! [X35,X33,X34,X32] : (para(X32,X33,X34,X35)) ) | (~spl37_30 | ~spl37_54 | ~spl37_60 | ~spl37_232 | ~spl37_238 | ~spl37_244 | ~spl37_650 | ~spl37_742)),
  inference(subsumption_resolution,[],[f35181,f35180])).
fof(f35180,plain,(
  ( ! [X30,X31] : (perp(X30,X31,sK22,sK24)) ) | (~spl37_30 | ~spl37_54 | ~spl37_60 | ~spl37_232 | ~spl37_238 | ~spl37_244 | ~spl37_650 | ~spl37_742)),
  inference(resolution,[],[f21475,f366])).
fof(f366,plain,(
  ( ! [X2,X0,X3,X1] : (~perp(X0,X1,X2,X3) | perp(X2,X3,X0,X1)) )),
  inference(cnf_transformation,[],[f177])).
fof(f177,plain,(
  ! [X0,X1,X2,X3] : (perp(X2,X3,X0,X1) | ~perp(X0,X1,X2,X3))),
  inference(ennf_transformation,[],[f8])).
fof(f8,axiom,(
  ! [X0,X1,X2,X3] : (perp(X0,X1,X2,X3) => perp(X2,X3,X0,X1))),
  file('./7_Book/00EE/01/E022-8.p',ruleD8)).
fof(f21475,plain,(
  ( ! [X10,X11] : (perp(sK22,sK24,X10,X11)) ) | (~spl37_30 | ~spl37_54 | ~spl37_60 | ~spl37_232 | ~spl37_238 | ~spl37_244 | ~spl37_650 | ~spl37_742)),
  inference(subsumption_resolution,[],[f21474,f21426])).
fof(f21426,plain,(
  ( ! [X3] : (midp(sK22,X3,sK24)) ) | (~spl37_30 | ~spl37_54 | ~spl37_232 | ~spl37_238 | ~spl37_244 | ~spl37_650 | ~spl37_742)),
  inference(subsumption_resolution,[],[f16380,f21415])).
fof(f21415,plain,(
  ( ! [X2] : (midp(sK22,X2,sK22)) ) | (~spl37_30 | ~spl37_54 | ~spl37_232 | ~spl37_238 | ~spl37_244 | ~spl37_650 | ~spl37_742)),
  inference(subsumption_resolution,[],[f16990,f18097])).
fof(f18097,plain,(
  ( ! [X2] : (midp(sK24,X2,sK22)) ) | (~spl37_30 | ~spl37_54 | ~spl37_232 | ~spl37_238 | ~spl37_244 | ~spl37_650 | ~spl37_742)),
  inference(subsumption_resolution,[],[f18096,f17753])).
fof(f17753,plain,(
  ( ! [X12] : (midp(sK22,X12,sK27)) ) | (~spl37_30 | ~spl37_232 | ~spl37_244 | ~spl37_650 | ~spl37_742)),
  inference(subsumption_resolution,[],[f17715,f6271])).
fof(f6271,plain,(
  ( ! [X2,X0,X1] : (coll(X0,X1,X2)) ) | ~spl37_232),
  inference(subsumption_resolution,[],[f6263,f6223])).
fof(f6223,plain,(
  ( ! [X6,X5] : (coll(X5,X5,X6)) ) | ~spl37_232),
  inference(resolution,[],[f6189,f359])).
fof(f359,plain,(
  ( ! [X2,X0,X1] : (~coll(X0,X1,X2) | coll(X0,X2,X1)) )),
  inference(cnf_transformation,[],[f168])).
fof(f168,plain,(
  ! [X0,X1,X2] : (coll(X0,X2,X1) | ~coll(X0,X1,X2))),
  inference(ennf_transformation,[],[f1])).
fof(f1,axiom,(
  ! [X0,X1,X2] : (coll(X0,X1,X2) => coll(X0,X2,X1))),
  file('./7_Book/00EE/01/E022-8.p',ruleD1)).
fof(f6189,plain,(
  ( ! [X4,X3] : (coll(X3,X4,X3)) ) | ~spl37_232),
  inference(resolution,[],[f6138,f360])).
fof(f360,plain,(
  ( ! [X2,X0,X1] : (~coll(X0,X1,X2) | coll(X1,X0,X2)) )),
  inference(cnf_transformation,[],[f169])).
fof(f169,plain,(
  ! [X0,X1,X2] : (coll(X1,X0,X2) | ~coll(X0,X1,X2))),
  inference(ennf_transformation,[],[f2])).
fof(f2,axiom,(
  ! [X0,X1,X2] : (coll(X0,X1,X2) => coll(X1,X0,X2))),
  file('./7_Book/00EE/01/E022-8.p',ruleD2)).
fof(f6138,plain,(
  ( ! [X4,X3] : (coll(X3,X4,X4)) ) | ~spl37_232),
  inference(resolution,[],[f2302,f425])).
fof(f425,plain,(
  ( ! [X2,X0,X1] : (~para(X0,X1,X0,X2) | coll(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f267])).
fof(f267,plain,(
  ! [X0,X1,X2] : (coll(X0,X1,X2) | ~para(X0,X1,X0,X2))),
  inference(ennf_transformation,[],[f67])).
fof(f67,axiom,(
  ! [X0,X1,X2] : (para(X0,X1,X0,X2) => coll(X0,X1,X2))),
  file('./7_Book/00EE/01/E022-8.p',ruleD66)).
fof(f2302,plain,(
  ( ! [X2,X3] : (para(X2,X3,X2,X3)) ) | ~spl37_232),
  inference(avatar_component_clause,[],[f2301])).
fof(f2301,plain,(
  spl37_232 <=> ! [X3,X2] : para(X2,X3,X2,X3)),
  introduced(avatar_definition,[new_symbols(naming,[spl37_232])])).
fof(f6263,plain,(
  ( ! [X2,X0,X1] : (coll(X0,X1,X2) | ~coll(X2,X2,X0)) ) | ~spl37_232),
  inference(resolution,[],[f6223,f361])).
fof(f361,plain,(
  ( ! [X2,X0,X3,X1] : (~coll(X0,X1,X3) | coll(X2,X3,X0) | ~coll(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f171])).
fof(f171,plain,(
  ! [X0,X1,X2,X3] : (coll(X2,X3,X0) | ~coll(X0,X1,X3) | ~coll(X0,X1,X2))),
  inference(flattening,[],[f170])).
fof(f170,plain,(
  ! [X0,X1,X2,X3] : (coll(X2,X3,X0) | (~coll(X0,X1,X3) | ~coll(X0,X1,X2)))),
  inference(ennf_transformation,[],[f3])).
fof(f3,axiom,(
  ! [X0,X1,X2,X3] : ((coll(X0,X1,X3) & coll(X0,X1,X2)) => coll(X2,X3,X0))),
  file('./7_Book/00EE/01/E022-8.p',ruleD3)).
fof(f17715,plain,(
  ( ! [X12] : (~coll(sK22,X12,sK27) | midp(sK22,X12,sK27)) ) | (~spl37_30 | ~spl37_232 | ~spl37_244 | ~spl37_650 | ~spl37_742)),
  inference(resolution,[],[f17367,f426])).
fof(f426,plain,(
  ( ! [X2,X0,X1] : (~cong(X0,X1,X0,X2) | ~coll(X0,X1,X2) | midp(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f269])).
fof(f269,plain,(
  ! [X0,X1,X2] : (midp(X0,X1,X2) | ~coll(X0,X1,X2) | ~cong(X0,X1,X0,X2))),
  inference(flattening,[],[f268])).
fof(f268,plain,(
  ! [X0,X1,X2] : (midp(X0,X1,X2) | (~coll(X0,X1,X2) | ~cong(X0,X1,X0,X2)))),
  inference(ennf_transformation,[],[f68])).
fof(f68,axiom,(
  ! [X0,X1,X2] : ((coll(X0,X1,X2) & cong(X0,X1,X0,X2)) => midp(X0,X1,X2))),
  file('./7_Book/00EE/01/E022-8.p',ruleD67)).
fof(f17367,plain,(
  ( ! [X0] : (cong(sK22,X0,sK22,sK27)) ) | (~spl37_30 | ~spl37_232 | ~spl37_244 | ~spl37_650 | ~spl37_742)),
  inference(subsumption_resolution,[],[f17351,f16155])).
fof(f16155,plain,(
  ( ! [X4] : (midp(X4,sK22,sK24)) ) | (~spl37_232 | ~spl37_650 | ~spl37_742)),
  inference(subsumption_resolution,[],[f16154,f10254])).
fof(f10254,plain,(
  ( ! [X30,X29] : (midp(X29,X30,X30)) ) | ~spl37_232),
  inference(resolution,[],[f8846,f6152])).
fof(f6152,plain,(
  ( ! [X12,X13,X11] : (~midp(X11,X13,X13) | midp(X11,X12,X12)) ) | ~spl37_232),
  inference(subsumption_resolution,[],[f6141,f2302])).
fof(f6141,plain,(
  ( ! [X12,X13,X11] : (midp(X11,X12,X12) | ~para(X13,X12,X13,X12) | ~midp(X11,X13,X13)) ) | ~spl37_232),
  inference(resolution,[],[f2302,f423])).
fof(f423,plain,(
  ( ! [X4,X2,X0,X3,X1] : (~para(X0,X3,X1,X2) | midp(X4,X2,X3) | ~para(X0,X2,X1,X3) | ~midp(X4,X0,X1)) )),
  inference(cnf_transformation,[],[f264])).
fof(f264,plain,(
  ! [X0,X1,X2,X3,X4] : (midp(X4,X2,X3) | ~para(X0,X3,X1,X2) | ~para(X0,X2,X1,X3) | ~midp(X4,X0,X1))),
  inference(flattening,[],[f263])).
fof(f263,plain,(
  ! [X0,X1,X2,X3,X4] : (midp(X4,X2,X3) | (~para(X0,X3,X1,X2) | ~para(X0,X2,X1,X3) | ~midp(X4,X0,X1)))),
  inference(ennf_transformation,[],[f143])).
fof(f143,plain,(
  ! [X0,X1,X2,X3,X4] : ((para(X0,X3,X1,X2) & para(X0,X2,X1,X3) & midp(X4,X0,X1)) => midp(X4,X2,X3))),
  inference(rectify,[],[f65])).
fof(f65,axiom,(
  ! [X0,X1,X2,X3,X6] : ((para(X0,X3,X1,X2) & para(X0,X2,X1,X3) & midp(X6,X0,X1)) => midp(X6,X2,X3))),
  file('./7_Book/00EE/01/E022-8.p',ruleD64)).
fof(f8846,plain,(
  ( ! [X14] : (midp(X14,X14,X14)) ) | ~spl37_232),
  inference(subsumption_resolution,[],[f8818,f6223])).
fof(f8818,plain,(
  ( ! [X14] : (~coll(X14,X14,X14) | midp(X14,X14,X14)) ) | ~spl37_232),
  inference(resolution,[],[f8716,f426])).
fof(f8716,plain,(
  ( ! [X35] : (cong(X35,X35,X35,X35)) ) | ~spl37_232),
  inference(subsumption_resolution,[],[f8712,f8682])).
fof(f8682,plain,(
  ( ! [X2,X0,X1] : (cyclic(X0,X0,X1,X2)) ) | ~spl37_232),
  inference(resolution,[],[f4346,f2302])).
fof(f4346,plain,(
  ( ! [X68,X66,X67] : (~para(X66,X67,X66,X67) | cyclic(X67,X67,X66,X68)) )),
  inference(subsumption_resolution,[],[f2465,f2466])).
fof(f2466,plain,(
  ( ! [X70,X71,X69] : (~para(X69,X70,X69,X70) | coll(X69,X71,X70) | cyclic(X70,X70,X69,X71)) )),
  inference(resolution,[],[f841,f400])).
fof(f400,plain,(
  ( ! [X2,X0,X3,X1] : (~eqangle(X2,X0,X2,X1,X3,X0,X3,X1) | coll(X2,X3,X0) | cyclic(X0,X1,X2,X3)) )),
  inference(cnf_transformation,[],[f222])).
fof(f222,plain,(
  ! [X0,X1,X2,X3] : (cyclic(X0,X1,X2,X3) | coll(X2,X3,X0) | ~eqangle(X2,X0,X2,X1,X3,X0,X3,X1))),
  inference(flattening,[],[f221])).
fof(f221,plain,(
  ! [X0,X1,X2,X3] : (cyclic(X0,X1,X2,X3) | (coll(X2,X3,X0) | ~eqangle(X2,X0,X2,X1,X3,X0,X3,X1)))),
  inference(ennf_transformation,[],[f121])).
fof(f121,plain,(
  ! [X0,X1,X2,X3] : ((~coll(X2,X3,X0) & eqangle(X2,X0,X2,X1,X3,X0,X3,X1)) => cyclic(X0,X1,X2,X3))),
  inference(rectify,[],[f42])).
fof(f42,axiom,(
  ! [X0,X1,X8,X9] : ((~coll(X8,X9,X0) & eqangle(X8,X0,X8,X1,X9,X0,X9,X1)) => cyclic(X0,X1,X8,X9))),
  file('./7_Book/00EE/01/E022-8.p',ruleD42a)).
fof(f841,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (eqangle(X0,X1,X2,X3,X4,X5,X4,X5) | ~para(X0,X1,X2,X3)) )),
  inference(resolution,[],[f379,f398])).
fof(f398,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (eqangle(X0,X1,X4,X5,X2,X3,X4,X5) | ~para(X0,X1,X2,X3)) )),
  inference(cnf_transformation,[],[f219])).
fof(f219,plain,(
  ! [X0,X1,X2,X3,X4,X5] : (eqangle(X0,X1,X4,X5,X2,X3,X4,X5) | ~para(X0,X1,X2,X3))),
  inference(ennf_transformation,[],[f119])).
fof(f119,plain,(
  ! [X0,X1,X2,X3,X4,X5] : (para(X0,X1,X2,X3) => eqangle(X0,X1,X4,X5,X2,X3,X4,X5))),
  inference(rectify,[],[f40])).
fof(f40,axiom,(
  ! [X0,X1,X2,X3,X8,X9] : (para(X0,X1,X2,X3) => eqangle(X0,X1,X8,X9,X2,X3,X8,X9))),
  file('./7_Book/00EE/01/E022-8.p',ruleD40)).
fof(f379,plain,(
  ( ! [X6,X4,X2,X0,X7,X5,X3,X1] : (~eqangle(X0,X1,X2,X3,X4,X5,X6,X7) | eqangle(X0,X1,X4,X5,X2,X3,X6,X7)) )),
  inference(cnf_transformation,[],[f195])).
fof(f195,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6,X7] : (eqangle(X0,X1,X4,X5,X2,X3,X6,X7) | ~eqangle(X0,X1,X2,X3,X4,X5,X6,X7))),
  inference(ennf_transformation,[],[f103])).
fof(f103,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6,X7] : (eqangle(X0,X1,X2,X3,X4,X5,X6,X7) => eqangle(X0,X1,X4,X5,X2,X3,X6,X7))),
  inference(rectify,[],[f21])).
fof(f21,axiom,(
  ! [X0,X1,X2,X3,X8,X9,X10,X11] : (eqangle(X0,X1,X2,X3,X8,X9,X10,X11) => eqangle(X0,X1,X8,X9,X2,X3,X10,X11))),
  file('./7_Book/00EE/01/E022-8.p',ruleD21)).
fof(f2465,plain,(
  ( ! [X68,X66,X67] : (~para(X66,X67,X66,X67) | ~coll(X66,X68,X67) | cyclic(X67,X67,X66,X68)) )),
  inference(resolution,[],[f841,f401])).
fof(f401,plain,(
  ( ! [X2,X0,X3,X1] : (~eqangle(X2,X0,X2,X1,X3,X0,X3,X1) | ~coll(X2,X3,X1) | cyclic(X0,X1,X2,X3)) )),
  inference(cnf_transformation,[],[f224])).
fof(f224,plain,(
  ! [X0,X1,X2,X3] : (cyclic(X0,X1,X2,X3) | ~coll(X2,X3,X1) | ~eqangle(X2,X0,X2,X1,X3,X0,X3,X1))),
  inference(flattening,[],[f223])).
fof(f223,plain,(
  ! [X0,X1,X2,X3] : (cyclic(X0,X1,X2,X3) | (~coll(X2,X3,X1) | ~eqangle(X2,X0,X2,X1,X3,X0,X3,X1)))),
  inference(ennf_transformation,[],[f122])).
fof(f122,plain,(
  ! [X0,X1,X2,X3] : ((coll(X2,X3,X1) & eqangle(X2,X0,X2,X1,X3,X0,X3,X1)) => cyclic(X0,X1,X2,X3))),
  inference(rectify,[],[f43])).
fof(f43,axiom,(
  ! [X0,X1,X8,X9] : ((coll(X8,X9,X1) & eqangle(X8,X0,X8,X1,X9,X0,X9,X1)) => cyclic(X0,X1,X8,X9))),
  file('./7_Book/00EE/01/E022-8.p',ruleD42b)).
fof(f8712,plain,(
  ( ! [X35,X36] : (cong(X35,X35,X35,X35) | ~cyclic(X35,X35,X36,X35)) ) | ~spl37_232),
  inference(duplicate_literal_removal,[],[f8699])).
fof(f8699,plain,(
  ( ! [X35,X36] : (cong(X35,X35,X35,X35) | ~cyclic(X35,X35,X36,X35) | ~cyclic(X35,X35,X36,X35)) ) | ~spl37_232),
  inference(resolution,[],[f8682,f1318])).
fof(f1318,plain,(
  ( ! [X10,X8,X7,X9] : (~cyclic(X7,X8,X9,X10) | cong(X7,X8,X7,X8) | ~cyclic(X7,X8,X9,X8) | ~cyclic(X7,X8,X9,X7)) )),
  inference(duplicate_literal_removal,[],[f1312])).
fof(f1312,plain,(
  ( ! [X10,X8,X7,X9] : (cong(X7,X8,X7,X8) | ~cyclic(X7,X8,X9,X10) | ~cyclic(X7,X8,X9,X8) | ~cyclic(X7,X8,X9,X7) | ~cyclic(X7,X8,X9,X10)) )),
  inference(resolution,[],[f402,f399])).
fof(f399,plain,(
  ( ! [X2,X0,X3,X1] : (eqangle(X2,X0,X2,X1,X3,X0,X3,X1) | ~cyclic(X0,X1,X2,X3)) )),
  inference(cnf_transformation,[],[f220])).
fof(f220,plain,(
  ! [X0,X1,X2,X3] : (eqangle(X2,X0,X2,X1,X3,X0,X3,X1) | ~cyclic(X0,X1,X2,X3))),
  inference(ennf_transformation,[],[f120])).
fof(f120,plain,(
  ! [X0,X1,X2,X3] : (cyclic(X0,X1,X2,X3) => eqangle(X2,X0,X2,X1,X3,X0,X3,X1))),
  inference(rectify,[],[f41])).
fof(f41,axiom,(
  ! [X0,X1,X8,X9] : (cyclic(X0,X1,X8,X9) => eqangle(X8,X0,X8,X1,X9,X0,X9,X1))),
  file('./7_Book/00EE/01/E022-8.p',ruleD41)).
fof(f402,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (~eqangle(X2,X0,X2,X1,X5,X3,X5,X4) | cong(X0,X1,X3,X4) | ~cyclic(X0,X1,X2,X5) | ~cyclic(X0,X1,X2,X4) | ~cyclic(X0,X1,X2,X3)) )),
  inference(cnf_transformation,[],[f226])).
fof(f226,plain,(
  ! [X0,X1,X2,X3,X4,X5] : (cong(X0,X1,X3,X4) | ~eqangle(X2,X0,X2,X1,X5,X3,X5,X4) | ~cyclic(X0,X1,X2,X5) | ~cyclic(X0,X1,X2,X4) | ~cyclic(X0,X1,X2,X3))),
  inference(flattening,[],[f225])).
fof(f225,plain,(
  ! [X0,X1,X2,X3,X4,X5] : (cong(X0,X1,X3,X4) | (~eqangle(X2,X0,X2,X1,X5,X3,X5,X4) | ~cyclic(X0,X1,X2,X5) | ~cyclic(X0,X1,X2,X4) | ~cyclic(X0,X1,X2,X3)))),
  inference(ennf_transformation,[],[f123])).
fof(f123,plain,(
  ! [X0,X1,X2,X3,X4,X5] : ((eqangle(X2,X0,X2,X1,X5,X3,X5,X4) & cyclic(X0,X1,X2,X5) & cyclic(X0,X1,X2,X4) & cyclic(X0,X1,X2,X3)) => cong(X0,X1,X3,X4))),
  inference(rectify,[],[f44])).
fof(f44,axiom,(
  ! [X0,X1,X2,X8,X9,X14] : ((eqangle(X2,X0,X2,X1,X14,X8,X14,X9) & cyclic(X0,X1,X2,X14) & cyclic(X0,X1,X2,X9) & cyclic(X0,X1,X2,X8)) => cong(X0,X1,X8,X9))),
  file('./7_Book/00EE/01/E022-8.p',ruleD43)).
fof(f16154,plain,(
  ( ! [X4] : (midp(X4,sK22,sK24) | ~midp(X4,sK22,sK22)) ) | (~spl37_650 | ~spl37_742)),
  inference(subsumption_resolution,[],[f16147,f9214])).
fof(f9214,plain,(
  para(sK22,sK22,sK22,sK24) | ~spl37_742),
  inference(avatar_component_clause,[],[f9213])).
fof(f9213,plain,(
  spl37_742 <=> para(sK22,sK22,sK22,sK24)),
  introduced(avatar_definition,[new_symbols(naming,[spl37_742])])).
fof(f16147,plain,(
  ( ! [X4] : (midp(X4,sK22,sK24) | ~para(sK22,sK22,sK22,sK24) | ~midp(X4,sK22,sK22)) ) | ~spl37_650),
  inference(resolution,[],[f8868,f423])).
fof(f8868,plain,(
  para(sK22,sK24,sK22,sK22) | ~spl37_650),
  inference(avatar_component_clause,[],[f8867])).
fof(f8867,plain,(
  spl37_650 <=> para(sK22,sK24,sK22,sK22)),
  introduced(avatar_definition,[new_symbols(naming,[spl37_650])])).
fof(f17351,plain,(
  ( ! [X0] : (cong(sK22,X0,sK22,sK27) | ~midp(X0,sK22,sK24)) ) | (~spl37_30 | ~spl37_244)),
  inference(resolution,[],[f17157,f2802])).
fof(f2802,plain,(
  ( ! [X14,X15,X13] : (cong(X14,X13,X14,X15) | ~midp(X13,X14,X15)) ) | ~spl37_244),
  inference(avatar_component_clause,[],[f2801])).
fof(f2801,plain,(
  spl37_244 <=> ! [X13,X15,X14] : (~midp(X13,X14,X15) | cong(X14,X13,X14,X15))),
  introduced(avatar_definition,[new_symbols(naming,[spl37_244])])).
fof(f17157,plain,(
  ( ! [X19,X20] : (~cong(X19,X20,sK22,sK24) | cong(X19,X20,sK22,sK27)) ) | ~spl37_30),
  inference(resolution,[],[f694,f383])).
fof(f383,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (~cong(X2,X3,X4,X5) | cong(X0,X1,X4,X5) | ~cong(X0,X1,X2,X3)) )),
  inference(cnf_transformation,[],[f201])).
fof(f201,plain,(
  ! [X0,X1,X2,X3,X4,X5] : (cong(X0,X1,X4,X5) | ~cong(X2,X3,X4,X5) | ~cong(X0,X1,X2,X3))),
  inference(flattening,[],[f200])).
fof(f200,plain,(
  ! [X0,X1,X2,X3,X4,X5] : (cong(X0,X1,X4,X5) | (~cong(X2,X3,X4,X5) | ~cong(X0,X1,X2,X3)))),
  inference(ennf_transformation,[],[f25])).
fof(f25,axiom,(
  ! [X0,X1,X2,X3,X4,X5] : ((cong(X2,X3,X4,X5) & cong(X0,X1,X2,X3)) => cong(X0,X1,X4,X5))),
  file('./7_Book/00EE/01/E022-8.p',ruleD25)).
fof(f694,plain,(
  cong(sK22,sK24,sK22,sK27) | ~spl37_30),
  inference(avatar_component_clause,[],[f693])).
fof(f693,plain,(
  spl37_30 <=> cong(sK22,sK24,sK22,sK27)),
  introduced(avatar_definition,[new_symbols(naming,[spl37_30])])).
fof(f18096,plain,(
  ( ! [X2] : (midp(sK24,X2,sK22) | ~midp(sK22,X2,sK27)) ) | (~spl37_54 | ~spl37_232 | ~spl37_238)),
  inference(subsumption_resolution,[],[f18089,f6271])).
fof(f18089,plain,(
  ( ! [X2] : (~coll(sK24,X2,sK22) | midp(sK24,X2,sK22) | ~midp(sK22,X2,sK27)) ) | (~spl37_54 | ~spl37_232 | ~spl37_238)),
  inference(resolution,[],[f16350,f404])).
fof(f404,plain,(
  ( ! [X4,X2,X0,X3,X1] : (~para(X3,X4,X1,X2) | ~coll(X4,X0,X2) | midp(X4,X0,X2) | ~midp(X3,X0,X1)) )),
  inference(cnf_transformation,[],[f230])).
fof(f230,plain,(
  ! [X0,X1,X2,X3,X4] : (midp(X4,X0,X2) | ~coll(X4,X0,X2) | ~para(X3,X4,X1,X2) | ~midp(X3,X0,X1))),
  inference(flattening,[],[f229])).
fof(f229,plain,(
  ! [X0,X1,X2,X3,X4] : (midp(X4,X0,X2) | (~coll(X4,X0,X2) | ~para(X3,X4,X1,X2) | ~midp(X3,X0,X1)))),
  inference(ennf_transformation,[],[f125])).
fof(f125,plain,(
  ! [X0,X1,X2,X3,X4] : ((coll(X4,X0,X2) & para(X3,X4,X1,X2) & midp(X3,X0,X1)) => midp(X4,X0,X2))),
  inference(rectify,[],[f46])).
fof(f46,axiom,(
  ! [X0,X1,X2,X4,X5] : ((coll(X5,X0,X2) & para(X4,X5,X1,X2) & midp(X4,X0,X1)) => midp(X5,X0,X2))),
  file('./7_Book/00EE/01/E022-8.p',ruleD45)).
fof(f16350,plain,(
  para(sK22,sK24,sK27,sK22) | (~spl37_54 | ~spl37_232 | ~spl37_238)),
  inference(subsumption_resolution,[],[f16324,f13134])).
fof(f13134,plain,(
  para(sK27,sK24,sK22,sK24) | (~spl37_232 | ~spl37_238)),
  inference(resolution,[],[f12321,f363])).
fof(f363,plain,(
  ( ! [X2,X0,X3,X1] : (~para(X0,X1,X2,X3) | para(X2,X3,X0,X1)) )),
  inference(cnf_transformation,[],[f173])).
fof(f173,plain,(
  ! [X0,X1,X2,X3] : (para(X2,X3,X0,X1) | ~para(X0,X1,X2,X3))),
  inference(ennf_transformation,[],[f5])).
fof(f5,axiom,(
  ! [X0,X1,X2,X3] : (para(X0,X1,X2,X3) => para(X2,X3,X0,X1))),
  file('./7_Book/00EE/01/E022-8.p',ruleD5)).
fof(f12321,plain,(
  para(sK22,sK24,sK27,sK24) | (~spl37_232 | ~spl37_238)),
  inference(resolution,[],[f11086,f793])).
fof(f793,plain,(
  ( ! [X8,X9] : (~perp(X8,X9,sK22,sK24) | para(X8,X9,sK27,sK24)) )),
  inference(resolution,[],[f367,f503])).
fof(f503,plain,(
  perp(sK22,sK24,sK27,sK24)),
  inference(resolution,[],[f365,f481])).
fof(f481,plain,(
  perp(sK22,sK24,sK24,sK27)),
  inference(cnf_transformation,[],[f358])).
fof(f365,plain,(
  ( ! [X2,X0,X3,X1] : (~perp(X0,X1,X2,X3) | perp(X0,X1,X3,X2)) )),
  inference(cnf_transformation,[],[f176])).
fof(f176,plain,(
  ! [X0,X1,X2,X3] : (perp(X0,X1,X3,X2) | ~perp(X0,X1,X2,X3))),
  inference(ennf_transformation,[],[f7])).
fof(f7,axiom,(
  ! [X0,X1,X2,X3] : (perp(X0,X1,X2,X3) => perp(X0,X1,X3,X2))),
  file('./7_Book/00EE/01/E022-8.p',ruleD7)).
fof(f367,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (~perp(X2,X3,X4,X5) | para(X0,X1,X4,X5) | ~perp(X0,X1,X2,X3)) )),
  inference(cnf_transformation,[],[f179])).
fof(f179,plain,(
  ! [X0,X1,X2,X3,X4,X5] : (para(X0,X1,X4,X5) | ~perp(X2,X3,X4,X5) | ~perp(X0,X1,X2,X3))),
  inference(flattening,[],[f178])).
fof(f178,plain,(
  ! [X0,X1,X2,X3,X4,X5] : (para(X0,X1,X4,X5) | (~perp(X2,X3,X4,X5) | ~perp(X0,X1,X2,X3)))),
  inference(ennf_transformation,[],[f9])).
fof(f9,axiom,(
  ! [X0,X1,X2,X3,X4,X5] : ((perp(X2,X3,X4,X5) & perp(X0,X1,X2,X3)) => para(X0,X1,X4,X5))),
  file('./7_Book/00EE/01/E022-8.p',ruleD9)).
fof(f11086,plain,(
  ( ! [X0,X1] : (perp(X0,X1,X0,X1)) ) | (~spl37_232 | ~spl37_238)),
  inference(resolution,[],[f2481,f2302])).
fof(f2481,plain,(
  ( ! [X47,X45,X46,X44] : (~para(X44,X45,X46,X47) | perp(X44,X45,X46,X47)) ) | ~spl37_238),
  inference(avatar_component_clause,[],[f2480])).
fof(f2480,plain,(
  spl37_238 <=> ! [X44,X46,X45,X47] : (~para(X44,X45,X46,X47) | perp(X44,X45,X46,X47))),
  introduced(avatar_definition,[new_symbols(naming,[spl37_238])])).
fof(f16324,plain,(
  ~para(sK27,sK24,sK22,sK24) | para(sK22,sK24,sK27,sK22) | ~spl37_54),
  inference(resolution,[],[f772,f2345])).
fof(f2345,plain,(
  ( ! [X21,X22,X20] : (~cong(X20,X21,X20,X22) | ~para(X20,X22,X21,X22) | para(X21,X22,X20,X21)) )),
  inference(resolution,[],[f829,f432])).
fof(f432,plain,(
  ( ! [X6,X4,X2,X0,X7,X5,X3,X1] : (~eqangle(X0,X1,X2,X3,X4,X5,X6,X7) | ~para(X4,X5,X6,X7) | para(X0,X1,X2,X3)) )),
  inference(cnf_transformation,[],[f279])).
fof(f279,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6,X7] : (para(X0,X1,X2,X3) | ~para(X4,X5,X6,X7) | ~eqangle(X0,X1,X2,X3,X4,X5,X6,X7))),
  inference(flattening,[],[f278])).
fof(f278,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6,X7] : (para(X0,X1,X2,X3) | (~para(X4,X5,X6,X7) | ~eqangle(X0,X1,X2,X3,X4,X5,X6,X7)))),
  inference(ennf_transformation,[],[f146])).
fof(f146,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6,X7] : ((para(X4,X5,X6,X7) & eqangle(X0,X1,X2,X3,X4,X5,X6,X7)) => para(X0,X1,X2,X3))),
  inference(rectify,[],[f74])).
fof(f74,axiom,(
  ! [X0,X1,X2,X3,X8,X9,X10,X11] : ((para(X8,X9,X10,X11) & eqangle(X0,X1,X2,X3,X8,X9,X10,X11)) => para(X0,X1,X2,X3))),
  file('./7_Book/00EE/01/E022-8.p',ruleD73)).
fof(f829,plain,(
  ( ! [X12,X10,X11] : (eqangle(X10,X11,X12,X10,X12,X11,X10,X11) | ~cong(X12,X10,X12,X11)) )),
  inference(resolution,[],[f377,f405])).
fof(f405,plain,(
  ( ! [X2,X0,X1] : (eqangle(X2,X0,X0,X1,X0,X1,X2,X1) | ~cong(X2,X0,X2,X1)) )),
  inference(cnf_transformation,[],[f231])).
fof(f231,plain,(
  ! [X0,X1,X2] : (eqangle(X2,X0,X0,X1,X0,X1,X2,X1) | ~cong(X2,X0,X2,X1))),
  inference(ennf_transformation,[],[f126])).
fof(f126,plain,(
  ! [X0,X1,X2] : (cong(X2,X0,X2,X1) => eqangle(X2,X0,X0,X1,X0,X1,X2,X1))),
  inference(rectify,[],[f47])).
fof(f47,axiom,(
  ! [X0,X1,X7] : (cong(X7,X0,X7,X1) => eqangle(X7,X0,X0,X1,X0,X1,X7,X1))),
  file('./7_Book/00EE/01/E022-8.p',ruleD46)).
fof(f377,plain,(
  ( ! [X6,X4,X2,X0,X7,X5,X3,X1] : (~eqangle(X0,X1,X2,X3,X4,X5,X6,X7) | eqangle(X2,X3,X0,X1,X6,X7,X4,X5)) )),
  inference(cnf_transformation,[],[f193])).
fof(f193,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6,X7] : (eqangle(X2,X3,X0,X1,X6,X7,X4,X5) | ~eqangle(X0,X1,X2,X3,X4,X5,X6,X7))),
  inference(ennf_transformation,[],[f101])).
fof(f101,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6,X7] : (eqangle(X0,X1,X2,X3,X4,X5,X6,X7) => eqangle(X2,X3,X0,X1,X6,X7,X4,X5))),
  inference(rectify,[],[f19])).
fof(f19,axiom,(
  ! [X0,X1,X2,X3,X8,X9,X10,X11] : (eqangle(X0,X1,X2,X3,X8,X9,X10,X11) => eqangle(X2,X3,X0,X1,X10,X11,X8,X9))),
  file('./7_Book/00EE/01/E022-8.p',ruleD19)).
fof(f772,plain,(
  cong(sK27,sK22,sK27,sK24) | ~spl37_54),
  inference(avatar_component_clause,[],[f771])).
fof(f771,plain,(
  spl37_54 <=> cong(sK27,sK22,sK27,sK24)),
  introduced(avatar_definition,[new_symbols(naming,[spl37_54])])).
fof(f16990,plain,(
  ( ! [X2] : (midp(sK22,X2,sK22) | ~midp(sK24,X2,sK22)) ) | (~spl37_232 | ~spl37_650)),
  inference(subsumption_resolution,[],[f16983,f6189])).
fof(f16983,plain,(
  ( ! [X2] : (~coll(sK22,X2,sK22) | midp(sK22,X2,sK22) | ~midp(sK24,X2,sK22)) ) | ~spl37_650),
  inference(resolution,[],[f16148,f404])).
fof(f16148,plain,(
  para(sK24,sK22,sK22,sK22) | ~spl37_650),
  inference(resolution,[],[f8868,f2191])).
fof(f2191,plain,(
  ( ! [X2,X0,X3,X1] : (~para(X0,X1,X2,X3) | para(X1,X0,X2,X3)) )),
  inference(resolution,[],[f823,f397])).
fof(f397,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (~eqangle(X0,X1,X4,X5,X2,X3,X4,X5) | para(X0,X1,X2,X3)) )),
  inference(cnf_transformation,[],[f218])).
fof(f218,plain,(
  ! [X0,X1,X2,X3,X4,X5] : (para(X0,X1,X2,X3) | ~eqangle(X0,X1,X4,X5,X2,X3,X4,X5))),
  inference(ennf_transformation,[],[f118])).
fof(f118,plain,(
  ! [X0,X1,X2,X3,X4,X5] : (eqangle(X0,X1,X4,X5,X2,X3,X4,X5) => para(X0,X1,X2,X3))),
  inference(rectify,[],[f39])).
fof(f39,axiom,(
  ! [X0,X1,X2,X3,X8,X9] : (eqangle(X0,X1,X8,X9,X2,X3,X8,X9) => para(X0,X1,X2,X3))),
  file('./7_Book/00EE/01/E022-8.p',ruleD39)).
fof(f823,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (eqangle(X0,X1,X2,X3,X4,X5,X2,X3) | ~para(X1,X0,X4,X5)) )),
  inference(resolution,[],[f376,f398])).
fof(f376,plain,(
  ( ! [X6,X4,X2,X0,X7,X5,X3,X1] : (~eqangle(X0,X1,X2,X3,X4,X5,X6,X7) | eqangle(X1,X0,X2,X3,X4,X5,X6,X7)) )),
  inference(cnf_transformation,[],[f192])).
fof(f192,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6,X7] : (eqangle(X1,X0,X2,X3,X4,X5,X6,X7) | ~eqangle(X0,X1,X2,X3,X4,X5,X6,X7))),
  inference(ennf_transformation,[],[f100])).
fof(f100,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6,X7] : (eqangle(X0,X1,X2,X3,X4,X5,X6,X7) => eqangle(X1,X0,X2,X3,X4,X5,X6,X7))),
  inference(rectify,[],[f18])).
fof(f18,axiom,(
  ! [X0,X1,X2,X3,X8,X9,X10,X11] : (eqangle(X0,X1,X2,X3,X8,X9,X10,X11) => eqangle(X1,X0,X2,X3,X8,X9,X10,X11))),
  file('./7_Book/00EE/01/E022-8.p',ruleD18)).
fof(f16380,plain,(
  ( ! [X3] : (midp(sK22,X3,sK24) | ~midp(sK22,X3,sK22)) ) | (~spl37_232 | ~spl37_742)),
  inference(subsumption_resolution,[],[f16373,f6271])).
fof(f16373,plain,(
  ( ! [X3] : (~coll(sK22,X3,sK24) | midp(sK22,X3,sK24) | ~midp(sK22,X3,sK22)) ) | ~spl37_742),
  inference(resolution,[],[f9214,f404])).
fof(f21474,plain,(
  ( ! [X10,X11] : (perp(sK22,sK24,X10,X11) | ~midp(sK22,X10,sK24)) ) | (~spl37_60 | ~spl37_232 | ~spl37_238)),
  inference(subsumption_resolution,[],[f10039,f15935])).
fof(f15935,plain,(
  ( ! [X0] : (midp(X0,sK22,sK27)) ) | (~spl37_232 | ~spl37_238)),
  inference(resolution,[],[f15923,f369])).
fof(f369,plain,(
  ( ! [X2,X0,X1] : (~midp(X2,X1,X0) | midp(X2,X0,X1)) )),
  inference(cnf_transformation,[],[f182])).
fof(f182,plain,(
  ! [X0,X1,X2] : (midp(X2,X0,X1) | ~midp(X2,X1,X0))),
  inference(ennf_transformation,[],[f97])).
fof(f97,plain,(
  ! [X0,X1,X2] : (midp(X2,X1,X0) => midp(X2,X0,X1))),
  inference(rectify,[],[f11])).
fof(f11,axiom,(
  ! [X0,X1,X6] : (midp(X6,X1,X0) => midp(X6,X0,X1))),
  file('./7_Book/00EE/01/E022-8.p',ruleD11)).
fof(f15923,plain,(
  ( ! [X4] : (midp(X4,sK27,sK22)) ) | (~spl37_232 | ~spl37_238)),
  inference(subsumption_resolution,[],[f15922,f10254])).
fof(f15922,plain,(
  ( ! [X4] : (midp(X4,sK27,sK22) | ~midp(X4,sK24,sK24)) ) | (~spl37_232 | ~spl37_238)),
  inference(subsumption_resolution,[],[f15915,f13282])).
fof(f13282,plain,(
  para(sK24,sK27,sK24,sK22) | (~spl37_232 | ~spl37_238)),
  inference(resolution,[],[f12324,f362])).
fof(f362,plain,(
  ( ! [X2,X0,X3,X1] : (~para(X0,X1,X2,X3) | para(X0,X1,X3,X2)) )),
  inference(cnf_transformation,[],[f172])).
fof(f172,plain,(
  ! [X0,X1,X2,X3] : (para(X0,X1,X3,X2) | ~para(X0,X1,X2,X3))),
  inference(ennf_transformation,[],[f4])).
fof(f4,axiom,(
  ! [X0,X1,X2,X3] : (para(X0,X1,X2,X3) => para(X0,X1,X3,X2))),
  file('./7_Book/00EE/01/E022-8.p',ruleD4)).
fof(f12324,plain,(
  para(sK24,sK27,sK22,sK24) | (~spl37_232 | ~spl37_238)),
  inference(resolution,[],[f11086,f794])).
fof(f794,plain,(
  ( ! [X10,X11] : (~perp(X10,X11,sK24,sK27) | para(X10,X11,sK22,sK24)) )),
  inference(resolution,[],[f367,f506])).
fof(f506,plain,(
  perp(sK24,sK27,sK22,sK24)),
  inference(resolution,[],[f366,f481])).
fof(f15915,plain,(
  ( ! [X4] : (midp(X4,sK27,sK22) | ~para(sK24,sK27,sK24,sK22) | ~midp(X4,sK24,sK24)) ) | (~spl37_232 | ~spl37_238)),
  inference(resolution,[],[f13186,f423])).
fof(f13186,plain,(
  para(sK24,sK22,sK24,sK27) | (~spl37_232 | ~spl37_238)),
  inference(resolution,[],[f12322,f2191])).
fof(f12322,plain,(
  para(sK22,sK24,sK24,sK27) | (~spl37_232 | ~spl37_238)),
  inference(resolution,[],[f11086,f792])).
fof(f792,plain,(
  ( ! [X6,X7] : (~perp(X6,X7,sK22,sK24) | para(X6,X7,sK24,sK27)) )),
  inference(resolution,[],[f367,f481])).
fof(f10039,plain,(
  ( ! [X10,X11] : (perp(sK22,sK24,X10,X11) | ~midp(X11,sK22,sK27) | ~midp(sK22,X10,sK24)) ) | ~spl37_60),
  inference(resolution,[],[f1379,f5567])).
fof(f5567,plain,(
  ( ! [X14,X12,X13] : (cong(X12,X13,X14,X13) | ~midp(X12,X13,X14)) ) | ~spl37_60),
  inference(resolution,[],[f1023,f381])).
fof(f381,plain,(
  ( ! [X2,X0,X3,X1] : (~cong(X0,X1,X2,X3) | cong(X0,X1,X3,X2)) )),
  inference(cnf_transformation,[],[f198])).
fof(f198,plain,(
  ! [X0,X1,X2,X3] : (cong(X0,X1,X3,X2) | ~cong(X0,X1,X2,X3))),
  inference(ennf_transformation,[],[f23])).
fof(f23,axiom,(
  ! [X0,X1,X2,X3] : (cong(X0,X1,X2,X3) => cong(X0,X1,X3,X2))),
  file('./7_Book/00EE/01/E022-8.p',ruleD23)).
fof(f1023,plain,(
  ( ! [X10,X11,X9] : (cong(X9,X10,X10,X11) | ~midp(X9,X10,X11)) ) | ~spl37_60),
  inference(avatar_component_clause,[],[f1022])).
fof(f1022,plain,(
  spl37_60 <=> ! [X9,X11,X10] : (cong(X9,X10,X10,X11) | ~midp(X9,X10,X11))),
  introduced(avatar_definition,[new_symbols(naming,[spl37_60])])).
fof(f1379,plain,(
  ( ! [X6,X5] : (~cong(sK22,X6,sK24,X6) | perp(sK22,sK24,X6,X5) | ~midp(X5,sK22,sK27)) )),
  inference(resolution,[],[f572,f415])).
fof(f415,plain,(
  ( ! [X2,X0,X3,X1] : (~cong(X0,X3,X1,X3) | perp(X0,X1,X2,X3) | ~cong(X0,X2,X1,X2)) )),
  inference(cnf_transformation,[],[f251])).
fof(f251,plain,(
  ! [X0,X1,X2,X3] : (perp(X0,X1,X2,X3) | ~cong(X0,X3,X1,X3) | ~cong(X0,X2,X1,X2))),
  inference(flattening,[],[f250])).
fof(f250,plain,(
  ! [X0,X1,X2,X3] : (perp(X0,X1,X2,X3) | (~cong(X0,X3,X1,X3) | ~cong(X0,X2,X1,X2)))),
  inference(ennf_transformation,[],[f135])).
fof(f135,plain,(
  ! [X0,X1,X2,X3] : ((cong(X0,X3,X1,X3) & cong(X0,X2,X1,X2)) => perp(X0,X1,X2,X3))),
  inference(rectify,[],[f57])).
fof(f57,axiom,(
  ! [X0,X1,X8,X9] : ((cong(X0,X9,X1,X9) & cong(X0,X8,X1,X8)) => perp(X0,X1,X8,X9))),
  file('./7_Book/00EE/01/E022-8.p',ruleD56)).
fof(f572,plain,(
  ( ! [X0] : (cong(sK22,X0,sK24,X0) | ~midp(X0,sK22,sK27)) )),
  inference(resolution,[],[f411,f481])).
fof(f411,plain,(
  ( ! [X2,X0,X3,X1] : (~perp(X0,X1,X1,X2) | ~midp(X3,X0,X2) | cong(X0,X3,X1,X3)) )),
  inference(cnf_transformation,[],[f243])).
fof(f243,plain,(
  ! [X0,X1,X2,X3] : (cong(X0,X3,X1,X3) | ~midp(X3,X0,X2) | ~perp(X0,X1,X1,X2))),
  inference(flattening,[],[f242])).
fof(f242,plain,(
  ! [X0,X1,X2,X3] : (cong(X0,X3,X1,X3) | (~midp(X3,X0,X2) | ~perp(X0,X1,X1,X2)))),
  inference(ennf_transformation,[],[f132])).
fof(f132,plain,(
  ! [X0,X1,X2,X3] : ((midp(X3,X0,X2) & perp(X0,X1,X1,X2)) => cong(X0,X3,X1,X3))),
  inference(rectify,[],[f53])).
fof(f53,axiom,(
  ! [X0,X1,X2,X6] : ((midp(X6,X0,X2) & perp(X0,X1,X1,X2)) => cong(X0,X6,X1,X6))),
  file('./7_Book/00EE/01/E022-8.p',ruleD52)).
fof(f35181,plain,(
  ( ! [X35,X33,X34,X32] : (para(X32,X33,X34,X35) | ~perp(X32,X33,sK22,sK24)) ) | (~spl37_30 | ~spl37_54 | ~spl37_60 | ~spl37_232 | ~spl37_238 | ~spl37_244 | ~spl37_650 | ~spl37_742)),
  inference(resolution,[],[f21475,f367])).
fof(f17122,plain,(
  spl37_29 | ~spl37_54 | ~spl37_232 | ~spl37_238),
  inference(avatar_contradiction_clause,[],[f17121])).
fof(f17121,plain,(
  $false | (~spl37_29 | ~spl37_54 | ~spl37_232 | ~spl37_238)),
  inference(resolution,[],[f17082,f688])).
fof(f688,plain,(
  ~midp(sK24,sK24,sK27) | ~spl37_29),
  inference(avatar_component_clause,[],[f687])).
fof(f687,plain,(
  spl37_29 <=> ~midp(sK24,sK24,sK27)),
  introduced(avatar_definition,[new_symbols(naming,[spl37_29])])).
fof(f17082,plain,(
  ( ! [X4] : (midp(X4,sK24,sK27)) ) | (~spl37_54 | ~spl37_232 | ~spl37_238)),
  inference(subsumption_resolution,[],[f17081,f10254])).
fof(f17081,plain,(
  ( ! [X4] : (midp(X4,sK24,sK27) | ~midp(X4,sK22,sK22)) ) | (~spl37_54 | ~spl37_232 | ~spl37_238)),
  inference(subsumption_resolution,[],[f17072,f17069])).
fof(f17069,plain,(
  para(sK22,sK24,sK22,sK27) | (~spl37_54 | ~spl37_232 | ~spl37_238)),
  inference(resolution,[],[f16348,f363])).
fof(f16348,plain,(
  para(sK22,sK27,sK22,sK24) | (~spl37_54 | ~spl37_232 | ~spl37_238)),
  inference(subsumption_resolution,[],[f16322,f12321])).
fof(f16322,plain,(
  ~para(sK22,sK24,sK27,sK24) | para(sK22,sK27,sK22,sK24) | ~spl37_54),
  inference(resolution,[],[f772,f2254])).
fof(f2254,plain,(
  ( ! [X23,X21,X22] : (~cong(X21,X22,X21,X23) | ~para(X22,X23,X21,X23) | para(X22,X21,X22,X23)) )),
  inference(resolution,[],[f825,f432])).
fof(f825,plain,(
  ( ! [X12,X10,X11] : (eqangle(X10,X11,X10,X12,X10,X12,X11,X12) | ~cong(X11,X10,X11,X12)) )),
  inference(resolution,[],[f376,f405])).
fof(f17072,plain,(
  ( ! [X4] : (midp(X4,sK24,sK27) | ~para(sK22,sK24,sK22,sK27) | ~midp(X4,sK22,sK22)) ) | (~spl37_54 | ~spl37_232 | ~spl37_238)),
  inference(resolution,[],[f16348,f423])).
fof(f16186,plain,(
  spl37_53 | ~spl37_232 | ~spl37_650 | ~spl37_742),
  inference(avatar_contradiction_clause,[],[f16183])).
fof(f16183,plain,(
  $false | (~spl37_53 | ~spl37_232 | ~spl37_650 | ~spl37_742)),
  inference(resolution,[],[f16155,f766])).
fof(f766,plain,(
  ~midp(sK24,sK22,sK24) | ~spl37_53),
  inference(avatar_component_clause,[],[f765])).
fof(f765,plain,(
  spl37_53 <=> ~midp(sK24,sK22,sK24)),
  introduced(avatar_definition,[new_symbols(naming,[spl37_53])])).
fof(f16025,plain,(
  spl37_650 | ~spl37_232 | ~spl37_238),
  inference(avatar_split_clause,[],[f16024,f2480,f2301,f8867])).
fof(f16024,plain,(
  para(sK22,sK24,sK22,sK22) | (~spl37_232 | ~spl37_238)),
  inference(subsumption_resolution,[],[f16023,f15935])).
fof(f16023,plain,(
  para(sK22,sK24,sK22,sK22) | ~midp(sK22,sK22,sK27) | ~spl37_232),
  inference(subsumption_resolution,[],[f14526,f2302])).
fof(f14526,plain,(
  ~para(sK22,sK24,sK22,sK24) | para(sK22,sK24,sK22,sK22) | ~midp(sK22,sK22,sK27)),
  inference(resolution,[],[f2345,f1378])).
fof(f1378,plain,(
  ( ! [X4] : (cong(sK22,X4,X4,sK24) | ~midp(X4,sK22,sK27)) )),
  inference(resolution,[],[f572,f381])).
fof(f16003,plain,(
  ~spl37_232 | ~spl37_238 | spl37_649),
  inference(avatar_contradiction_clause,[],[f15998])).
fof(f15998,plain,(
  $false | (~spl37_232 | ~spl37_238 | ~spl37_649)),
  inference(resolution,[],[f15935,f8862])).
fof(f8862,plain,(
  ~midp(sK22,sK22,sK27) | ~spl37_649),
  inference(avatar_component_clause,[],[f8861])).
fof(f8861,plain,(
  spl37_649 <=> ~midp(sK22,sK22,sK27)),
  introduced(avatar_definition,[new_symbols(naming,[spl37_649])])).
fof(f11055,plain,(
  ~spl37_232 | ~spl37_236 | ~spl37_424),
  inference(avatar_contradiction_clause,[],[f11054])).
fof(f11054,plain,(
  $false | (~spl37_232 | ~spl37_236 | ~spl37_424)),
  inference(subsumption_resolution,[],[f11046,f8846])).
fof(f11046,plain,(
  ~midp(sK24,sK24,sK24) | (~spl37_232 | ~spl37_236 | ~spl37_424)),
  inference(resolution,[],[f9680,f1443])).
fof(f1443,plain,(
  ( ! [X0,X1] : (para(X0,X0,X0,X1) | ~midp(X0,X1,X0)) )),
  inference(resolution,[],[f570,f518])).
fof(f518,plain,(
  ( ! [X2,X0,X1] : (cong(X0,X1,X0,X2) | ~midp(X0,X2,X1)) )),
  inference(resolution,[],[f382,f427])).
fof(f427,plain,(
  ( ! [X2,X0,X1] : (cong(X0,X1,X0,X2) | ~midp(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f270])).
fof(f270,plain,(
  ! [X0,X1,X2] : (cong(X0,X1,X0,X2) | ~midp(X0,X1,X2))),
  inference(ennf_transformation,[],[f69])).
fof(f69,axiom,(
  ! [X0,X1,X2] : (midp(X0,X1,X2) => cong(X0,X1,X0,X2))),
  file('./7_Book/00EE/01/E022-8.p',ruleD68)).
fof(f382,plain,(
  ( ! [X2,X0,X3,X1] : (~cong(X0,X1,X2,X3) | cong(X2,X3,X0,X1)) )),
  inference(cnf_transformation,[],[f199])).
fof(f199,plain,(
  ! [X0,X1,X2,X3] : (cong(X2,X3,X0,X1) | ~cong(X0,X1,X2,X3))),
  inference(ennf_transformation,[],[f24])).
fof(f24,axiom,(
  ! [X0,X1,X2,X3] : (cong(X0,X1,X2,X3) => cong(X2,X3,X0,X1))),
  file('./7_Book/00EE/01/E022-8.p',ruleD24)).
fof(f570,plain,(
  ( ! [X0,X1] : (~cong(X0,X0,X0,X1) | para(X0,X0,X0,X1)) )),
  inference(resolution,[],[f405,f397])).
fof(f9680,plain,(
  ( ! [X11] : (~para(sK24,X11,sK24,sK24)) ) | (~spl37_232 | ~spl37_236 | ~spl37_424)),
  inference(subsumption_resolution,[],[f9671,f8766])).
fof(f8766,plain,(
  ( ! [X24,X23,X22] : (~para(X22,X24,X22,X23) | ~perp(X22,X23,X22,X24)) ) | (~spl37_232 | ~spl37_236)),
  inference(subsumption_resolution,[],[f8750,f2478])).
fof(f2478,plain,(
  ( ! [X48,X49] : (~perp(X48,X49,X48,X49)) ) | ~spl37_236),
  inference(avatar_component_clause,[],[f2477])).
fof(f2477,plain,(
  spl37_236 <=> ! [X49,X48] : ~perp(X48,X49,X48,X49)),
  introduced(avatar_definition,[new_symbols(naming,[spl37_236])])).
fof(f8750,plain,(
  ( ! [X24,X23,X22] : (perp(X22,X23,X22,X23) | ~para(X22,X24,X22,X23) | ~perp(X22,X23,X22,X24)) ) | ~spl37_232),
  inference(resolution,[],[f8689,f1011])).
fof(f1011,plain,(
  ( ! [X6,X8,X7,X9] : (~cyclic(X9,X8,X6,X7) | perp(X9,X7,X6,X7) | ~para(X9,X8,X6,X7) | ~perp(X6,X7,X6,X8)) )),
  inference(resolution,[],[f433,f413])).
fof(f413,plain,(
  ( ! [X2,X0,X3,X1] : (eqangle(X0,X3,X2,X3,X2,X3,X2,X1) | ~para(X0,X1,X2,X3) | ~cyclic(X0,X1,X2,X3)) )),
  inference(cnf_transformation,[],[f247])).
fof(f247,plain,(
  ! [X0,X1,X2,X3] : (eqangle(X0,X3,X2,X3,X2,X3,X2,X1) | ~para(X0,X1,X2,X3) | ~cyclic(X0,X1,X2,X3))),
  inference(flattening,[],[f246])).
fof(f246,plain,(
  ! [X0,X1,X2,X3] : (eqangle(X0,X3,X2,X3,X2,X3,X2,X1) | (~para(X0,X1,X2,X3) | ~cyclic(X0,X1,X2,X3)))),
  inference(ennf_transformation,[],[f55])).
fof(f55,axiom,(
  ! [X0,X1,X2,X3] : ((para(X0,X1,X2,X3) & cyclic(X0,X1,X2,X3)) => eqangle(X0,X3,X2,X3,X2,X3,X2,X1))),
  file('./7_Book/00EE/01/E022-8.p',ruleD54)).
fof(f433,plain,(
  ( ! [X6,X4,X2,X0,X7,X5,X3,X1] : (~eqangle(X0,X1,X2,X3,X4,X5,X6,X7) | ~perp(X4,X5,X6,X7) | perp(X0,X1,X2,X3)) )),
  inference(cnf_transformation,[],[f281])).
fof(f281,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6,X7] : (perp(X0,X1,X2,X3) | ~perp(X4,X5,X6,X7) | ~eqangle(X0,X1,X2,X3,X4,X5,X6,X7))),
  inference(flattening,[],[f280])).
fof(f280,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6,X7] : (perp(X0,X1,X2,X3) | (~perp(X4,X5,X6,X7) | ~eqangle(X0,X1,X2,X3,X4,X5,X6,X7)))),
  inference(ennf_transformation,[],[f147])).
fof(f147,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6,X7] : ((perp(X4,X5,X6,X7) & eqangle(X0,X1,X2,X3,X4,X5,X6,X7)) => perp(X0,X1,X2,X3))),
  inference(rectify,[],[f75])).
fof(f75,axiom,(
  ! [X0,X1,X2,X3,X8,X9,X10,X11] : ((perp(X8,X9,X10,X11) & eqangle(X0,X1,X2,X3,X8,X9,X10,X11)) => perp(X0,X1,X2,X3))),
  file('./7_Book/00EE/01/E022-8.p',ruleD74)).
fof(f8689,plain,(
  ( ! [X4,X5,X3] : (cyclic(X3,X4,X3,X5)) ) | ~spl37_232),
  inference(resolution,[],[f8682,f373])).
fof(f373,plain,(
  ( ! [X2,X0,X3,X1] : (~cyclic(X0,X1,X2,X3) | cyclic(X0,X2,X1,X3)) )),
  inference(cnf_transformation,[],[f188])).
fof(f188,plain,(
  ! [X0,X1,X2,X3] : (cyclic(X0,X2,X1,X3) | ~cyclic(X0,X1,X2,X3))),
  inference(ennf_transformation,[],[f15])).
fof(f15,axiom,(
  ! [X0,X1,X2,X3] : (cyclic(X0,X1,X2,X3) => cyclic(X0,X2,X1,X3))),
  file('./7_Book/00EE/01/E022-8.p',ruleD15)).
fof(f9671,plain,(
  ( ! [X11] : (perp(sK24,sK24,sK24,X11) | ~para(sK24,X11,sK24,sK24)) ) | ~spl37_424),
  inference(resolution,[],[f4471,f965])).
fof(f965,plain,(
  ( ! [X2,X0,X3,X1] : (~circle(X0,X1,X3,X1) | perp(X0,X1,X1,X2) | ~para(X1,X2,X1,X1)) )),
  inference(resolution,[],[f408,f398])).
fof(f408,plain,(
  ( ! [X4,X2,X0,X3,X1] : (~eqangle(X0,X4,X0,X1,X2,X0,X2,X1) | perp(X3,X0,X0,X4) | ~circle(X3,X0,X1,X2)) )),
  inference(cnf_transformation,[],[f237])).
fof(f237,plain,(
  ! [X0,X1,X2,X3,X4] : (perp(X3,X0,X0,X4) | ~eqangle(X0,X4,X0,X1,X2,X0,X2,X1) | ~circle(X3,X0,X1,X2))),
  inference(flattening,[],[f236])).
fof(f236,plain,(
  ! [X0,X1,X2,X3,X4] : (perp(X3,X0,X0,X4) | (~eqangle(X0,X4,X0,X1,X2,X0,X2,X1) | ~circle(X3,X0,X1,X2)))),
  inference(ennf_transformation,[],[f129])).
fof(f129,plain,(
  ! [X0,X1,X2,X3,X4] : ((eqangle(X0,X4,X0,X1,X2,X0,X2,X1) & circle(X3,X0,X1,X2)) => perp(X3,X0,X0,X4))),
  inference(rectify,[],[f50])).
fof(f50,axiom,(
  ! [X0,X1,X2,X7,X15] : ((eqangle(X0,X15,X0,X1,X2,X0,X2,X1) & circle(X7,X0,X1,X2)) => perp(X7,X0,X0,X15))),
  file('./7_Book/00EE/01/E022-8.p',ruleD49)).
fof(f4471,plain,(
  circle(sK24,sK24,sK24,sK24) | ~spl37_424),
  inference(avatar_component_clause,[],[f4470])).
fof(f4470,plain,(
  spl37_424 <=> circle(sK24,sK24,sK24,sK24)),
  introduced(avatar_definition,[new_symbols(naming,[spl37_424])])).
fof(f9655,plain,(
  ~spl37_232 | spl37_425),
  inference(avatar_contradiction_clause,[],[f9654])).
fof(f9654,plain,(
  $false | (~spl37_232 | ~spl37_425)),
  inference(subsumption_resolution,[],[f9644,f8846])).
fof(f9644,plain,(
  ~midp(sK24,sK24,sK24) | ~spl37_425),
  inference(duplicate_literal_removal,[],[f9639])).
fof(f9639,plain,(
  ~midp(sK24,sK24,sK24) | ~midp(sK24,sK24,sK24) | ~spl37_425),
  inference(resolution,[],[f2164,f4474])).
fof(f4474,plain,(
  ~circle(sK24,sK24,sK24,sK24) | ~spl37_425),
  inference(avatar_component_clause,[],[f4473])).
fof(f4473,plain,(
  spl37_425 <=> ~circle(sK24,sK24,sK24,sK24)),
  introduced(avatar_definition,[new_symbols(naming,[spl37_425])])).
fof(f2164,plain,(
  ( ! [X2,X0,X3,X1] : (circle(X0,X1,X2,X3) | ~midp(X0,X1,X3) | ~midp(X0,X2,X1)) )),
  inference(resolution,[],[f809,f518])).
fof(f809,plain,(
  ( ! [X2,X0,X3,X1] : (~cong(X0,X1,X0,X2) | circle(X0,X1,X2,X3) | ~midp(X0,X1,X3)) )),
  inference(resolution,[],[f370,f427])).
fof(f370,plain,(
  ( ! [X2,X0,X3,X1] : (~cong(X3,X0,X3,X2) | circle(X3,X0,X1,X2) | ~cong(X3,X0,X3,X1)) )),
  inference(cnf_transformation,[],[f184])).
fof(f184,plain,(
  ! [X0,X1,X2,X3] : (circle(X3,X0,X1,X2) | ~cong(X3,X0,X3,X2) | ~cong(X3,X0,X3,X1))),
  inference(flattening,[],[f183])).
fof(f183,plain,(
  ! [X0,X1,X2,X3] : (circle(X3,X0,X1,X2) | (~cong(X3,X0,X3,X2) | ~cong(X3,X0,X3,X1)))),
  inference(ennf_transformation,[],[f98])).
fof(f98,plain,(
  ! [X0,X1,X2,X3] : ((cong(X3,X0,X3,X2) & cong(X3,X0,X3,X1)) => circle(X3,X0,X1,X2))),
  inference(rectify,[],[f12])).
fof(f12,axiom,(
  ! [X0,X1,X2,X7] : ((cong(X7,X0,X7,X2) & cong(X7,X0,X7,X1)) => circle(X7,X0,X1,X2))),
  file('./7_Book/00EE/01/E022-8.p',ruleD12)).
fof(f9215,plain,(
  spl37_742 | ~spl37_649),
  inference(avatar_split_clause,[],[f7062,f8861,f9213])).
fof(f7062,plain,(
  ~midp(sK22,sK22,sK27) | para(sK22,sK22,sK22,sK24)),
  inference(resolution,[],[f1378,f570])).
fof(f5159,plain,(
  ~spl37_234),
  inference(avatar_contradiction_clause,[],[f5158])).
fof(f5158,plain,(
  $false | ~spl37_234),
  inference(resolution,[],[f2305,f1429])).
fof(f1429,plain,(
  para(sK24,sK27,sK24,sK27)),
  inference(resolution,[],[f792,f506])).
fof(f2305,plain,(
  ( ! [X0,X1] : (~para(X0,X1,X0,X1)) ) | ~spl37_234),
  inference(avatar_component_clause,[],[f2304])).
fof(f2304,plain,(
  spl37_234 <=> ! [X1,X0] : ~para(X0,X1,X0,X1)),
  introduced(avatar_definition,[new_symbols(naming,[spl37_234])])).
fof(f4117,plain,(
  ~spl37_62 | ~spl37_232),
  inference(avatar_contradiction_clause,[],[f4116])).
fof(f4116,plain,(
  $false | (~spl37_62 | ~spl37_232)),
  inference(subsumption_resolution,[],[f4115,f2819])).
fof(f2819,plain,(
  ( ! [X0,X1] : (~midp(X0,X1,X0)) ) | ~spl37_62),
  inference(subsumption_resolution,[],[f2813,f369])).
fof(f2813,plain,(
  ( ! [X0,X1] : (~midp(X0,X0,X1) | ~midp(X0,X1,X0)) ) | ~spl37_62),
  inference(resolution,[],[f1026,f518])).
fof(f1026,plain,(
  ( ! [X6,X8,X7] : (~cong(X6,X7,X7,X8) | ~midp(X6,X7,X8)) ) | ~spl37_62),
  inference(avatar_component_clause,[],[f1025])).
fof(f1025,plain,(
  spl37_62 <=> ! [X7,X8,X6] : (~cong(X6,X7,X7,X8) | ~midp(X6,X7,X8))),
  introduced(avatar_definition,[new_symbols(naming,[spl37_62])])).
fof(f4115,plain,(
  ( ! [X14] : (midp(X14,X14,X14)) ) | ~spl37_232),
  inference(subsumption_resolution,[],[f4096,f3071])).
fof(f3071,plain,(
  ( ! [X6,X5] : (coll(X5,X5,X6)) ) | ~spl37_232),
  inference(resolution,[],[f3034,f359])).
fof(f3034,plain,(
  ( ! [X4,X3] : (coll(X3,X4,X3)) ) | ~spl37_232),
  inference(resolution,[],[f3020,f360])).
fof(f3020,plain,(
  ( ! [X0,X1] : (coll(X0,X1,X1)) ) | ~spl37_232),
  inference(resolution,[],[f2302,f425])).
fof(f4096,plain,(
  ( ! [X14] : (~coll(X14,X14,X14) | midp(X14,X14,X14)) ) | ~spl37_232),
  inference(resolution,[],[f3993,f426])).
fof(f3993,plain,(
  ( ! [X0] : (cong(X0,X0,X0,X0)) ) | ~spl37_232),
  inference(subsumption_resolution,[],[f3992,f2486])).
fof(f2486,plain,(
  ( ! [X70,X71,X69] : (cyclic(X70,X70,X69,X71)) ) | ~spl37_232),
  inference(subsumption_resolution,[],[f2485,f2484])).
fof(f2484,plain,(
  ( ! [X68,X66,X67] : (~coll(X66,X68,X67) | cyclic(X67,X67,X66,X68)) ) | ~spl37_232),
  inference(subsumption_resolution,[],[f2465,f2302])).
fof(f2485,plain,(
  ( ! [X70,X71,X69] : (coll(X69,X71,X70) | cyclic(X70,X70,X69,X71)) ) | ~spl37_232),
  inference(subsumption_resolution,[],[f2466,f2302])).
fof(f3992,plain,(
  ( ! [X0,X1] : (cong(X0,X0,X0,X0) | ~cyclic(X0,X0,X1,X0)) ) | ~spl37_232),
  inference(duplicate_literal_removal,[],[f3988])).
fof(f3988,plain,(
  ( ! [X0,X1] : (cong(X0,X0,X0,X0) | ~cyclic(X0,X0,X1,X0) | ~cyclic(X0,X0,X1,X0)) ) | ~spl37_232),
  inference(resolution,[],[f1318,f2486])).
fof(f2807,plain,(
  spl37_244 | spl37_62),
  inference(avatar_split_clause,[],[f2794,f1025,f2801])).
fof(f2794,plain,(
  ( ! [X14,X12,X10,X15,X13,X11] : (~midp(X10,X11,X12) | ~midp(X13,X14,X15) | ~cong(X10,X11,X11,X12) | cong(X14,X13,X14,X15)) )),
  inference(resolution,[],[f849,f434])).
fof(f434,plain,(
  ( ! [X6,X4,X2,X0,X7,X5,X3,X1] : (~eqratio(X0,X1,X2,X3,X4,X5,X6,X7) | ~cong(X4,X5,X6,X7) | cong(X0,X1,X2,X3)) )),
  inference(cnf_transformation,[],[f283])).
fof(f283,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6,X7] : (cong(X0,X1,X2,X3) | ~cong(X4,X5,X6,X7) | ~eqratio(X0,X1,X2,X3,X4,X5,X6,X7))),
  inference(flattening,[],[f282])).
fof(f282,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6,X7] : (cong(X0,X1,X2,X3) | (~cong(X4,X5,X6,X7) | ~eqratio(X0,X1,X2,X3,X4,X5,X6,X7)))),
  inference(ennf_transformation,[],[f148])).
fof(f148,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6,X7] : ((cong(X4,X5,X6,X7) & eqratio(X0,X1,X2,X3,X4,X5,X6,X7)) => cong(X0,X1,X2,X3))),
  inference(rectify,[],[f76])).
fof(f76,axiom,(
  ! [X0,X1,X2,X3,X8,X9,X10,X11] : ((cong(X8,X9,X10,X11) & eqratio(X0,X1,X2,X3,X8,X9,X10,X11)) => cong(X0,X1,X2,X3))),
  file('./7_Book/00EE/01/E022-8.p',ruleD75)).
fof(f849,plain,(
  ( ! [X6,X10,X8,X7,X11,X9] : (eqratio(X6,X7,X6,X8,X9,X10,X10,X11) | ~midp(X9,X10,X11) | ~midp(X7,X6,X8)) )),
  inference(resolution,[],[f384,f429])).
fof(f429,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (eqratio(X4,X0,X0,X1,X5,X2,X2,X3) | ~midp(X5,X2,X3) | ~midp(X4,X0,X1)) )),
  inference(cnf_transformation,[],[f273])).
fof(f273,plain,(
  ! [X0,X1,X2,X3,X4,X5] : (eqratio(X4,X0,X0,X1,X5,X2,X2,X3) | ~midp(X5,X2,X3) | ~midp(X4,X0,X1))),
  inference(flattening,[],[f272])).
fof(f272,plain,(
  ! [X0,X1,X2,X3,X4,X5] : (eqratio(X4,X0,X0,X1,X5,X2,X2,X3) | (~midp(X5,X2,X3) | ~midp(X4,X0,X1)))),
  inference(ennf_transformation,[],[f145])).
fof(f145,plain,(
  ! [X0,X1,X2,X3,X4,X5] : ((midp(X5,X2,X3) & midp(X4,X0,X1)) => eqratio(X4,X0,X0,X1,X5,X2,X2,X3))),
  inference(rectify,[],[f71])).
fof(f71,axiom,(
  ! [X0,X1,X2,X3,X6,X16] : ((midp(X16,X2,X3) & midp(X6,X0,X1)) => eqratio(X6,X0,X0,X1,X16,X2,X2,X3))),
  file('./7_Book/00EE/01/E022-8.p',ruleD70)).
fof(f384,plain,(
  ( ! [X6,X4,X2,X0,X7,X5,X3,X1] : (~eqratio(X0,X1,X2,X3,X4,X5,X6,X7) | eqratio(X1,X0,X2,X3,X4,X5,X6,X7)) )),
  inference(cnf_transformation,[],[f202])).
fof(f202,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6,X7] : (eqratio(X1,X0,X2,X3,X4,X5,X6,X7) | ~eqratio(X0,X1,X2,X3,X4,X5,X6,X7))),
  inference(ennf_transformation,[],[f105])).
fof(f105,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6,X7] : (eqratio(X0,X1,X2,X3,X4,X5,X6,X7) => eqratio(X1,X0,X2,X3,X4,X5,X6,X7))),
  inference(rectify,[],[f26])).
fof(f26,axiom,(
  ! [X0,X1,X2,X3,X8,X9,X10,X11] : (eqratio(X0,X1,X2,X3,X8,X9,X10,X11) => eqratio(X1,X0,X2,X3,X8,X9,X10,X11))),
  file('./7_Book/00EE/01/E022-8.p',ruleD26)).
fof(f2482,plain,(
  spl37_236 | spl37_238),
  inference(avatar_split_clause,[],[f2459,f2480,f2477])).
fof(f2459,plain,(
  ( ! [X47,X45,X48,X46,X44,X49] : (~para(X44,X45,X46,X47) | ~perp(X48,X49,X48,X49) | perp(X44,X45,X46,X47)) )),
  inference(resolution,[],[f841,f433])).
fof(f2306,plain,(
  spl37_232 | spl37_234),
  inference(avatar_split_clause,[],[f2277,f2304,f2301])).
fof(f2277,plain,(
  ( ! [X2,X0,X3,X1] : (~para(X0,X1,X0,X1) | para(X2,X3,X2,X3)) )),
  inference(resolution,[],[f827,f397])).
fof(f827,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (eqangle(X0,X1,X2,X3,X0,X1,X4,X5) | ~para(X2,X3,X4,X5)) )),
  inference(resolution,[],[f377,f398])).
fof(f1027,plain,(
  spl37_60 | spl37_62),
  inference(avatar_split_clause,[],[f1020,f1025,f1022])).
fof(f1020,plain,(
  ( ! [X6,X10,X8,X7,X11,X9] : (~cong(X6,X7,X7,X8) | cong(X9,X10,X10,X11) | ~midp(X6,X7,X8) | ~midp(X9,X10,X11)) )),
  inference(resolution,[],[f434,f429])).
fof(f773,plain,(
  ~spl37_53 | spl37_54),
  inference(avatar_split_clause,[],[f681,f771,f765])).
fof(f681,plain,(
  cong(sK27,sK22,sK27,sK24) | ~midp(sK24,sK22,sK24)),
  inference(resolution,[],[f414,f509])).
fof(f509,plain,(
  perp(sK27,sK24,sK22,sK24)),
  inference(resolution,[],[f503,f366])).
fof(f414,plain,(
  ( ! [X2,X0,X3,X1] : (~perp(X3,X2,X0,X1) | cong(X3,X0,X3,X1) | ~midp(X2,X0,X1)) )),
  inference(cnf_transformation,[],[f249])).
fof(f249,plain,(
  ! [X0,X1,X2,X3] : (cong(X3,X0,X3,X1) | ~perp(X3,X2,X0,X1) | ~midp(X2,X0,X1))),
  inference(flattening,[],[f248])).
fof(f248,plain,(
  ! [X0,X1,X2,X3] : (cong(X3,X0,X3,X1) | (~perp(X3,X2,X0,X1) | ~midp(X2,X0,X1)))),
  inference(ennf_transformation,[],[f134])).
fof(f134,plain,(
  ! [X0,X1,X2,X3] : ((perp(X3,X2,X0,X1) & midp(X2,X0,X1)) => cong(X3,X0,X3,X1))),
  inference(rectify,[],[f56])).
fof(f56,axiom,(
  ! [X0,X1,X6,X7] : ((perp(X7,X6,X0,X1) & midp(X6,X0,X1)) => cong(X7,X0,X7,X1))),
  file('./7_Book/00EE/01/E022-8.p',ruleD55)).
fof(f695,plain,(
  ~spl37_29 | spl37_30),
  inference(avatar_split_clause,[],[f675,f693,f687])).
fof(f675,plain,(
  cong(sK22,sK24,sK22,sK27) | ~midp(sK24,sK24,sK27)),
  inference(resolution,[],[f414,f481])).
% SZS output end Proof for E022-8
% ------------------------------
% Version: Vampire 4.4.0 (commit 7916d27 on 2019-08-23 08:50:16 +0100)
% Termination reason: Refutation

% Memory used [KB]: 32877
% Time elapsed: 1.224 s
% ------------------------------
% ------------------------------
