% Refutation found. Thanks to Tanya!
% SZS status Theorem for E072-12
% SZS output start Proof for E072-12
9. ! [X0,X1,X2,X3,X4,X5] : ((perp(X2,X3,X4,X5) & perp(X0,X1,X2,X3)) => para(X0,X1,X4,X5)) [input]
18. ! [X0,X1,X2,X3,X8,X9,X10,X11] : (eqangle(X0,X1,X2,X3,X8,X9,X10,X11) => eqangle(X1,X0,X2,X3,X8,X9,X10,X11)) [input]
19. ! [X0,X1,X2,X3,X8,X9,X10,X11] : (eqangle(X0,X1,X2,X3,X8,X9,X10,X11) => eqangle(X2,X3,X0,X1,X10,X11,X8,X9)) [input]
21. ! [X0,X1,X2,X3,X8,X9,X10,X11] : (eqangle(X0,X1,X2,X3,X8,X9,X10,X11) => eqangle(X0,X1,X8,X9,X2,X3,X10,X11)) [input]
22. ! [X0,X1,X2,X3,X8,X9,X10,X11,X4,X5,X12,X13] : ((eqangle(X8,X9,X10,X11,X4,X5,X12,X13) & eqangle(X0,X1,X2,X3,X8,X9,X10,X11)) => eqangle(X0,X1,X2,X3,X4,X5,X12,X13)) [input]
40. ! [X0,X1,X2,X3,X8,X9] : (para(X0,X1,X2,X3) => eqangle(X0,X1,X8,X9,X2,X3,X8,X9)) [input]
94. ! [X0,X1,X2,X3,X4,X5,X12,X13,X18,X19,X20,X21,X22,X23,X24,X25,X26,X27] : (coll(X13,X1,X5) & circle(X0,X1,X13,X27) & coll(X12,X1,X4) & circle(X2,X1,X12,X26) & coll(X5,X3,X4) & circle(X2,X3,X5,X25) & circle(X0,X1,X4,X24) & perp(X1,X3,X3,X4) & circle(X2,X1,X3,X23) & circle(X0,X1,X3,X22) & circle(X2,X1,X20,X21) & circle(X0,X1,X18,X19)) => eqangle(X13,X3,X3,X1,X1,X3,X3,X12) [input]
95. ~! [X12,X3,X1,X13] : (! [X0,X1,X2,X3,X4,X5,X12,X13,X18,X19,X20,X21,X22,X23,X24,X25,X26,X27] : (coll(X13,X1,X5) & circle(X0,X1,X13,X27) & coll(X12,X1,X4) & circle(X2,X1,X12,X26) & coll(X5,X3,X4) & circle(X2,X3,X5,X25) & circle(X0,X1,X4,X24) & perp(X1,X3,X3,X4) & circle(X2,X1,X3,X23) & circle(X0,X1,X3,X22) & circle(X2,X1,X20,X21) & circle(X0,X1,X18,X19)) => eqangle(X13,X3,X3,X1,X1,X3,X3,X12)) [negated conjecture 94]
99. ! [X0,X1,X2,X3,X4,X5,X6,X7] : (eqangle(X0,X1,X2,X3,X4,X5,X6,X7) => eqangle(X1,X0,X2,X3,X4,X5,X6,X7)) [rectify 18]
100. ! [X0,X1,X2,X3,X4,X5,X6,X7] : (eqangle(X0,X1,X2,X3,X4,X5,X6,X7) => eqangle(X2,X3,X0,X1,X6,X7,X4,X5)) [rectify 19]
102. ! [X0,X1,X2,X3,X4,X5,X6,X7] : (eqangle(X0,X1,X2,X3,X4,X5,X6,X7) => eqangle(X0,X1,X4,X5,X2,X3,X6,X7)) [rectify 21]
103. ! [X0,X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11] : ((eqangle(X4,X5,X6,X7,X8,X9,X10,X11) & eqangle(X0,X1,X2,X3,X4,X5,X6,X7)) => eqangle(X0,X1,X2,X3,X8,X9,X10,X11)) [rectify 22]
118. ! [X0,X1,X2,X3,X4,X5] : (para(X0,X1,X2,X3) => eqangle(X0,X1,X4,X5,X2,X3,X4,X5)) [rectify 40]
165. ~! [X0,X1,X2,X3] : (! [X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21] : (coll(X11,X5,X9) & circle(X4,X5,X11,X21) & coll(X10,X5,X8) & circle(X6,X5,X10,X20) & coll(X9,X7,X8) & circle(X6,X7,X9,X19) & circle(X4,X5,X8,X18) & perp(X5,X7,X7,X8) & circle(X6,X5,X7,X17) & circle(X4,X5,X7,X16) & circle(X6,X5,X14,X15) & circle(X4,X5,X12,X13)) => eqangle(X3,X1,X1,X2,X2,X1,X1,X0)) [rectify 95]
176. ! [X0,X1,X2,X3,X4,X5] : (para(X0,X1,X4,X5) | (~perp(X2,X3,X4,X5) | ~perp(X0,X1,X2,X3))) [ennf transformation 9]
177. ! [X0,X1,X2,X3,X4,X5] : (para(X0,X1,X4,X5) | ~perp(X2,X3,X4,X5) | ~perp(X0,X1,X2,X3)) [flattening 176]
190. ! [X0,X1,X2,X3,X4,X5,X6,X7] : (eqangle(X1,X0,X2,X3,X4,X5,X6,X7) | ~eqangle(X0,X1,X2,X3,X4,X5,X6,X7)) [ennf transformation 99]
191. ! [X0,X1,X2,X3,X4,X5,X6,X7] : (eqangle(X2,X3,X0,X1,X6,X7,X4,X5) | ~eqangle(X0,X1,X2,X3,X4,X5,X6,X7)) [ennf transformation 100]
193. ! [X0,X1,X2,X3,X4,X5,X6,X7] : (eqangle(X0,X1,X4,X5,X2,X3,X6,X7) | ~eqangle(X0,X1,X2,X3,X4,X5,X6,X7)) [ennf transformation 102]
194. ! [X0,X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11] : (eqangle(X0,X1,X2,X3,X8,X9,X10,X11) | (~eqangle(X4,X5,X6,X7,X8,X9,X10,X11) | ~eqangle(X0,X1,X2,X3,X4,X5,X6,X7))) [ennf transformation 103]
195. ! [X0,X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11] : (eqangle(X0,X1,X2,X3,X8,X9,X10,X11) | ~eqangle(X4,X5,X6,X7,X8,X9,X10,X11) | ~eqangle(X0,X1,X2,X3,X4,X5,X6,X7)) [flattening 194]
217. ! [X0,X1,X2,X3,X4,X5] : (eqangle(X0,X1,X4,X5,X2,X3,X4,X5) | ~para(X0,X1,X2,X3)) [ennf transformation 118]
315. ? [X0,X1,X2,X3] : (~eqangle(X3,X1,X1,X2,X2,X1,X1,X0) & ! [X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21] : (coll(X11,X5,X9) & circle(X4,X5,X11,X21) & coll(X10,X5,X8) & circle(X6,X5,X10,X20) & coll(X9,X7,X8) & circle(X6,X7,X9,X19) & circle(X4,X5,X8,X18) & perp(X5,X7,X7,X8) & circle(X6,X5,X7,X17) & circle(X4,X5,X7,X16) & circle(X6,X5,X14,X15) & circle(X4,X5,X12,X13))) [ennf transformation 165]
352. ? [X0,X1,X2,X3] : (~eqangle(X3,X1,X1,X2,X2,X1,X1,X0) & ! [X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21] : (coll(X11,X5,X9) & circle(X4,X5,X11,X21) & coll(X10,X5,X8) & circle(X6,X5,X10,X20) & coll(X9,X7,X8) & circle(X6,X7,X9,X19) & circle(X4,X5,X8,X18) & perp(X5,X7,X7,X8) & circle(X6,X5,X7,X17) & circle(X4,X5,X7,X16) & circle(X6,X5,X14,X15) & circle(X4,X5,X12,X13))) => (~eqangle(sK23,sK21,sK21,sK22,sK22,sK21,sK21,sK20) & ! [X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21] : (coll(X11,X5,X9) & circle(X4,X5,X11,X21) & coll(X10,X5,X8) & circle(X6,X5,X10,X20) & coll(X9,X7,X8) & circle(X6,X7,X9,X19) & circle(X4,X5,X8,X18) & perp(X5,X7,X7,X8) & circle(X6,X5,X7,X17) & circle(X4,X5,X7,X16) & circle(X6,X5,X14,X15) & circle(X4,X5,X12,X13))) [choice axiom]
353. ~eqangle(sK23,sK21,sK21,sK22,sK22,sK21,sK21,sK20) & ! [X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21] : (coll(X11,X5,X9) & circle(X4,X5,X11,X21) & coll(X10,X5,X8) & circle(X6,X5,X10,X20) & coll(X9,X7,X8) & circle(X6,X7,X9,X19) & circle(X4,X5,X8,X18) & perp(X5,X7,X7,X8) & circle(X6,X5,X7,X17) & circle(X4,X5,X7,X16) & circle(X6,X5,X14,X15) & circle(X4,X5,X12,X13)) [skolemisation 315,352]
362. ~perp(X2,X3,X4,X5) | para(X0,X1,X4,X5) | ~perp(X0,X1,X2,X3) [cnf transformation 177]
371. ~eqangle(X0,X1,X2,X3,X4,X5,X6,X7) | eqangle(X1,X0,X2,X3,X4,X5,X6,X7) [cnf transformation 190]
372. ~eqangle(X0,X1,X2,X3,X4,X5,X6,X7) | eqangle(X2,X3,X0,X1,X6,X7,X4,X5) [cnf transformation 191]
374. ~eqangle(X0,X1,X2,X3,X4,X5,X6,X7) | eqangle(X0,X1,X4,X5,X2,X3,X6,X7) [cnf transformation 193]
375. ~eqangle(X4,X5,X6,X7,X8,X9,X10,X11) | eqangle(X0,X1,X2,X3,X8,X9,X10,X11) | ~eqangle(X0,X1,X2,X3,X4,X5,X6,X7) [cnf transformation 195]
393. eqangle(X0,X1,X4,X5,X2,X3,X4,X5) | ~para(X0,X1,X2,X3) [cnf transformation 217]
471. perp(X5,X7,X7,X8) [cnf transformation 353]
479. ~eqangle(sK23,sK21,sK21,sK22,sK22,sK21,sK21,sK20) [cnf transformation 353]
544. ~perp(X0,X1,X4,X2) | para(X0,X1,X2,X3) [resolution 362,471]
624. para(X0,X1,X2,X3) [resolution 544,471]
638. eqangle(X0,X1,X2,X3,X4,X5,X2,X3) | ~para(X1,X0,X4,X5) [resolution 371,393]
642. eqangle(X0,X1,X2,X3,X4,X5,X2,X3) [subsumption resolution 638,624]
643. eqangle(X0,X1,X2,X3,X0,X1,X4,X5) | ~para(X2,X3,X4,X5) [resolution 372,393]
647. eqangle(X0,X1,X2,X3,X0,X1,X4,X5) [subsumption resolution 643,624]
660. eqangle(X0,X1,X2,X3,X4,X5,X4,X5) [resolution 374,642]
662. eqangle(X12,X13,X12,X13,X14,X15,X16,X17) [resolution 374,647]
1650. eqangle(X191,X192,X193,X194,X195,X196,X197,X198) | ~eqangle(X191,X192,X193,X194,X199,X200,X199,X200) [resolution 375,662]
1663. eqangle(X191,X192,X193,X194,X195,X196,X197,X198) [subsumption resolution 1650,660]
1739. $false [resolution 1663,479]
% SZS output end Proof for E072-12
% ------------------------------
% Version: Vampire 4.4.0 (commit 7916d27 on 2019-08-23 08:50:16 +0100)
% Termination reason: Refutation

% Memory used [KB]: 5500
% Time elapsed: 0.039 s
% ------------------------------
% ------------------------------
