% Refutation found. Thanks to Tanya!
% SZS status Theorem for E074-23
% SZS output start Proof for E074-23
fof(f3052,plain,(
  $false),
  inference(resolution,[],[f2413,f481])).
fof(f481,plain,(
  ~eqangle(sK23,sK21,sK21,sK22,sK22,sK21,sK21,sK20)),
  inference(cnf_transformation,[],[f357])).
fof(f357,plain,(
  ~eqangle(sK23,sK21,sK21,sK22,sK22,sK21,sK21,sK20) & ! [X7,X8,X9,X10,X11,X12,X13] : (circle(X11,X7,X12,X13) & eqangle(X12,X7,X7,X9,X12,X7,X7,X8) & perp(X7,X9,sK24,X11) & midp(sK24,X7,X9) & perp(X8,X9,sK25,X11) & midp(sK25,X8,X9) & perp(X8,X7,sK26,X11) & midp(sK26,X8,X7) & coll(X10,X8,X9) & perp(X10,X7,X8,X9))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK20,sK21,sK22,sK23,sK24,sK25,sK26])],[f319,f356])).
fof(f356,plain,(
  ? [X0,X1,X2,X3,X4,X5,X6] : (~eqangle(X3,X1,X1,X2,X2,X1,X1,X0) & ! [X7,X8,X9,X10,X11,X12,X13] : (circle(X11,X7,X12,X13) & eqangle(X12,X7,X7,X9,X12,X7,X7,X8) & perp(X7,X9,X4,X11) & midp(X4,X7,X9) & perp(X8,X9,X5,X11) & midp(X5,X8,X9) & perp(X8,X7,X6,X11) & midp(X6,X8,X7) & coll(X10,X8,X9) & perp(X10,X7,X8,X9))) => (~eqangle(sK23,sK21,sK21,sK22,sK22,sK21,sK21,sK20) & ! [X13,X12,X11,X10,X9,X8,X7] : (circle(X11,X7,X12,X13) & eqangle(X12,X7,X7,X9,X12,X7,X7,X8) & perp(X7,X9,sK24,X11) & midp(sK24,X7,X9) & perp(X8,X9,sK25,X11) & midp(sK25,X8,X9) & perp(X8,X7,sK26,X11) & midp(sK26,X8,X7) & coll(X10,X8,X9) & perp(X10,X7,X8,X9)))),
  introduced(choice_axiom,[])).
fof(f319,plain,(
  ? [X0,X1,X2,X3,X4,X5,X6] : (~eqangle(X3,X1,X1,X2,X2,X1,X1,X0) & ! [X7,X8,X9,X10,X11,X12,X13] : (circle(X11,X7,X12,X13) & eqangle(X12,X7,X7,X9,X12,X7,X7,X8) & perp(X7,X9,X4,X11) & midp(X4,X7,X9) & perp(X8,X9,X5,X11) & midp(X5,X8,X9) & perp(X8,X7,X6,X11) & midp(X6,X8,X7) & coll(X10,X8,X9) & perp(X10,X7,X8,X9)))),
  inference(ennf_transformation,[],[f167])).
fof(f167,plain,(
  ~! [X0,X1,X2,X3,X4,X5,X6] : (! [X7,X8,X9,X10,X11,X12,X13] : (circle(X11,X7,X12,X13) & eqangle(X12,X7,X7,X9,X12,X7,X7,X8) & perp(X7,X9,X4,X11) & midp(X4,X7,X9) & perp(X8,X9,X5,X11) & midp(X5,X8,X9) & perp(X8,X7,X6,X11) & midp(X6,X8,X7) & coll(X10,X8,X9) & perp(X10,X7,X8,X9)) => eqangle(X3,X1,X1,X2,X2,X1,X1,X0))),
  inference(rectify,[],[f96])).
fof(f96,negated_conjecture,(
  ~! [X3,X0,X5,X4,X21,X20,X19] : (! [X0,X1,X2,X3,X4,X5,X18] : (circle(X4,X0,X5,X18) & eqangle(X5,X0,X0,X2,X5,X0,X0,X1) & perp(X0,X2,X21,X4) & midp(X21,X0,X2) & perp(X1,X2,X20,X4) & midp(X20,X1,X2) & perp(X1,X0,X19,X4) & midp(X19,X1,X0) & coll(X3,X1,X2) & perp(X3,X0,X1,X2)) => eqangle(X4,X0,X0,X5,X5,X0,X0,X3))),
  inference(negated_conjecture,[],[f95])).
fof(f95,conjecture,(
  ( ! [X4,X0,X21,X19,X5,X3,X20] : (! [X0,X1,X2,X3,X4,X5,X18] : (circle(X4,X0,X5,X18) & eqangle(X5,X0,X0,X2,X5,X0,X0,X1) & perp(X0,X2,X21,X4) & midp(X21,X0,X2) & perp(X1,X2,X20,X4) & midp(X20,X1,X2) & perp(X1,X0,X19,X4) & midp(X19,X1,X0) & coll(X3,X1,X2) & perp(X3,X0,X1,X2)) => eqangle(X4,X0,X0,X5,X5,X0,X0,X3)) )),
  file('./7_Book/00EE/10/E074-23.p',exemplo7Book00EE10E07423)).
fof(f2413,plain,(
  ( ! [X30,X28,X26,X23,X31,X29,X27,X25] : (eqangle(X26,X27,X28,X29,X30,X25,X31,X23)) )),
  inference(subsumption_resolution,[],[f2412,f2016])).
fof(f2016,plain,(
  ( ! [X24,X23,X21,X19,X25,X22,X20] : (eqangle(X19,X20,X21,X22,X23,X24,X25,X24)) )),
  inference(subsumption_resolution,[],[f1997,f1850])).
fof(f1850,plain,(
  ( ! [X24,X23,X21,X19,X22,X20,X18] : (eqangle(X18,X19,X20,X21,X22,X23,X22,X24)) )),
  inference(subsumption_resolution,[],[f1834,f768])).
fof(f768,plain,(
  ( ! [X6,X4,X8,X7,X5,X9] : (eqangle(X4,X5,X6,X7,X8,X9,X8,X9)) )),
  inference(subsumption_resolution,[],[f764,f516])).
fof(f516,plain,(
  ( ! [X2,X0,X3,X1] : (para(X0,X1,X2,X3)) )),
  inference(subsumption_resolution,[],[f513,f473])).
fof(f473,plain,(
  ( ! [X8,X7] : (midp(sK26,X8,X7)) )),
  inference(cnf_transformation,[],[f357])).
fof(f513,plain,(
  ( ! [X2,X0,X3,X1] : (para(X0,X1,X2,X3) | ~midp(sK26,X0,X2)) )),
  inference(resolution,[],[f421,f473])).
fof(f421,plain,(
  ( ! [X4,X2,X0,X3,X1] : (~midp(X4,X2,X3) | para(X0,X2,X1,X3) | ~midp(X4,X0,X1)) )),
  inference(cnf_transformation,[],[f262])).
fof(f262,plain,(
  ! [X0,X1,X2,X3,X4] : (para(X0,X2,X1,X3) | ~midp(X4,X2,X3) | ~midp(X4,X0,X1))),
  inference(flattening,[],[f261])).
fof(f261,plain,(
  ! [X0,X1,X2,X3,X4] : (para(X0,X2,X1,X3) | (~midp(X4,X2,X3) | ~midp(X4,X0,X1)))),
  inference(ennf_transformation,[],[f142])).
fof(f142,plain,(
  ! [X0,X1,X2,X3,X4] : ((midp(X4,X2,X3) & midp(X4,X0,X1)) => para(X0,X2,X1,X3))),
  inference(rectify,[],[f64])).
fof(f64,axiom,(
  ! [X0,X1,X2,X3,X6] : ((midp(X6,X2,X3) & midp(X6,X0,X1)) => para(X0,X2,X1,X3))),
  file('./7_Book/00EE/10/E074-23.p',ruleD63)).
fof(f764,plain,(
  ( ! [X6,X4,X8,X7,X5,X9] : (eqangle(X4,X5,X6,X7,X8,X9,X8,X9) | ~para(X4,X5,X6,X7)) )),
  inference(resolution,[],[f378,f397])).
fof(f397,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (eqangle(X0,X1,X4,X5,X2,X3,X4,X5) | ~para(X0,X1,X2,X3)) )),
  inference(cnf_transformation,[],[f219])).
fof(f219,plain,(
  ! [X0,X1,X2,X3,X4,X5] : (eqangle(X0,X1,X4,X5,X2,X3,X4,X5) | ~para(X0,X1,X2,X3))),
  inference(ennf_transformation,[],[f119])).
fof(f119,plain,(
  ! [X0,X1,X2,X3,X4,X5] : (para(X0,X1,X2,X3) => eqangle(X0,X1,X4,X5,X2,X3,X4,X5))),
  inference(rectify,[],[f40])).
fof(f40,axiom,(
  ! [X0,X1,X2,X3,X8,X9] : (para(X0,X1,X2,X3) => eqangle(X0,X1,X8,X9,X2,X3,X8,X9))),
  file('./7_Book/00EE/10/E074-23.p',ruleD40)).
fof(f378,plain,(
  ( ! [X6,X4,X2,X0,X7,X5,X3,X1] : (~eqangle(X0,X1,X2,X3,X4,X5,X6,X7) | eqangle(X0,X1,X4,X5,X2,X3,X6,X7)) )),
  inference(cnf_transformation,[],[f195])).
fof(f195,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6,X7] : (eqangle(X0,X1,X4,X5,X2,X3,X6,X7) | ~eqangle(X0,X1,X2,X3,X4,X5,X6,X7))),
  inference(ennf_transformation,[],[f103])).
fof(f103,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6,X7] : (eqangle(X0,X1,X2,X3,X4,X5,X6,X7) => eqangle(X0,X1,X4,X5,X2,X3,X6,X7))),
  inference(rectify,[],[f21])).
fof(f21,axiom,(
  ! [X0,X1,X2,X3,X8,X9,X10,X11] : (eqangle(X0,X1,X2,X3,X8,X9,X10,X11) => eqangle(X0,X1,X8,X9,X2,X3,X10,X11))),
  file('./7_Book/00EE/10/E074-23.p',ruleD21)).
fof(f1834,plain,(
  ( ! [X24,X23,X21,X19,X25,X22,X20,X18] : (eqangle(X18,X19,X20,X21,X22,X23,X22,X24) | ~eqangle(X18,X19,X20,X21,X25,X22,X25,X22)) )),
  inference(resolution,[],[f763,f379])).
fof(f379,plain,(
  ( ! [X6,X4,X2,X0,X10,X8,X7,X5,X3,X1,X11,X9] : (~eqangle(X4,X5,X6,X7,X8,X9,X10,X11) | eqangle(X0,X1,X2,X3,X8,X9,X10,X11) | ~eqangle(X0,X1,X2,X3,X4,X5,X6,X7)) )),
  inference(cnf_transformation,[],[f197])).
fof(f197,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11] : (eqangle(X0,X1,X2,X3,X8,X9,X10,X11) | ~eqangle(X4,X5,X6,X7,X8,X9,X10,X11) | ~eqangle(X0,X1,X2,X3,X4,X5,X6,X7))),
  inference(flattening,[],[f196])).
fof(f196,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11] : (eqangle(X0,X1,X2,X3,X8,X9,X10,X11) | (~eqangle(X4,X5,X6,X7,X8,X9,X10,X11) | ~eqangle(X0,X1,X2,X3,X4,X5,X6,X7)))),
  inference(ennf_transformation,[],[f104])).
fof(f104,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11] : ((eqangle(X4,X5,X6,X7,X8,X9,X10,X11) & eqangle(X0,X1,X2,X3,X4,X5,X6,X7)) => eqangle(X0,X1,X2,X3,X8,X9,X10,X11))),
  inference(rectify,[],[f22])).
fof(f22,axiom,(
  ! [X0,X1,X2,X3,X8,X9,X10,X11,X4,X5,X12,X13] : ((eqangle(X8,X9,X10,X11,X4,X5,X12,X13) & eqangle(X0,X1,X2,X3,X8,X9,X10,X11)) => eqangle(X0,X1,X2,X3,X4,X5,X12,X13))),
  file('./7_Book/00EE/10/E074-23.p',ruleD22)).
fof(f763,plain,(
  ( ! [X2,X0,X3,X1] : (eqangle(X0,X1,X0,X1,X1,X2,X1,X3)) )),
  inference(resolution,[],[f378,f479])).
fof(f479,plain,(
  ( ! [X12,X8,X7,X9] : (eqangle(X12,X7,X7,X9,X12,X7,X7,X8)) )),
  inference(cnf_transformation,[],[f357])).
fof(f1997,plain,(
  ( ! [X26,X24,X23,X21,X19,X25,X22,X20] : (eqangle(X19,X20,X21,X22,X23,X24,X25,X24) | ~eqangle(X19,X20,X21,X22,X25,X24,X25,X26)) )),
  inference(resolution,[],[f1003,f379])).
fof(f1003,plain,(
  ( ! [X10,X8,X7,X9] : (eqangle(X9,X10,X9,X8,X7,X10,X9,X10)) )),
  inference(subsumption_resolution,[],[f1002,f933])).
fof(f933,plain,(
  ( ! [X2,X0,X3,X1] : (cyclic(X0,X1,X2,X3)) )),
  inference(subsumption_resolution,[],[f929,f905])).
fof(f905,plain,(
  ( ! [X2,X0,X1] : (cyclic(X0,X0,X1,X2)) )),
  inference(subsumption_resolution,[],[f901,f872])).
fof(f872,plain,(
  ( ! [X12,X10,X11] : (cyclic(X10,X11,X11,X12)) )),
  inference(resolution,[],[f844,f371])).
fof(f371,plain,(
  ( ! [X2,X0,X3,X1] : (~cyclic(X0,X1,X2,X3) | cyclic(X0,X1,X3,X2)) )),
  inference(cnf_transformation,[],[f187])).
fof(f187,plain,(
  ! [X0,X1,X2,X3] : (cyclic(X0,X1,X3,X2) | ~cyclic(X0,X1,X2,X3))),
  inference(ennf_transformation,[],[f14])).
fof(f14,axiom,(
  ! [X0,X1,X2,X3] : (cyclic(X0,X1,X2,X3) => cyclic(X0,X1,X3,X2))),
  file('./7_Book/00EE/10/E074-23.p',ruleD14)).
fof(f844,plain,(
  ( ! [X8,X7,X9] : (cyclic(X7,X8,X9,X8)) )),
  inference(resolution,[],[f832,f372])).
fof(f372,plain,(
  ( ! [X2,X0,X3,X1] : (~cyclic(X0,X1,X2,X3) | cyclic(X0,X2,X1,X3)) )),
  inference(cnf_transformation,[],[f188])).
fof(f188,plain,(
  ! [X0,X1,X2,X3] : (cyclic(X0,X2,X1,X3) | ~cyclic(X0,X1,X2,X3))),
  inference(ennf_transformation,[],[f15])).
fof(f15,axiom,(
  ! [X0,X1,X2,X3] : (cyclic(X0,X1,X2,X3) => cyclic(X0,X2,X1,X3))),
  file('./7_Book/00EE/10/E074-23.p',ruleD15)).
fof(f832,plain,(
  ( ! [X4,X2,X3] : (cyclic(X4,X3,X2,X2)) )),
  inference(subsumption_resolution,[],[f831,f516])).
fof(f831,plain,(
  ( ! [X4,X2,X3] : (cyclic(X4,X3,X2,X2) | ~para(X2,X4,X2,X4)) )),
  inference(subsumption_resolution,[],[f826,f472])).
fof(f472,plain,(
  ( ! [X10,X8,X9] : (coll(X10,X8,X9)) )),
  inference(cnf_transformation,[],[f357])).
fof(f826,plain,(
  ( ! [X4,X2,X3] : (~coll(X2,X2,X3) | cyclic(X4,X3,X2,X2) | ~para(X2,X4,X2,X4)) )),
  inference(resolution,[],[f400,f397])).
fof(f400,plain,(
  ( ! [X2,X0,X3,X1] : (~eqangle(X2,X0,X2,X1,X3,X0,X3,X1) | ~coll(X2,X3,X1) | cyclic(X0,X1,X2,X3)) )),
  inference(cnf_transformation,[],[f224])).
fof(f224,plain,(
  ! [X0,X1,X2,X3] : (cyclic(X0,X1,X2,X3) | ~coll(X2,X3,X1) | ~eqangle(X2,X0,X2,X1,X3,X0,X3,X1))),
  inference(flattening,[],[f223])).
fof(f223,plain,(
  ! [X0,X1,X2,X3] : (cyclic(X0,X1,X2,X3) | (~coll(X2,X3,X1) | ~eqangle(X2,X0,X2,X1,X3,X0,X3,X1)))),
  inference(ennf_transformation,[],[f122])).
fof(f122,plain,(
  ! [X0,X1,X2,X3] : ((coll(X2,X3,X1) & eqangle(X2,X0,X2,X1,X3,X0,X3,X1)) => cyclic(X0,X1,X2,X3))),
  inference(rectify,[],[f43])).
fof(f43,axiom,(
  ! [X0,X1,X8,X9] : ((coll(X8,X9,X1) & eqangle(X8,X0,X8,X1,X9,X0,X9,X1)) => cyclic(X0,X1,X8,X9))),
  file('./7_Book/00EE/10/E074-23.p',ruleD42b)).
fof(f901,plain,(
  ( ! [X2,X0,X1] : (cyclic(X0,X0,X1,X2) | ~cyclic(X0,X0,X0,X1)) )),
  inference(resolution,[],[f863,f374])).
fof(f374,plain,(
  ( ! [X4,X2,X0,X3,X1] : (~cyclic(X0,X1,X2,X4) | cyclic(X1,X2,X3,X4) | ~cyclic(X0,X1,X2,X3)) )),
  inference(cnf_transformation,[],[f191])).
fof(f191,plain,(
  ! [X0,X1,X2,X3,X4] : (cyclic(X1,X2,X3,X4) | ~cyclic(X0,X1,X2,X4) | ~cyclic(X0,X1,X2,X3))),
  inference(flattening,[],[f190])).
fof(f190,plain,(
  ! [X0,X1,X2,X3,X4] : (cyclic(X1,X2,X3,X4) | (~cyclic(X0,X1,X2,X4) | ~cyclic(X0,X1,X2,X3)))),
  inference(ennf_transformation,[],[f17])).
fof(f17,axiom,(
  ! [X0,X1,X2,X3,X4] : ((cyclic(X0,X1,X2,X4) & cyclic(X0,X1,X2,X3)) => cyclic(X1,X2,X3,X4))),
  file('./7_Book/00EE/10/E074-23.p',ruleD17)).
fof(f863,plain,(
  ( ! [X8,X7] : (cyclic(X7,X7,X7,X8)) )),
  inference(resolution,[],[f835,f371])).
fof(f835,plain,(
  ( ! [X6,X5] : (cyclic(X5,X5,X6,X5)) )),
  inference(resolution,[],[f830,f372])).
fof(f830,plain,(
  ( ! [X0,X1] : (cyclic(X0,X1,X0,X0)) )),
  inference(subsumption_resolution,[],[f825,f472])).
fof(f825,plain,(
  ( ! [X0,X1] : (~coll(X0,X0,X1) | cyclic(X0,X1,X0,X0)) )),
  inference(resolution,[],[f400,f479])).
fof(f929,plain,(
  ( ! [X2,X0,X3,X1] : (cyclic(X0,X1,X2,X3) | ~cyclic(X0,X0,X1,X2)) )),
  inference(resolution,[],[f905,f374])).
fof(f1002,plain,(
  ( ! [X10,X8,X7,X9] : (~cyclic(X7,X8,X9,X10) | eqangle(X9,X10,X9,X8,X7,X10,X9,X10)) )),
  inference(subsumption_resolution,[],[f995,f516])).
fof(f995,plain,(
  ( ! [X10,X8,X7,X9] : (~para(X7,X8,X9,X10) | ~cyclic(X7,X8,X9,X10) | eqangle(X9,X10,X9,X8,X7,X10,X9,X10)) )),
  inference(resolution,[],[f412,f377])).
fof(f377,plain,(
  ( ! [X6,X4,X2,X0,X7,X5,X3,X1] : (~eqangle(X0,X1,X2,X3,X4,X5,X6,X7) | eqangle(X4,X5,X6,X7,X0,X1,X2,X3)) )),
  inference(cnf_transformation,[],[f194])).
fof(f194,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6,X7] : (eqangle(X4,X5,X6,X7,X0,X1,X2,X3) | ~eqangle(X0,X1,X2,X3,X4,X5,X6,X7))),
  inference(ennf_transformation,[],[f102])).
fof(f102,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6,X7] : (eqangle(X0,X1,X2,X3,X4,X5,X6,X7) => eqangle(X4,X5,X6,X7,X0,X1,X2,X3))),
  inference(rectify,[],[f20])).
fof(f20,axiom,(
  ! [X0,X1,X2,X3,X8,X9,X10,X11] : (eqangle(X0,X1,X2,X3,X8,X9,X10,X11) => eqangle(X8,X9,X10,X11,X0,X1,X2,X3))),
  file('./7_Book/00EE/10/E074-23.p',ruleD20)).
fof(f412,plain,(
  ( ! [X2,X0,X3,X1] : (eqangle(X0,X3,X2,X3,X2,X3,X2,X1) | ~para(X0,X1,X2,X3) | ~cyclic(X0,X1,X2,X3)) )),
  inference(cnf_transformation,[],[f247])).
fof(f247,plain,(
  ! [X0,X1,X2,X3] : (eqangle(X0,X3,X2,X3,X2,X3,X2,X1) | ~para(X0,X1,X2,X3) | ~cyclic(X0,X1,X2,X3))),
  inference(flattening,[],[f246])).
fof(f246,plain,(
  ! [X0,X1,X2,X3] : (eqangle(X0,X3,X2,X3,X2,X3,X2,X1) | (~para(X0,X1,X2,X3) | ~cyclic(X0,X1,X2,X3)))),
  inference(ennf_transformation,[],[f55])).
fof(f55,axiom,(
  ! [X0,X1,X2,X3] : ((para(X0,X1,X2,X3) & cyclic(X0,X1,X2,X3)) => eqangle(X0,X3,X2,X3,X2,X3,X2,X1))),
  file('./7_Book/00EE/10/E074-23.p',ruleD54)).
fof(f2412,plain,(
  ( ! [X30,X28,X26,X24,X23,X31,X29,X27,X25] : (eqangle(X26,X27,X28,X29,X30,X25,X31,X23) | ~eqangle(X26,X27,X28,X29,X30,X24,X31,X24)) )),
  inference(subsumption_resolution,[],[f2392,f1434])).
fof(f1434,plain,(
  ( ! [X6,X7,X5] : (midp(X5,X6,X7)) )),
  inference(subsumption_resolution,[],[f1433,f480])).
fof(f480,plain,(
  ( ! [X12,X7,X13,X11] : (circle(X11,X7,X12,X13)) )),
  inference(cnf_transformation,[],[f357])).
fof(f1433,plain,(
  ( ! [X6,X7,X5] : (midp(X5,X6,X7) | ~circle(X6,X6,X6,X7)) )),
  inference(subsumption_resolution,[],[f1424,f472])).
fof(f1424,plain,(
  ( ! [X6,X7,X5] : (midp(X5,X6,X7) | ~coll(X5,X6,X7) | ~circle(X6,X6,X6,X7)) )),
  inference(resolution,[],[f409,f479])).
fof(f409,plain,(
  ( ! [X4,X2,X0,X3,X1] : (~eqangle(X0,X1,X0,X2,X3,X1,X3,X4) | midp(X4,X1,X2) | ~coll(X4,X1,X2) | ~circle(X3,X0,X1,X2)) )),
  inference(cnf_transformation,[],[f241])).
fof(f241,plain,(
  ! [X0,X1,X2,X3,X4] : (midp(X4,X1,X2) | ~eqangle(X0,X1,X0,X2,X3,X1,X3,X4) | ~coll(X4,X1,X2) | ~circle(X3,X0,X1,X2))),
  inference(flattening,[],[f240])).
fof(f240,plain,(
  ! [X0,X1,X2,X3,X4] : (midp(X4,X1,X2) | (~eqangle(X0,X1,X0,X2,X3,X1,X3,X4) | ~coll(X4,X1,X2) | ~circle(X3,X0,X1,X2)))),
  inference(ennf_transformation,[],[f131])).
fof(f131,plain,(
  ! [X0,X1,X2,X3,X4] : ((eqangle(X0,X1,X0,X2,X3,X1,X3,X4) & coll(X4,X1,X2) & circle(X3,X0,X1,X2)) => midp(X4,X1,X2))),
  inference(rectify,[],[f52])).
fof(f52,axiom,(
  ! [X0,X1,X2,X7,X6] : ((eqangle(X0,X1,X0,X2,X7,X1,X7,X6) & coll(X6,X1,X2) & circle(X7,X0,X1,X2)) => midp(X6,X1,X2))),
  file('./7_Book/00EE/10/E074-23.p',ruleD51)).
fof(f2392,plain,(
  ( ! [X30,X28,X26,X24,X23,X31,X29,X27,X25] : (~midp(X23,X24,X25) | eqangle(X26,X27,X28,X29,X30,X25,X31,X23) | ~eqangle(X26,X27,X28,X29,X30,X24,X31,X24)) )),
  inference(resolution,[],[f854,f379])).
fof(f854,plain,(
  ( ! [X6,X4,X7,X5,X3] : (eqangle(X7,X4,X6,X4,X7,X5,X6,X3) | ~midp(X3,X4,X5)) )),
  inference(subsumption_resolution,[],[f847,f480])).
fof(f847,plain,(
  ( ! [X6,X4,X7,X5,X3] : (~midp(X3,X4,X5) | ~circle(X6,X7,X4,X5) | eqangle(X7,X4,X6,X4,X7,X5,X6,X3)) )),
  inference(resolution,[],[f408,f378])).
fof(f408,plain,(
  ( ! [X4,X2,X0,X3,X1] : (eqangle(X0,X1,X0,X2,X3,X1,X3,X4) | ~midp(X4,X1,X2) | ~circle(X3,X0,X1,X2)) )),
  inference(cnf_transformation,[],[f239])).
fof(f239,plain,(
  ! [X0,X1,X2,X3,X4] : (eqangle(X0,X1,X0,X2,X3,X1,X3,X4) | ~midp(X4,X1,X2) | ~circle(X3,X0,X1,X2))),
  inference(flattening,[],[f238])).
fof(f238,plain,(
  ! [X0,X1,X2,X3,X4] : (eqangle(X0,X1,X0,X2,X3,X1,X3,X4) | (~midp(X4,X1,X2) | ~circle(X3,X0,X1,X2)))),
  inference(ennf_transformation,[],[f130])).
fof(f130,plain,(
  ! [X0,X1,X2,X3,X4] : ((midp(X4,X1,X2) & circle(X3,X0,X1,X2)) => eqangle(X0,X1,X0,X2,X3,X1,X3,X4))),
  inference(rectify,[],[f51])).
fof(f51,axiom,(
  ! [X0,X1,X2,X7,X6] : ((midp(X6,X1,X2) & circle(X7,X0,X1,X2)) => eqangle(X0,X1,X0,X2,X7,X1,X7,X6))),
  file('./7_Book/00EE/10/E074-23.p',ruleD50)).
% SZS output end Proof for E074-23
% ------------------------------
% Version: Vampire 4.4.0 (commit 7916d27 on 2019-08-23 08:50:16 +0100)
% Termination reason: Refutation

% Memory used [KB]: 5884
% Time elapsed: 0.047 s
% ------------------------------
% ------------------------------
