% Refutation found. Thanks to Tanya!
% SZS status Theorem for E051-32
% SZS output start Proof for E051-32
fof(f24153,plain,(
  $false),
  inference(avatar_sat_refutation,[],[f946,f971,f1443,f1765,f1766,f1822,f2003,f2451,f2920,f3181,f7394,f13195,f13480,f13489,f24136])).
fof(f24136,plain,(
  ~spl26_62 | ~spl26_134 | ~spl26_136),
  inference(avatar_contradiction_clause,[],[f24129])).
fof(f24129,plain,(
  $false | (~spl26_62 | ~spl26_134 | ~spl26_136)),
  inference(resolution,[],[f18111,f477])).
fof(f477,plain,(
  ~para(sK24,sK23,sK22,sK25)),
  inference(cnf_transformation,[],[f357])).
fof(f357,plain,(
  ~para(sK24,sK23,sK22,sK25) & eqangle(sK21,sK22,sK22,sK25,sK22,sK25,sK25,sK21) & eqangle(sK21,sK22,sK22,sK25,sK22,sK21,sK21,sK25) & eqangle(sK20,sK22,sK22,sK24,sK22,sK24,sK24,sK20) & eqangle(sK20,sK22,sK22,sK24,sK22,sK20,sK20,sK24) & eqangle(sK21,sK20,sK20,sK23,sK20,sK23,sK23,sK21) & eqangle(sK21,sK20,sK20,sK23,sK20,sK21,sK21,sK23)),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK20,sK21,sK22,sK23,sK24,sK25])],[f319,f356])).
fof(f356,plain,(
  ? [X0,X1,X2,X3,X4,X5] : (~para(X4,X3,X2,X5) & eqangle(X1,X2,X2,X5,X2,X5,X5,X1) & eqangle(X1,X2,X2,X5,X2,X1,X1,X5) & eqangle(X0,X2,X2,X4,X2,X4,X4,X0) & eqangle(X0,X2,X2,X4,X2,X0,X0,X4) & eqangle(X1,X0,X0,X3,X0,X3,X3,X1) & eqangle(X1,X0,X0,X3,X0,X1,X1,X3)) => (~para(sK24,sK23,sK22,sK25) & eqangle(sK21,sK22,sK22,sK25,sK22,sK25,sK25,sK21) & eqangle(sK21,sK22,sK22,sK25,sK22,sK21,sK21,sK25) & eqangle(sK20,sK22,sK22,sK24,sK22,sK24,sK24,sK20) & eqangle(sK20,sK22,sK22,sK24,sK22,sK20,sK20,sK24) & eqangle(sK21,sK20,sK20,sK23,sK20,sK23,sK23,sK21) & eqangle(sK21,sK20,sK20,sK23,sK20,sK21,sK21,sK23))),
  introduced(choice_axiom,[])).
fof(f319,plain,(
  ? [X0,X1,X2,X3,X4,X5] : (~para(X4,X3,X2,X5) & eqangle(X1,X2,X2,X5,X2,X5,X5,X1) & eqangle(X1,X2,X2,X5,X2,X1,X1,X5) & eqangle(X0,X2,X2,X4,X2,X4,X4,X0) & eqangle(X0,X2,X2,X4,X2,X0,X0,X4) & eqangle(X1,X0,X0,X3,X0,X3,X3,X1) & eqangle(X1,X0,X0,X3,X0,X1,X1,X3))),
  inference(flattening,[],[f318])).
fof(f318,plain,(
  ? [X0,X1,X2,X3,X4,X5] : (~para(X4,X3,X2,X5) & (eqangle(X1,X2,X2,X5,X2,X5,X5,X1) & eqangle(X1,X2,X2,X5,X2,X1,X1,X5) & eqangle(X0,X2,X2,X4,X2,X4,X4,X0) & eqangle(X0,X2,X2,X4,X2,X0,X0,X4) & eqangle(X1,X0,X0,X3,X0,X3,X3,X1) & eqangle(X1,X0,X0,X3,X0,X1,X1,X3)))),
  inference(ennf_transformation,[],[f96])).
fof(f96,negated_conjecture,(
  ~! [X0,X1,X2,X3,X4,X5] : ((eqangle(X1,X2,X2,X5,X2,X5,X5,X1) & eqangle(X1,X2,X2,X5,X2,X1,X1,X5) & eqangle(X0,X2,X2,X4,X2,X4,X4,X0) & eqangle(X0,X2,X2,X4,X2,X0,X0,X4) & eqangle(X1,X0,X0,X3,X0,X3,X3,X1) & eqangle(X1,X0,X0,X3,X0,X1,X1,X3)) => para(X4,X3,X2,X5))),
  inference(negated_conjecture,[],[f95])).
fof(f95,conjecture,(
  ! [X0,X1,X2,X3,X4,X5] : ((eqangle(X1,X2,X2,X5,X2,X5,X5,X1) & eqangle(X1,X2,X2,X5,X2,X1,X1,X5) & eqangle(X0,X2,X2,X4,X2,X4,X4,X0) & eqangle(X0,X2,X2,X4,X2,X0,X0,X4) & eqangle(X1,X0,X0,X3,X0,X3,X3,X1) & eqangle(X1,X0,X0,X3,X0,X1,X1,X3)) => para(X4,X3,X2,X5))),
  file('./7_Book/00EE/06/E051-32.p',exemplo7Book00EE06E05132)).
fof(f18111,plain,(
  ( ! [X28,X26,X27,X25] : (para(X25,X26,X27,X28)) ) | (~spl26_62 | ~spl26_134 | ~spl26_136)),
  inference(subsumption_resolution,[],[f18095,f18094])).
fof(f18094,plain,(
  ( ! [X24,X23] : (perp(X23,X24,sK20,sK21)) ) | (~spl26_62 | ~spl26_134 | ~spl26_136)),
  inference(resolution,[],[f14095,f365])).
fof(f365,plain,(
  ( ! [X2,X0,X3,X1] : (~perp(X0,X1,X2,X3) | perp(X2,X3,X0,X1)) )),
  inference(cnf_transformation,[],[f176])).
fof(f176,plain,(
  ! [X0,X1,X2,X3] : (perp(X2,X3,X0,X1) | ~perp(X0,X1,X2,X3))),
  inference(ennf_transformation,[],[f8])).
fof(f8,axiom,(
  ! [X0,X1,X2,X3] : (perp(X0,X1,X2,X3) => perp(X2,X3,X0,X1))),
  file('./7_Book/00EE/06/E051-32.p',ruleD8)).
fof(f14095,plain,(
  ( ! [X12,X13] : (perp(sK20,sK21,X12,X13)) ) | (~spl26_62 | ~spl26_134 | ~spl26_136)),
  inference(subsumption_resolution,[],[f14092,f13704])).
fof(f13704,plain,(
  ( ! [X6] : (cong(sK20,X6,sK21,X6)) ) | (~spl26_62 | ~spl26_134 | ~spl26_136)),
  inference(subsumption_resolution,[],[f13697,f8830])).
fof(f8830,plain,(
  ( ! [X45] : (midp(X45,sK20,sK23)) ) | (~spl26_134 | ~spl26_136)),
  inference(resolution,[],[f8048,f1764])).
fof(f1764,plain,(
  ( ! [X0] : (~midp(X0,sK21,sK21) | midp(X0,sK20,sK23)) ) | ~spl26_134),
  inference(avatar_component_clause,[],[f1763])).
fof(f1763,plain,(
  spl26_134 <=> ! [X0] : (midp(X0,sK20,sK23) | ~midp(X0,sK21,sK21))),
  introduced(avatar_definition,[new_symbols(naming,[spl26_134])])).
fof(f8048,plain,(
  ( ! [X26,X25] : (midp(X25,X26,X26)) ) | ~spl26_136),
  inference(resolution,[],[f7392,f3211])).
fof(f3211,plain,(
  ( ! [X10,X8,X9] : (~midp(X8,X10,X10) | midp(X8,X9,X9)) ) | ~spl26_136),
  inference(subsumption_resolution,[],[f3205,f1818])).
fof(f1818,plain,(
  ( ! [X2,X3] : (para(X2,X3,X2,X3)) ) | ~spl26_136),
  inference(avatar_component_clause,[],[f1817])).
fof(f1817,plain,(
  spl26_136 <=> ! [X3,X2] : para(X2,X3,X2,X3)),
  introduced(avatar_definition,[new_symbols(naming,[spl26_136])])).
fof(f3205,plain,(
  ( ! [X10,X8,X9] : (midp(X8,X9,X9) | ~para(X10,X9,X10,X9) | ~midp(X8,X10,X10)) ) | ~spl26_136),
  inference(resolution,[],[f1818,f422])).
fof(f422,plain,(
  ( ! [X4,X2,X0,X3,X1] : (~para(X0,X3,X1,X2) | midp(X4,X2,X3) | ~para(X0,X2,X1,X3) | ~midp(X4,X0,X1)) )),
  inference(cnf_transformation,[],[f263])).
fof(f263,plain,(
  ! [X0,X1,X2,X3,X4] : (midp(X4,X2,X3) | ~para(X0,X3,X1,X2) | ~para(X0,X2,X1,X3) | ~midp(X4,X0,X1))),
  inference(flattening,[],[f262])).
fof(f262,plain,(
  ! [X0,X1,X2,X3,X4] : (midp(X4,X2,X3) | (~para(X0,X3,X1,X2) | ~para(X0,X2,X1,X3) | ~midp(X4,X0,X1)))),
  inference(ennf_transformation,[],[f143])).
fof(f143,plain,(
  ! [X0,X1,X2,X3,X4] : ((para(X0,X3,X1,X2) & para(X0,X2,X1,X3) & midp(X4,X0,X1)) => midp(X4,X2,X3))),
  inference(rectify,[],[f65])).
fof(f65,axiom,(
  ! [X0,X1,X2,X3,X6] : ((para(X0,X3,X1,X2) & para(X0,X2,X1,X3) & midp(X6,X0,X1)) => midp(X6,X2,X3))),
  file('./7_Book/00EE/06/E051-32.p',ruleD64)).
fof(f7392,plain,(
  ( ! [X14] : (midp(X14,X14,X14)) ) | ~spl26_136),
  inference(subsumption_resolution,[],[f7371,f3237])).
fof(f3237,plain,(
  ( ! [X4,X3] : (coll(X3,X4,X3)) ) | ~spl26_136),
  inference(resolution,[],[f3202,f359])).
fof(f359,plain,(
  ( ! [X2,X0,X1] : (~coll(X0,X1,X2) | coll(X1,X0,X2)) )),
  inference(cnf_transformation,[],[f168])).
fof(f168,plain,(
  ! [X0,X1,X2] : (coll(X1,X0,X2) | ~coll(X0,X1,X2))),
  inference(ennf_transformation,[],[f2])).
fof(f2,axiom,(
  ! [X0,X1,X2] : (coll(X0,X1,X2) => coll(X1,X0,X2))),
  file('./7_Book/00EE/06/E051-32.p',ruleD2)).
fof(f3202,plain,(
  ( ! [X0,X1] : (coll(X0,X1,X1)) ) | ~spl26_136),
  inference(resolution,[],[f1818,f424])).
fof(f424,plain,(
  ( ! [X2,X0,X1] : (~para(X0,X1,X0,X2) | coll(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f266])).
fof(f266,plain,(
  ! [X0,X1,X2] : (coll(X0,X1,X2) | ~para(X0,X1,X0,X2))),
  inference(ennf_transformation,[],[f67])).
fof(f67,axiom,(
  ! [X0,X1,X2] : (para(X0,X1,X0,X2) => coll(X0,X1,X2))),
  file('./7_Book/00EE/06/E051-32.p',ruleD66)).
fof(f7371,plain,(
  ( ! [X14] : (~coll(X14,X14,X14) | midp(X14,X14,X14)) ) | ~spl26_136),
  inference(resolution,[],[f7213,f425])).
fof(f425,plain,(
  ( ! [X2,X0,X1] : (~cong(X0,X1,X0,X2) | ~coll(X0,X1,X2) | midp(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f268])).
fof(f268,plain,(
  ! [X0,X1,X2] : (midp(X0,X1,X2) | ~coll(X0,X1,X2) | ~cong(X0,X1,X0,X2))),
  inference(flattening,[],[f267])).
fof(f267,plain,(
  ! [X0,X1,X2] : (midp(X0,X1,X2) | (~coll(X0,X1,X2) | ~cong(X0,X1,X0,X2)))),
  inference(ennf_transformation,[],[f68])).
fof(f68,axiom,(
  ! [X0,X1,X2] : ((coll(X0,X1,X2) & cong(X0,X1,X0,X2)) => midp(X0,X1,X2))),
  file('./7_Book/00EE/06/E051-32.p',ruleD67)).
fof(f7213,plain,(
  ( ! [X35] : (cong(X35,X35,X35,X35)) ) | ~spl26_136),
  inference(subsumption_resolution,[],[f7209,f7177])).
fof(f7177,plain,(
  ( ! [X10,X8,X9] : (cyclic(X8,X8,X9,X10)) ) | ~spl26_136),
  inference(resolution,[],[f2993,f1818])).
fof(f2993,plain,(
  ( ! [X68,X66,X67] : (~para(X66,X67,X66,X67) | cyclic(X67,X67,X66,X68)) )),
  inference(subsumption_resolution,[],[f1986,f1987])).
fof(f1987,plain,(
  ( ! [X70,X71,X69] : (~para(X69,X70,X69,X70) | coll(X69,X71,X70) | cyclic(X70,X70,X69,X71)) )),
  inference(resolution,[],[f649,f399])).
fof(f399,plain,(
  ( ! [X2,X0,X3,X1] : (~eqangle(X2,X0,X2,X1,X3,X0,X3,X1) | coll(X2,X3,X0) | cyclic(X0,X1,X2,X3)) )),
  inference(cnf_transformation,[],[f221])).
fof(f221,plain,(
  ! [X0,X1,X2,X3] : (cyclic(X0,X1,X2,X3) | coll(X2,X3,X0) | ~eqangle(X2,X0,X2,X1,X3,X0,X3,X1))),
  inference(flattening,[],[f220])).
fof(f220,plain,(
  ! [X0,X1,X2,X3] : (cyclic(X0,X1,X2,X3) | (coll(X2,X3,X0) | ~eqangle(X2,X0,X2,X1,X3,X0,X3,X1)))),
  inference(ennf_transformation,[],[f121])).
fof(f121,plain,(
  ! [X0,X1,X2,X3] : ((~coll(X2,X3,X0) & eqangle(X2,X0,X2,X1,X3,X0,X3,X1)) => cyclic(X0,X1,X2,X3))),
  inference(rectify,[],[f42])).
fof(f42,axiom,(
  ! [X0,X1,X8,X9] : ((~coll(X8,X9,X0) & eqangle(X8,X0,X8,X1,X9,X0,X9,X1)) => cyclic(X0,X1,X8,X9))),
  file('./7_Book/00EE/06/E051-32.p',ruleD42a)).
fof(f649,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (eqangle(X0,X1,X2,X3,X4,X5,X4,X5) | ~para(X0,X1,X2,X3)) )),
  inference(resolution,[],[f378,f397])).
fof(f397,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (eqangle(X0,X1,X4,X5,X2,X3,X4,X5) | ~para(X0,X1,X2,X3)) )),
  inference(cnf_transformation,[],[f218])).
fof(f218,plain,(
  ! [X0,X1,X2,X3,X4,X5] : (eqangle(X0,X1,X4,X5,X2,X3,X4,X5) | ~para(X0,X1,X2,X3))),
  inference(ennf_transformation,[],[f119])).
fof(f119,plain,(
  ! [X0,X1,X2,X3,X4,X5] : (para(X0,X1,X2,X3) => eqangle(X0,X1,X4,X5,X2,X3,X4,X5))),
  inference(rectify,[],[f40])).
fof(f40,axiom,(
  ! [X0,X1,X2,X3,X8,X9] : (para(X0,X1,X2,X3) => eqangle(X0,X1,X8,X9,X2,X3,X8,X9))),
  file('./7_Book/00EE/06/E051-32.p',ruleD40)).
fof(f378,plain,(
  ( ! [X6,X4,X2,X0,X7,X5,X3,X1] : (~eqangle(X0,X1,X2,X3,X4,X5,X6,X7) | eqangle(X0,X1,X4,X5,X2,X3,X6,X7)) )),
  inference(cnf_transformation,[],[f194])).
fof(f194,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6,X7] : (eqangle(X0,X1,X4,X5,X2,X3,X6,X7) | ~eqangle(X0,X1,X2,X3,X4,X5,X6,X7))),
  inference(ennf_transformation,[],[f103])).
fof(f103,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6,X7] : (eqangle(X0,X1,X2,X3,X4,X5,X6,X7) => eqangle(X0,X1,X4,X5,X2,X3,X6,X7))),
  inference(rectify,[],[f21])).
fof(f21,axiom,(
  ! [X0,X1,X2,X3,X8,X9,X10,X11] : (eqangle(X0,X1,X2,X3,X8,X9,X10,X11) => eqangle(X0,X1,X8,X9,X2,X3,X10,X11))),
  file('./7_Book/00EE/06/E051-32.p',ruleD21)).
fof(f1986,plain,(
  ( ! [X68,X66,X67] : (~para(X66,X67,X66,X67) | ~coll(X66,X68,X67) | cyclic(X67,X67,X66,X68)) )),
  inference(resolution,[],[f649,f400])).
fof(f400,plain,(
  ( ! [X2,X0,X3,X1] : (~eqangle(X2,X0,X2,X1,X3,X0,X3,X1) | ~coll(X2,X3,X1) | cyclic(X0,X1,X2,X3)) )),
  inference(cnf_transformation,[],[f223])).
fof(f223,plain,(
  ! [X0,X1,X2,X3] : (cyclic(X0,X1,X2,X3) | ~coll(X2,X3,X1) | ~eqangle(X2,X0,X2,X1,X3,X0,X3,X1))),
  inference(flattening,[],[f222])).
fof(f222,plain,(
  ! [X0,X1,X2,X3] : (cyclic(X0,X1,X2,X3) | (~coll(X2,X3,X1) | ~eqangle(X2,X0,X2,X1,X3,X0,X3,X1)))),
  inference(ennf_transformation,[],[f122])).
fof(f122,plain,(
  ! [X0,X1,X2,X3] : ((coll(X2,X3,X1) & eqangle(X2,X0,X2,X1,X3,X0,X3,X1)) => cyclic(X0,X1,X2,X3))),
  inference(rectify,[],[f43])).
fof(f43,axiom,(
  ! [X0,X1,X8,X9] : ((coll(X8,X9,X1) & eqangle(X8,X0,X8,X1,X9,X0,X9,X1)) => cyclic(X0,X1,X8,X9))),
  file('./7_Book/00EE/06/E051-32.p',ruleD42b)).
fof(f7209,plain,(
  ( ! [X35,X36] : (cong(X35,X35,X35,X35) | ~cyclic(X35,X35,X36,X35)) ) | ~spl26_136),
  inference(duplicate_literal_removal,[],[f7196])).
fof(f7196,plain,(
  ( ! [X35,X36] : (cong(X35,X35,X35,X35) | ~cyclic(X35,X35,X36,X35) | ~cyclic(X35,X35,X36,X35)) ) | ~spl26_136),
  inference(resolution,[],[f7177,f1597])).
fof(f1597,plain,(
  ( ! [X10,X8,X7,X9] : (~cyclic(X7,X8,X9,X10) | cong(X7,X8,X7,X8) | ~cyclic(X7,X8,X9,X8) | ~cyclic(X7,X8,X9,X7)) )),
  inference(duplicate_literal_removal,[],[f1591])).
fof(f1591,plain,(
  ( ! [X10,X8,X7,X9] : (cong(X7,X8,X7,X8) | ~cyclic(X7,X8,X9,X10) | ~cyclic(X7,X8,X9,X8) | ~cyclic(X7,X8,X9,X7) | ~cyclic(X7,X8,X9,X10)) )),
  inference(resolution,[],[f401,f398])).
fof(f398,plain,(
  ( ! [X2,X0,X3,X1] : (eqangle(X2,X0,X2,X1,X3,X0,X3,X1) | ~cyclic(X0,X1,X2,X3)) )),
  inference(cnf_transformation,[],[f219])).
fof(f219,plain,(
  ! [X0,X1,X2,X3] : (eqangle(X2,X0,X2,X1,X3,X0,X3,X1) | ~cyclic(X0,X1,X2,X3))),
  inference(ennf_transformation,[],[f120])).
fof(f120,plain,(
  ! [X0,X1,X2,X3] : (cyclic(X0,X1,X2,X3) => eqangle(X2,X0,X2,X1,X3,X0,X3,X1))),
  inference(rectify,[],[f41])).
fof(f41,axiom,(
  ! [X0,X1,X8,X9] : (cyclic(X0,X1,X8,X9) => eqangle(X8,X0,X8,X1,X9,X0,X9,X1))),
  file('./7_Book/00EE/06/E051-32.p',ruleD41)).
fof(f401,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (~eqangle(X2,X0,X2,X1,X5,X3,X5,X4) | cong(X0,X1,X3,X4) | ~cyclic(X0,X1,X2,X5) | ~cyclic(X0,X1,X2,X4) | ~cyclic(X0,X1,X2,X3)) )),
  inference(cnf_transformation,[],[f225])).
fof(f225,plain,(
  ! [X0,X1,X2,X3,X4,X5] : (cong(X0,X1,X3,X4) | ~eqangle(X2,X0,X2,X1,X5,X3,X5,X4) | ~cyclic(X0,X1,X2,X5) | ~cyclic(X0,X1,X2,X4) | ~cyclic(X0,X1,X2,X3))),
  inference(flattening,[],[f224])).
fof(f224,plain,(
  ! [X0,X1,X2,X3,X4,X5] : (cong(X0,X1,X3,X4) | (~eqangle(X2,X0,X2,X1,X5,X3,X5,X4) | ~cyclic(X0,X1,X2,X5) | ~cyclic(X0,X1,X2,X4) | ~cyclic(X0,X1,X2,X3)))),
  inference(ennf_transformation,[],[f123])).
fof(f123,plain,(
  ! [X0,X1,X2,X3,X4,X5] : ((eqangle(X2,X0,X2,X1,X5,X3,X5,X4) & cyclic(X0,X1,X2,X5) & cyclic(X0,X1,X2,X4) & cyclic(X0,X1,X2,X3)) => cong(X0,X1,X3,X4))),
  inference(rectify,[],[f44])).
fof(f44,axiom,(
  ! [X0,X1,X2,X8,X9,X14] : ((eqangle(X2,X0,X2,X1,X14,X8,X14,X9) & cyclic(X0,X1,X2,X14) & cyclic(X0,X1,X2,X9) & cyclic(X0,X1,X2,X8)) => cong(X0,X1,X8,X9))),
  file('./7_Book/00EE/06/E051-32.p',ruleD43)).
fof(f13697,plain,(
  ( ! [X6] : (~midp(X6,sK20,sK23) | cong(sK20,X6,sK21,X6)) ) | ~spl26_62),
  inference(resolution,[],[f1141,f410])).
fof(f410,plain,(
  ( ! [X2,X0,X3,X1] : (~perp(X0,X1,X1,X2) | ~midp(X3,X0,X2) | cong(X0,X3,X1,X3)) )),
  inference(cnf_transformation,[],[f242])).
fof(f242,plain,(
  ! [X0,X1,X2,X3] : (cong(X0,X3,X1,X3) | ~midp(X3,X0,X2) | ~perp(X0,X1,X1,X2))),
  inference(flattening,[],[f241])).
fof(f241,plain,(
  ! [X0,X1,X2,X3] : (cong(X0,X3,X1,X3) | (~midp(X3,X0,X2) | ~perp(X0,X1,X1,X2)))),
  inference(ennf_transformation,[],[f132])).
fof(f132,plain,(
  ! [X0,X1,X2,X3] : ((midp(X3,X0,X2) & perp(X0,X1,X1,X2)) => cong(X0,X3,X1,X3))),
  inference(rectify,[],[f53])).
fof(f53,axiom,(
  ! [X0,X1,X2,X6] : ((midp(X6,X0,X2) & perp(X0,X1,X1,X2)) => cong(X0,X6,X1,X6))),
  file('./7_Book/00EE/06/E051-32.p',ruleD52)).
fof(f1141,plain,(
  perp(sK20,sK21,sK21,sK23) | ~spl26_62),
  inference(avatar_component_clause,[],[f1140])).
fof(f1140,plain,(
  spl26_62 <=> perp(sK20,sK21,sK21,sK23)),
  introduced(avatar_definition,[new_symbols(naming,[spl26_62])])).
fof(f14092,plain,(
  ( ! [X12,X13] : (perp(sK20,sK21,X12,X13) | ~cong(sK20,X12,sK21,X12)) ) | (~spl26_62 | ~spl26_134 | ~spl26_136)),
  inference(resolution,[],[f13704,f414])).
fof(f414,plain,(
  ( ! [X2,X0,X3,X1] : (~cong(X0,X3,X1,X3) | perp(X0,X1,X2,X3) | ~cong(X0,X2,X1,X2)) )),
  inference(cnf_transformation,[],[f250])).
fof(f250,plain,(
  ! [X0,X1,X2,X3] : (perp(X0,X1,X2,X3) | ~cong(X0,X3,X1,X3) | ~cong(X0,X2,X1,X2))),
  inference(flattening,[],[f249])).
fof(f249,plain,(
  ! [X0,X1,X2,X3] : (perp(X0,X1,X2,X3) | (~cong(X0,X3,X1,X3) | ~cong(X0,X2,X1,X2)))),
  inference(ennf_transformation,[],[f135])).
fof(f135,plain,(
  ! [X0,X1,X2,X3] : ((cong(X0,X3,X1,X3) & cong(X0,X2,X1,X2)) => perp(X0,X1,X2,X3))),
  inference(rectify,[],[f57])).
fof(f57,axiom,(
  ! [X0,X1,X8,X9] : ((cong(X0,X9,X1,X9) & cong(X0,X8,X1,X8)) => perp(X0,X1,X8,X9))),
  file('./7_Book/00EE/06/E051-32.p',ruleD56)).
fof(f18095,plain,(
  ( ! [X28,X26,X27,X25] : (para(X25,X26,X27,X28) | ~perp(X25,X26,sK20,sK21)) ) | (~spl26_62 | ~spl26_134 | ~spl26_136)),
  inference(resolution,[],[f14095,f366])).
fof(f366,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (~perp(X2,X3,X4,X5) | para(X0,X1,X4,X5) | ~perp(X0,X1,X2,X3)) )),
  inference(cnf_transformation,[],[f178])).
fof(f178,plain,(
  ! [X0,X1,X2,X3,X4,X5] : (para(X0,X1,X4,X5) | ~perp(X2,X3,X4,X5) | ~perp(X0,X1,X2,X3))),
  inference(flattening,[],[f177])).
fof(f177,plain,(
  ! [X0,X1,X2,X3,X4,X5] : (para(X0,X1,X4,X5) | (~perp(X2,X3,X4,X5) | ~perp(X0,X1,X2,X3)))),
  inference(ennf_transformation,[],[f9])).
fof(f9,axiom,(
  ! [X0,X1,X2,X3,X4,X5] : ((perp(X2,X3,X4,X5) & perp(X0,X1,X2,X3)) => para(X0,X1,X4,X5))),
  file('./7_Book/00EE/06/E051-32.p',ruleD9)).
fof(f13489,plain,(
  spl26_62 | ~spl26_65),
  inference(avatar_split_clause,[],[f6573,f1150,f1140])).
fof(f1150,plain,(
  spl26_65 <=> ~perp(sK20,sK23,sK23,sK21)),
  introduced(avatar_definition,[new_symbols(naming,[spl26_65])])).
fof(f6573,plain,(
  ~perp(sK20,sK23,sK23,sK21) | perp(sK20,sK21,sK21,sK23)),
  inference(resolution,[],[f2625,f432])).
fof(f432,plain,(
  ( ! [X6,X4,X2,X0,X7,X5,X3,X1] : (~eqangle(X0,X1,X2,X3,X4,X5,X6,X7) | ~perp(X4,X5,X6,X7) | perp(X0,X1,X2,X3)) )),
  inference(cnf_transformation,[],[f280])).
fof(f280,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6,X7] : (perp(X0,X1,X2,X3) | ~perp(X4,X5,X6,X7) | ~eqangle(X0,X1,X2,X3,X4,X5,X6,X7))),
  inference(flattening,[],[f279])).
fof(f279,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6,X7] : (perp(X0,X1,X2,X3) | (~perp(X4,X5,X6,X7) | ~eqangle(X0,X1,X2,X3,X4,X5,X6,X7)))),
  inference(ennf_transformation,[],[f147])).
fof(f147,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6,X7] : ((perp(X4,X5,X6,X7) & eqangle(X0,X1,X2,X3,X4,X5,X6,X7)) => perp(X0,X1,X2,X3))),
  inference(rectify,[],[f75])).
fof(f75,axiom,(
  ! [X0,X1,X2,X3,X8,X9,X10,X11] : ((perp(X8,X9,X10,X11) & eqangle(X0,X1,X2,X3,X8,X9,X10,X11)) => perp(X0,X1,X2,X3))),
  file('./7_Book/00EE/06/E051-32.p',ruleD74)).
fof(f2625,plain,(
  eqangle(sK20,sK21,sK21,sK23,sK20,sK23,sK23,sK21)),
  inference(resolution,[],[f1544,f634])).
fof(f634,plain,(
  eqangle(sK20,sK21,sK21,sK23,sK21,sK20,sK20,sK23)),
  inference(resolution,[],[f377,f471])).
fof(f471,plain,(
  eqangle(sK21,sK20,sK20,sK23,sK20,sK21,sK21,sK23)),
  inference(cnf_transformation,[],[f357])).
fof(f377,plain,(
  ( ! [X6,X4,X2,X0,X7,X5,X3,X1] : (~eqangle(X0,X1,X2,X3,X4,X5,X6,X7) | eqangle(X4,X5,X6,X7,X0,X1,X2,X3)) )),
  inference(cnf_transformation,[],[f193])).
fof(f193,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6,X7] : (eqangle(X4,X5,X6,X7,X0,X1,X2,X3) | ~eqangle(X0,X1,X2,X3,X4,X5,X6,X7))),
  inference(ennf_transformation,[],[f102])).
fof(f102,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6,X7] : (eqangle(X0,X1,X2,X3,X4,X5,X6,X7) => eqangle(X4,X5,X6,X7,X0,X1,X2,X3))),
  inference(rectify,[],[f20])).
fof(f20,axiom,(
  ! [X0,X1,X2,X3,X8,X9,X10,X11] : (eqangle(X0,X1,X2,X3,X8,X9,X10,X11) => eqangle(X8,X9,X10,X11,X0,X1,X2,X3))),
  file('./7_Book/00EE/06/E051-32.p',ruleD20)).
fof(f1544,plain,(
  ( ! [X99,X97,X98,X96] : (~eqangle(X96,X97,X98,X99,sK21,sK20,sK20,sK23) | eqangle(X96,X97,X98,X99,sK20,sK23,sK23,sK21)) )),
  inference(resolution,[],[f379,f472])).
fof(f472,plain,(
  eqangle(sK21,sK20,sK20,sK23,sK20,sK23,sK23,sK21)),
  inference(cnf_transformation,[],[f357])).
fof(f379,plain,(
  ( ! [X6,X4,X2,X0,X10,X8,X7,X5,X3,X1,X11,X9] : (~eqangle(X4,X5,X6,X7,X8,X9,X10,X11) | eqangle(X0,X1,X2,X3,X8,X9,X10,X11) | ~eqangle(X0,X1,X2,X3,X4,X5,X6,X7)) )),
  inference(cnf_transformation,[],[f196])).
fof(f196,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11] : (eqangle(X0,X1,X2,X3,X8,X9,X10,X11) | ~eqangle(X4,X5,X6,X7,X8,X9,X10,X11) | ~eqangle(X0,X1,X2,X3,X4,X5,X6,X7))),
  inference(flattening,[],[f195])).
fof(f195,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11] : (eqangle(X0,X1,X2,X3,X8,X9,X10,X11) | (~eqangle(X4,X5,X6,X7,X8,X9,X10,X11) | ~eqangle(X0,X1,X2,X3,X4,X5,X6,X7)))),
  inference(ennf_transformation,[],[f104])).
fof(f104,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11] : ((eqangle(X4,X5,X6,X7,X8,X9,X10,X11) & eqangle(X0,X1,X2,X3,X4,X5,X6,X7)) => eqangle(X0,X1,X2,X3,X8,X9,X10,X11))),
  inference(rectify,[],[f22])).
fof(f22,axiom,(
  ! [X0,X1,X2,X3,X8,X9,X10,X11,X4,X5,X12,X13] : ((eqangle(X8,X9,X10,X11,X4,X5,X12,X13) & eqangle(X0,X1,X2,X3,X8,X9,X10,X11)) => eqangle(X0,X1,X2,X3,X4,X5,X12,X13))),
  file('./7_Book/00EE/06/E051-32.p',ruleD22)).
fof(f13480,plain,(
  spl26_64 | ~spl26_142),
  inference(avatar_split_clause,[],[f13248,f2001,f1147])).
fof(f1147,plain,(
  spl26_64 <=> perp(sK20,sK23,sK23,sK21)),
  introduced(avatar_definition,[new_symbols(naming,[spl26_64])])).
fof(f2001,plain,(
  spl26_142 <=> ! [X44,X46,X45,X47] : (~para(X44,X45,X46,X47) | perp(X44,X45,X46,X47))),
  introduced(avatar_definition,[new_symbols(naming,[spl26_142])])).
fof(f13248,plain,(
  perp(sK20,sK23,sK23,sK21) | ~spl26_142),
  inference(resolution,[],[f2002,f837])).
fof(f837,plain,(
  para(sK20,sK23,sK23,sK21)),
  inference(resolution,[],[f819,f361])).
fof(f361,plain,(
  ( ! [X2,X0,X3,X1] : (~para(X0,X1,X2,X3) | para(X0,X1,X3,X2)) )),
  inference(cnf_transformation,[],[f171])).
fof(f171,plain,(
  ! [X0,X1,X2,X3] : (para(X0,X1,X3,X2) | ~para(X0,X1,X2,X3))),
  inference(ennf_transformation,[],[f4])).
fof(f4,axiom,(
  ! [X0,X1,X2,X3] : (para(X0,X1,X2,X3) => para(X0,X1,X3,X2))),
  file('./7_Book/00EE/06/E051-32.p',ruleD4)).
fof(f819,plain,(
  para(sK20,sK23,sK21,sK23)),
  inference(resolution,[],[f627,f396])).
fof(f396,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (~eqangle(X0,X1,X4,X5,X2,X3,X4,X5) | para(X0,X1,X2,X3)) )),
  inference(cnf_transformation,[],[f217])).
fof(f217,plain,(
  ! [X0,X1,X2,X3,X4,X5] : (para(X0,X1,X2,X3) | ~eqangle(X0,X1,X4,X5,X2,X3,X4,X5))),
  inference(ennf_transformation,[],[f118])).
fof(f118,plain,(
  ! [X0,X1,X2,X3,X4,X5] : (eqangle(X0,X1,X4,X5,X2,X3,X4,X5) => para(X0,X1,X2,X3))),
  inference(rectify,[],[f39])).
fof(f39,axiom,(
  ! [X0,X1,X2,X3,X8,X9] : (eqangle(X0,X1,X8,X9,X2,X3,X8,X9) => para(X0,X1,X2,X3))),
  file('./7_Book/00EE/06/E051-32.p',ruleD39)).
fof(f627,plain,(
  eqangle(sK20,sK23,sK20,sK21,sK21,sK23,sK20,sK21)),
  inference(resolution,[],[f376,f610])).
fof(f610,plain,(
  eqangle(sK20,sK21,sK20,sK23,sK20,sK21,sK21,sK23)),
  inference(resolution,[],[f375,f471])).
fof(f375,plain,(
  ( ! [X6,X4,X2,X0,X7,X5,X3,X1] : (~eqangle(X0,X1,X2,X3,X4,X5,X6,X7) | eqangle(X1,X0,X2,X3,X4,X5,X6,X7)) )),
  inference(cnf_transformation,[],[f191])).
fof(f191,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6,X7] : (eqangle(X1,X0,X2,X3,X4,X5,X6,X7) | ~eqangle(X0,X1,X2,X3,X4,X5,X6,X7))),
  inference(ennf_transformation,[],[f100])).
fof(f100,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6,X7] : (eqangle(X0,X1,X2,X3,X4,X5,X6,X7) => eqangle(X1,X0,X2,X3,X4,X5,X6,X7))),
  inference(rectify,[],[f18])).
fof(f18,axiom,(
  ! [X0,X1,X2,X3,X8,X9,X10,X11] : (eqangle(X0,X1,X2,X3,X8,X9,X10,X11) => eqangle(X1,X0,X2,X3,X8,X9,X10,X11))),
  file('./7_Book/00EE/06/E051-32.p',ruleD18)).
fof(f376,plain,(
  ( ! [X6,X4,X2,X0,X7,X5,X3,X1] : (~eqangle(X0,X1,X2,X3,X4,X5,X6,X7) | eqangle(X2,X3,X0,X1,X6,X7,X4,X5)) )),
  inference(cnf_transformation,[],[f192])).
fof(f192,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6,X7] : (eqangle(X2,X3,X0,X1,X6,X7,X4,X5) | ~eqangle(X0,X1,X2,X3,X4,X5,X6,X7))),
  inference(ennf_transformation,[],[f101])).
fof(f101,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6,X7] : (eqangle(X0,X1,X2,X3,X4,X5,X6,X7) => eqangle(X2,X3,X0,X1,X6,X7,X4,X5))),
  inference(rectify,[],[f19])).
fof(f19,axiom,(
  ! [X0,X1,X2,X3,X8,X9,X10,X11] : (eqangle(X0,X1,X2,X3,X8,X9,X10,X11) => eqangle(X2,X3,X0,X1,X10,X11,X8,X9))),
  file('./7_Book/00EE/06/E051-32.p',ruleD19)).
fof(f2002,plain,(
  ( ! [X47,X45,X46,X44] : (~para(X44,X45,X46,X47) | perp(X44,X45,X46,X47)) ) | ~spl26_142),
  inference(avatar_component_clause,[],[f2001])).
fof(f13195,plain,(
  ~spl26_136 | ~spl26_140 | ~spl26_168),
  inference(avatar_contradiction_clause,[],[f13176])).
fof(f13176,plain,(
  $false | (~spl26_136 | ~spl26_140 | ~spl26_168)),
  inference(resolution,[],[f11339,f1999])).
fof(f1999,plain,(
  ( ! [X48,X49] : (~perp(X48,X49,X48,X49)) ) | ~spl26_140),
  inference(avatar_component_clause,[],[f1998])).
fof(f1998,plain,(
  spl26_140 <=> ! [X49,X48] : ~perp(X48,X49,X48,X49)),
  introduced(avatar_definition,[new_symbols(naming,[spl26_140])])).
fof(f11339,plain,(
  ( ! [X2,X0,X1] : (perp(X0,X0,X1,X2)) ) | (~spl26_136 | ~spl26_168)),
  inference(subsumption_resolution,[],[f11338,f9454])).
fof(f9454,plain,(
  ( ! [X2,X3] : (midp(X2,X3,X2)) ) | ~spl26_136),
  inference(resolution,[],[f8817,f368])).
fof(f368,plain,(
  ( ! [X2,X0,X1] : (~midp(X2,X1,X0) | midp(X2,X0,X1)) )),
  inference(cnf_transformation,[],[f181])).
fof(f181,plain,(
  ! [X0,X1,X2] : (midp(X2,X0,X1) | ~midp(X2,X1,X0))),
  inference(ennf_transformation,[],[f97])).
fof(f97,plain,(
  ! [X0,X1,X2] : (midp(X2,X1,X0) => midp(X2,X0,X1))),
  inference(rectify,[],[f11])).
fof(f11,axiom,(
  ! [X0,X1,X6] : (midp(X6,X1,X0) => midp(X6,X0,X1))),
  file('./7_Book/00EE/06/E051-32.p',ruleD11)).
fof(f8817,plain,(
  ( ! [X23,X22] : (midp(X22,X22,X23)) ) | ~spl26_136),
  inference(resolution,[],[f8048,f3294])).
fof(f3294,plain,(
  ( ! [X10,X8,X9] : (~midp(X10,X9,X8) | midp(X8,X9,X10)) ) | ~spl26_136),
  inference(subsumption_resolution,[],[f3289,f3263])).
fof(f3263,plain,(
  ( ! [X2,X0,X1] : (coll(X0,X1,X2)) ) | ~spl26_136),
  inference(subsumption_resolution,[],[f3260,f3250])).
fof(f3250,plain,(
  ( ! [X6,X5] : (coll(X5,X5,X6)) ) | ~spl26_136),
  inference(resolution,[],[f3237,f358])).
fof(f358,plain,(
  ( ! [X2,X0,X1] : (~coll(X0,X1,X2) | coll(X0,X2,X1)) )),
  inference(cnf_transformation,[],[f167])).
fof(f167,plain,(
  ! [X0,X1,X2] : (coll(X0,X2,X1) | ~coll(X0,X1,X2))),
  inference(ennf_transformation,[],[f1])).
fof(f1,axiom,(
  ! [X0,X1,X2] : (coll(X0,X1,X2) => coll(X0,X2,X1))),
  file('./7_Book/00EE/06/E051-32.p',ruleD1)).
fof(f3260,plain,(
  ( ! [X2,X0,X1] : (coll(X0,X1,X2) | ~coll(X2,X2,X0)) ) | ~spl26_136),
  inference(resolution,[],[f3250,f360])).
fof(f360,plain,(
  ( ! [X2,X0,X3,X1] : (~coll(X0,X1,X3) | coll(X2,X3,X0) | ~coll(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f170])).
fof(f170,plain,(
  ! [X0,X1,X2,X3] : (coll(X2,X3,X0) | ~coll(X0,X1,X3) | ~coll(X0,X1,X2))),
  inference(flattening,[],[f169])).
fof(f169,plain,(
  ! [X0,X1,X2,X3] : (coll(X2,X3,X0) | (~coll(X0,X1,X3) | ~coll(X0,X1,X2)))),
  inference(ennf_transformation,[],[f3])).
fof(f3,axiom,(
  ! [X0,X1,X2,X3] : ((coll(X0,X1,X3) & coll(X0,X1,X2)) => coll(X2,X3,X0))),
  file('./7_Book/00EE/06/E051-32.p',ruleD3)).
fof(f3289,plain,(
  ( ! [X10,X8,X9] : (~coll(X8,X9,X10) | midp(X8,X9,X10) | ~midp(X10,X9,X8)) ) | ~spl26_136),
  inference(resolution,[],[f3209,f403])).
fof(f403,plain,(
  ( ! [X4,X2,X0,X3,X1] : (~para(X3,X4,X1,X2) | ~coll(X4,X0,X2) | midp(X4,X0,X2) | ~midp(X3,X0,X1)) )),
  inference(cnf_transformation,[],[f229])).
fof(f229,plain,(
  ! [X0,X1,X2,X3,X4] : (midp(X4,X0,X2) | ~coll(X4,X0,X2) | ~para(X3,X4,X1,X2) | ~midp(X3,X0,X1))),
  inference(flattening,[],[f228])).
fof(f228,plain,(
  ! [X0,X1,X2,X3,X4] : (midp(X4,X0,X2) | (~coll(X4,X0,X2) | ~para(X3,X4,X1,X2) | ~midp(X3,X0,X1)))),
  inference(ennf_transformation,[],[f125])).
fof(f125,plain,(
  ! [X0,X1,X2,X3,X4] : ((coll(X4,X0,X2) & para(X3,X4,X1,X2) & midp(X3,X0,X1)) => midp(X4,X0,X2))),
  inference(rectify,[],[f46])).
fof(f46,axiom,(
  ! [X0,X1,X2,X4,X5] : ((coll(X5,X0,X2) & para(X4,X5,X1,X2) & midp(X4,X0,X1)) => midp(X5,X0,X2))),
  file('./7_Book/00EE/06/E051-32.p',ruleD45)).
fof(f3209,plain,(
  ( ! [X21,X20] : (para(X20,X21,X21,X20)) ) | ~spl26_136),
  inference(resolution,[],[f1818,f361])).
fof(f11338,plain,(
  ( ! [X2,X0,X1] : (perp(X0,X0,X1,X2) | ~midp(X1,X0,X1)) ) | (~spl26_136 | ~spl26_168)),
  inference(subsumption_resolution,[],[f11322,f8048])).
fof(f11322,plain,(
  ( ! [X2,X0,X1] : (perp(X0,X0,X1,X2) | ~midp(X0,X2,X2) | ~midp(X1,X0,X1)) ) | ~spl26_168),
  inference(resolution,[],[f567,f2446])).
fof(f2446,plain,(
  ( ! [X14,X15,X13] : (cong(X14,X13,X14,X15) | ~midp(X13,X14,X15)) ) | ~spl26_168),
  inference(avatar_component_clause,[],[f2445])).
fof(f2445,plain,(
  spl26_168 <=> ! [X13,X15,X14] : (~midp(X13,X14,X15) | cong(X14,X13,X14,X15))),
  introduced(avatar_definition,[new_symbols(naming,[spl26_168])])).
fof(f567,plain,(
  ( ! [X2,X0,X1] : (~cong(X0,X1,X0,X1) | perp(X0,X0,X1,X2) | ~midp(X0,X2,X2)) )),
  inference(resolution,[],[f414,f479])).
fof(f479,plain,(
  ( ! [X2,X0,X1] : (cong(X0,X1,X0,X2) | ~midp(X0,X2,X1)) )),
  inference(resolution,[],[f381,f426])).
fof(f426,plain,(
  ( ! [X2,X0,X1] : (cong(X0,X1,X0,X2) | ~midp(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f269])).
fof(f269,plain,(
  ! [X0,X1,X2] : (cong(X0,X1,X0,X2) | ~midp(X0,X1,X2))),
  inference(ennf_transformation,[],[f69])).
fof(f69,axiom,(
  ! [X0,X1,X2] : (midp(X0,X1,X2) => cong(X0,X1,X0,X2))),
  file('./7_Book/00EE/06/E051-32.p',ruleD68)).
fof(f381,plain,(
  ( ! [X2,X0,X3,X1] : (~cong(X0,X1,X2,X3) | cong(X2,X3,X0,X1)) )),
  inference(cnf_transformation,[],[f198])).
fof(f198,plain,(
  ! [X0,X1,X2,X3] : (cong(X2,X3,X0,X1) | ~cong(X0,X1,X2,X3))),
  inference(ennf_transformation,[],[f24])).
fof(f24,axiom,(
  ! [X0,X1,X2,X3] : (cong(X0,X1,X2,X3) => cong(X2,X3,X0,X1))),
  file('./7_Book/00EE/06/E051-32.p',ruleD24)).
fof(f7394,plain,(
  ~spl26_136 | ~spl26_170),
  inference(avatar_contradiction_clause,[],[f7393])).
fof(f7393,plain,(
  $false | (~spl26_136 | ~spl26_170)),
  inference(subsumption_resolution,[],[f7392,f2449])).
fof(f2449,plain,(
  ( ! [X12,X10,X11] : (~midp(X10,X11,X12)) ) | ~spl26_170),
  inference(avatar_component_clause,[],[f2448])).
fof(f2448,plain,(
  spl26_170 <=> ! [X11,X10,X12] : ~midp(X10,X11,X12)),
  introduced(avatar_definition,[new_symbols(naming,[spl26_170])])).
fof(f3181,plain,(
  ~spl26_138),
  inference(avatar_contradiction_clause,[],[f3180])).
fof(f3180,plain,(
  $false | ~spl26_138),
  inference(resolution,[],[f1821,f1519])).
fof(f1519,plain,(
  para(sK21,sK20,sK21,sK20)),
  inference(resolution,[],[f972,f361])).
fof(f972,plain,(
  para(sK21,sK20,sK20,sK21)),
  inference(subsumption_resolution,[],[f884,f819])).
fof(f884,plain,(
  ~para(sK20,sK23,sK21,sK23) | para(sK21,sK20,sK20,sK21)),
  inference(resolution,[],[f431,f657])).
fof(f657,plain,(
  eqangle(sK21,sK20,sK20,sK21,sK20,sK23,sK21,sK23)),
  inference(resolution,[],[f378,f471])).
fof(f431,plain,(
  ( ! [X6,X4,X2,X0,X7,X5,X3,X1] : (~eqangle(X0,X1,X2,X3,X4,X5,X6,X7) | ~para(X4,X5,X6,X7) | para(X0,X1,X2,X3)) )),
  inference(cnf_transformation,[],[f278])).
fof(f278,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6,X7] : (para(X0,X1,X2,X3) | ~para(X4,X5,X6,X7) | ~eqangle(X0,X1,X2,X3,X4,X5,X6,X7))),
  inference(flattening,[],[f277])).
fof(f277,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6,X7] : (para(X0,X1,X2,X3) | (~para(X4,X5,X6,X7) | ~eqangle(X0,X1,X2,X3,X4,X5,X6,X7)))),
  inference(ennf_transformation,[],[f146])).
fof(f146,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6,X7] : ((para(X4,X5,X6,X7) & eqangle(X0,X1,X2,X3,X4,X5,X6,X7)) => para(X0,X1,X2,X3))),
  inference(rectify,[],[f74])).
fof(f74,axiom,(
  ! [X0,X1,X2,X3,X8,X9,X10,X11] : ((para(X8,X9,X10,X11) & eqangle(X0,X1,X2,X3,X8,X9,X10,X11)) => para(X0,X1,X2,X3))),
  file('./7_Book/00EE/06/E051-32.p',ruleD73)).
fof(f1821,plain,(
  ( ! [X0,X1] : (~para(X0,X1,X0,X1)) ) | ~spl26_138),
  inference(avatar_component_clause,[],[f1820])).
fof(f1820,plain,(
  spl26_138 <=> ! [X1,X0] : ~para(X0,X1,X0,X1)),
  introduced(avatar_definition,[new_symbols(naming,[spl26_138])])).
fof(f2920,plain,(
  spl26_170 | ~spl26_108 | ~spl26_136),
  inference(avatar_split_clause,[],[f2919,f1817,f1302,f2448])).
fof(f1302,plain,(
  spl26_108 <=> ! [X7,X8,X6] : (~cong(X6,X7,X7,X8) | ~midp(X6,X7,X8))),
  introduced(avatar_definition,[new_symbols(naming,[spl26_108])])).
fof(f2919,plain,(
  ( ! [X4,X5,X3] : (~midp(X3,X5,X4)) ) | (~spl26_108 | ~spl26_136)),
  inference(subsumption_resolution,[],[f2914,f368])).
fof(f2914,plain,(
  ( ! [X4,X5,X3] : (~midp(X3,X4,X5) | ~midp(X3,X5,X4)) ) | (~spl26_108 | ~spl26_136)),
  inference(resolution,[],[f2910,f595])).
fof(f595,plain,(
  ( ! [X2,X0,X3,X1] : (circle(X0,X1,X2,X3) | ~midp(X0,X1,X3) | ~midp(X0,X2,X1)) )),
  inference(resolution,[],[f540,f479])).
fof(f540,plain,(
  ( ! [X6,X4,X7,X5] : (~cong(X4,X5,X4,X6) | circle(X4,X5,X6,X7) | ~midp(X4,X5,X7)) )),
  inference(resolution,[],[f369,f426])).
fof(f369,plain,(
  ( ! [X2,X0,X3,X1] : (~cong(X3,X0,X3,X2) | circle(X3,X0,X1,X2) | ~cong(X3,X0,X3,X1)) )),
  inference(cnf_transformation,[],[f183])).
fof(f183,plain,(
  ! [X0,X1,X2,X3] : (circle(X3,X0,X1,X2) | ~cong(X3,X0,X3,X2) | ~cong(X3,X0,X3,X1))),
  inference(flattening,[],[f182])).
fof(f182,plain,(
  ! [X0,X1,X2,X3] : (circle(X3,X0,X1,X2) | (~cong(X3,X0,X3,X2) | ~cong(X3,X0,X3,X1)))),
  inference(ennf_transformation,[],[f98])).
fof(f98,plain,(
  ! [X0,X1,X2,X3] : ((cong(X3,X0,X3,X2) & cong(X3,X0,X3,X1)) => circle(X3,X0,X1,X2))),
  inference(rectify,[],[f12])).
fof(f12,axiom,(
  ! [X0,X1,X2,X7] : ((cong(X7,X0,X7,X2) & cong(X7,X0,X7,X1)) => circle(X7,X0,X1,X2))),
  file('./7_Book/00EE/06/E051-32.p',ruleD12)).
fof(f2910,plain,(
  ( ! [X2,X0,X1] : (~circle(X1,X2,X0,X0)) ) | (~spl26_108 | ~spl26_136)),
  inference(subsumption_resolution,[],[f2909,f2467])).
fof(f2467,plain,(
  ( ! [X0,X1] : (~midp(X0,X1,X0)) ) | ~spl26_108),
  inference(subsumption_resolution,[],[f2457,f368])).
fof(f2457,plain,(
  ( ! [X0,X1] : (~midp(X0,X0,X1) | ~midp(X0,X1,X0)) ) | ~spl26_108),
  inference(resolution,[],[f1303,f479])).
fof(f1303,plain,(
  ( ! [X6,X8,X7] : (~cong(X6,X7,X7,X8) | ~midp(X6,X7,X8)) ) | ~spl26_108),
  inference(avatar_component_clause,[],[f1302])).
fof(f2909,plain,(
  ( ! [X2,X0,X1] : (~circle(X1,X2,X0,X0) | midp(X0,X0,X0)) ) | ~spl26_136),
  inference(subsumption_resolution,[],[f2907,f1856])).
fof(f1856,plain,(
  ( ! [X4,X3] : (coll(X3,X4,X3)) ) | ~spl26_136),
  inference(resolution,[],[f1825,f359])).
fof(f1825,plain,(
  ( ! [X0,X1] : (coll(X0,X1,X1)) ) | ~spl26_136),
  inference(resolution,[],[f1818,f424])).
fof(f2907,plain,(
  ( ! [X2,X0,X1] : (~coll(X0,X0,X0) | ~circle(X1,X2,X0,X0) | midp(X0,X0,X0)) ) | ~spl26_136),
  inference(resolution,[],[f1384,f2005])).
fof(f2005,plain,(
  ( ! [X68,X66,X67] : (cyclic(X67,X67,X66,X68)) ) | ~spl26_136),
  inference(subsumption_resolution,[],[f2004,f1913])).
fof(f1913,plain,(
  ( ! [X2,X0,X1] : (coll(X0,X1,X2)) ) | ~spl26_136),
  inference(subsumption_resolution,[],[f1910,f1886])).
fof(f1886,plain,(
  ( ! [X6,X5] : (coll(X5,X5,X6)) ) | ~spl26_136),
  inference(resolution,[],[f1856,f358])).
fof(f1910,plain,(
  ( ! [X2,X0,X1] : (coll(X0,X1,X2) | ~coll(X2,X2,X0)) ) | ~spl26_136),
  inference(resolution,[],[f1886,f360])).
fof(f2004,plain,(
  ( ! [X68,X66,X67] : (~coll(X66,X68,X67) | cyclic(X67,X67,X66,X68)) ) | ~spl26_136),
  inference(subsumption_resolution,[],[f1986,f1818])).
fof(f1384,plain,(
  ( ! [X6,X8,X7,X9] : (~cyclic(X7,X6,X9,X8) | ~coll(X6,X7,X6) | ~circle(X8,X9,X7,X6) | midp(X6,X7,X6)) )),
  inference(resolution,[],[f409,f398])).
fof(f409,plain,(
  ( ! [X4,X2,X0,X3,X1] : (~eqangle(X0,X1,X0,X2,X3,X1,X3,X4) | midp(X4,X1,X2) | ~coll(X4,X1,X2) | ~circle(X3,X0,X1,X2)) )),
  inference(cnf_transformation,[],[f240])).
fof(f240,plain,(
  ! [X0,X1,X2,X3,X4] : (midp(X4,X1,X2) | ~eqangle(X0,X1,X0,X2,X3,X1,X3,X4) | ~coll(X4,X1,X2) | ~circle(X3,X0,X1,X2))),
  inference(flattening,[],[f239])).
fof(f239,plain,(
  ! [X0,X1,X2,X3,X4] : (midp(X4,X1,X2) | (~eqangle(X0,X1,X0,X2,X3,X1,X3,X4) | ~coll(X4,X1,X2) | ~circle(X3,X0,X1,X2)))),
  inference(ennf_transformation,[],[f131])).
fof(f131,plain,(
  ! [X0,X1,X2,X3,X4] : ((eqangle(X0,X1,X0,X2,X3,X1,X3,X4) & coll(X4,X1,X2) & circle(X3,X0,X1,X2)) => midp(X4,X1,X2))),
  inference(rectify,[],[f52])).
fof(f52,axiom,(
  ! [X0,X1,X2,X7,X6] : ((eqangle(X0,X1,X0,X2,X7,X1,X7,X6) & coll(X6,X1,X2) & circle(X7,X0,X1,X2)) => midp(X6,X1,X2))),
  file('./7_Book/00EE/06/E051-32.p',ruleD51)).
fof(f2451,plain,(
  spl26_168 | spl26_108),
  inference(avatar_split_clause,[],[f2438,f1302,f2445])).
fof(f2438,plain,(
  ( ! [X14,X12,X10,X15,X13,X11] : (~midp(X10,X11,X12) | ~midp(X13,X14,X15) | ~cong(X10,X11,X11,X12) | cong(X14,X13,X14,X15)) )),
  inference(resolution,[],[f667,f433])).
fof(f433,plain,(
  ( ! [X6,X4,X2,X0,X7,X5,X3,X1] : (~eqratio(X0,X1,X2,X3,X4,X5,X6,X7) | ~cong(X4,X5,X6,X7) | cong(X0,X1,X2,X3)) )),
  inference(cnf_transformation,[],[f282])).
fof(f282,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6,X7] : (cong(X0,X1,X2,X3) | ~cong(X4,X5,X6,X7) | ~eqratio(X0,X1,X2,X3,X4,X5,X6,X7))),
  inference(flattening,[],[f281])).
fof(f281,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6,X7] : (cong(X0,X1,X2,X3) | (~cong(X4,X5,X6,X7) | ~eqratio(X0,X1,X2,X3,X4,X5,X6,X7)))),
  inference(ennf_transformation,[],[f148])).
fof(f148,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6,X7] : ((cong(X4,X5,X6,X7) & eqratio(X0,X1,X2,X3,X4,X5,X6,X7)) => cong(X0,X1,X2,X3))),
  inference(rectify,[],[f76])).
fof(f76,axiom,(
  ! [X0,X1,X2,X3,X8,X9,X10,X11] : ((cong(X8,X9,X10,X11) & eqratio(X0,X1,X2,X3,X8,X9,X10,X11)) => cong(X0,X1,X2,X3))),
  file('./7_Book/00EE/06/E051-32.p',ruleD75)).
fof(f667,plain,(
  ( ! [X6,X10,X8,X7,X11,X9] : (eqratio(X6,X7,X6,X8,X9,X10,X10,X11) | ~midp(X9,X10,X11) | ~midp(X7,X6,X8)) )),
  inference(resolution,[],[f383,f428])).
fof(f428,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (eqratio(X4,X0,X0,X1,X5,X2,X2,X3) | ~midp(X5,X2,X3) | ~midp(X4,X0,X1)) )),
  inference(cnf_transformation,[],[f272])).
fof(f272,plain,(
  ! [X0,X1,X2,X3,X4,X5] : (eqratio(X4,X0,X0,X1,X5,X2,X2,X3) | ~midp(X5,X2,X3) | ~midp(X4,X0,X1))),
  inference(flattening,[],[f271])).
fof(f271,plain,(
  ! [X0,X1,X2,X3,X4,X5] : (eqratio(X4,X0,X0,X1,X5,X2,X2,X3) | (~midp(X5,X2,X3) | ~midp(X4,X0,X1)))),
  inference(ennf_transformation,[],[f145])).
fof(f145,plain,(
  ! [X0,X1,X2,X3,X4,X5] : ((midp(X5,X2,X3) & midp(X4,X0,X1)) => eqratio(X4,X0,X0,X1,X5,X2,X2,X3))),
  inference(rectify,[],[f71])).
fof(f71,axiom,(
  ! [X0,X1,X2,X3,X6,X16] : ((midp(X16,X2,X3) & midp(X6,X0,X1)) => eqratio(X6,X0,X0,X1,X16,X2,X2,X3))),
  file('./7_Book/00EE/06/E051-32.p',ruleD70)).
fof(f383,plain,(
  ( ! [X6,X4,X2,X0,X7,X5,X3,X1] : (~eqratio(X0,X1,X2,X3,X4,X5,X6,X7) | eqratio(X1,X0,X2,X3,X4,X5,X6,X7)) )),
  inference(cnf_transformation,[],[f201])).
fof(f201,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6,X7] : (eqratio(X1,X0,X2,X3,X4,X5,X6,X7) | ~eqratio(X0,X1,X2,X3,X4,X5,X6,X7))),
  inference(ennf_transformation,[],[f105])).
fof(f105,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6,X7] : (eqratio(X0,X1,X2,X3,X4,X5,X6,X7) => eqratio(X1,X0,X2,X3,X4,X5,X6,X7))),
  inference(rectify,[],[f26])).
fof(f26,axiom,(
  ! [X0,X1,X2,X3,X8,X9,X10,X11] : (eqratio(X0,X1,X2,X3,X8,X9,X10,X11) => eqratio(X1,X0,X2,X3,X8,X9,X10,X11))),
  file('./7_Book/00EE/06/E051-32.p',ruleD26)).
fof(f2003,plain,(
  spl26_140 | spl26_142),
  inference(avatar_split_clause,[],[f1980,f2001,f1998])).
fof(f1980,plain,(
  ( ! [X47,X45,X48,X46,X44,X49] : (~para(X44,X45,X46,X47) | ~perp(X48,X49,X48,X49) | perp(X44,X45,X46,X47)) )),
  inference(resolution,[],[f649,f432])).
fof(f1822,plain,(
  spl26_136 | spl26_138),
  inference(avatar_split_clause,[],[f1793,f1820,f1817])).
fof(f1793,plain,(
  ( ! [X2,X0,X3,X1] : (~para(X0,X1,X0,X1) | para(X2,X3,X2,X3)) )),
  inference(resolution,[],[f628,f396])).
fof(f628,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (eqangle(X0,X1,X2,X3,X0,X1,X4,X5) | ~para(X2,X3,X4,X5)) )),
  inference(resolution,[],[f376,f397])).
fof(f1766,plain,(
  spl26_132 | ~spl26_20),
  inference(avatar_split_clause,[],[f1754,f954,f1757])).
fof(f1757,plain,(
  spl26_132 <=> para(sK21,sK20,sK21,sK23)),
  introduced(avatar_definition,[new_symbols(naming,[spl26_132])])).
fof(f954,plain,(
  spl26_20 <=> para(sK21,sK23,sK20,sK21)),
  introduced(avatar_definition,[new_symbols(naming,[spl26_20])])).
fof(f1754,plain,(
  para(sK21,sK20,sK21,sK23) | ~spl26_20),
  inference(resolution,[],[f1498,f362])).
fof(f362,plain,(
  ( ! [X2,X0,X3,X1] : (~para(X0,X1,X2,X3) | para(X2,X3,X0,X1)) )),
  inference(cnf_transformation,[],[f172])).
fof(f172,plain,(
  ! [X0,X1,X2,X3] : (para(X2,X3,X0,X1) | ~para(X0,X1,X2,X3))),
  inference(ennf_transformation,[],[f5])).
fof(f5,axiom,(
  ! [X0,X1,X2,X3] : (para(X0,X1,X2,X3) => para(X2,X3,X0,X1))),
  file('./7_Book/00EE/06/E051-32.p',ruleD5)).
fof(f1498,plain,(
  para(sK21,sK23,sK21,sK20) | ~spl26_20),
  inference(resolution,[],[f955,f361])).
fof(f955,plain,(
  para(sK21,sK23,sK20,sK21) | ~spl26_20),
  inference(avatar_component_clause,[],[f954])).
fof(f1765,plain,(
  ~spl26_133 | spl26_134 | ~spl26_20),
  inference(avatar_split_clause,[],[f1751,f954,f1763,f1760])).
fof(f1760,plain,(
  spl26_133 <=> ~para(sK21,sK20,sK21,sK23)),
  introduced(avatar_definition,[new_symbols(naming,[spl26_133])])).
fof(f1751,plain,(
  ( ! [X0] : (midp(X0,sK20,sK23) | ~para(sK21,sK20,sK21,sK23) | ~midp(X0,sK21,sK21)) ) | ~spl26_20),
  inference(resolution,[],[f1498,f422])).
fof(f1443,plain,(
  spl26_20 | ~spl26_14),
  inference(avatar_split_clause,[],[f1429,f938,f954])).
fof(f938,plain,(
  spl26_14 <=> para(sK20,sK21,sK21,sK23)),
  introduced(avatar_definition,[new_symbols(naming,[spl26_14])])).
fof(f1429,plain,(
  para(sK21,sK23,sK20,sK21) | ~spl26_14),
  inference(resolution,[],[f939,f362])).
fof(f939,plain,(
  para(sK20,sK21,sK21,sK23) | ~spl26_14),
  inference(avatar_component_clause,[],[f938])).
fof(f971,plain,(
  spl26_16),
  inference(avatar_split_clause,[],[f970,f941])).
fof(f941,plain,(
  spl26_16 <=> para(sK21,sK20,sK20,sK23)),
  introduced(avatar_definition,[new_symbols(naming,[spl26_16])])).
fof(f970,plain,(
  para(sK21,sK20,sK20,sK23)),
  inference(subsumption_resolution,[],[f883,f837])).
fof(f883,plain,(
  ~para(sK20,sK23,sK23,sK21) | para(sK21,sK20,sK20,sK23)),
  inference(resolution,[],[f431,f472])).
fof(f946,plain,(
  spl26_14 | ~spl26_17),
  inference(avatar_split_clause,[],[f877,f944,f938])).
fof(f944,plain,(
  spl26_17 <=> ~para(sK21,sK20,sK20,sK23)),
  introduced(avatar_definition,[new_symbols(naming,[spl26_17])])).
fof(f877,plain,(
  ~para(sK21,sK20,sK20,sK23) | para(sK20,sK21,sK21,sK23)),
  inference(resolution,[],[f431,f634])).
% SZS output end Proof for E051-32
% ------------------------------
% Version: Vampire 4.4.0 (commit 7916d27 on 2019-08-23 08:50:16 +0100)
% Termination reason: Refutation

% Memory used [KB]: 17398
% Time elapsed: 0.436 s
% ------------------------------
% ------------------------------
