% Refutation found. Thanks to Tanya!
% SZS status Theorem for E051-29
% SZS output start Proof for E051-29
fof(f1463,plain,(
  $false),
  inference(resolution,[],[f1345,f516])).
fof(f516,plain,(
  ~midp(sK21,sK22,sK20)),
  inference(resolution,[],[f505,f481])).
fof(f481,plain,(
  ~cong(sK22,sK21,sK21,sK20)),
  inference(cnf_transformation,[],[f357])).
fof(f357,plain,(
  ~cong(sK22,sK21,sK21,sK20) & ! [X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16] : (perp(X10,X3,X10,X9) & circle(X3,X4,X10,X16) & coll(X9,X6,X7) & para(X4,X5,X9,X8) & coll(X8,X5,X6) & coll(X8,X4,X7) & circle(X3,X4,X7,X15) & circle(X3,X4,X6,X14) & circle(X3,X4,X5,X13) & circle(X3,X4,X11,X12))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK20,sK21,sK22])],[f319,f356])).
fof(f356,plain,(
  ? [X0,X1,X2] : (~cong(X2,X1,X1,X0) & ! [X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16] : (perp(X10,X3,X10,X9) & circle(X3,X4,X10,X16) & coll(X9,X6,X7) & para(X4,X5,X9,X8) & coll(X8,X5,X6) & coll(X8,X4,X7) & circle(X3,X4,X7,X15) & circle(X3,X4,X6,X14) & circle(X3,X4,X5,X13) & circle(X3,X4,X11,X12))) => (~cong(sK22,sK21,sK21,sK20) & ! [X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16] : (perp(X10,X3,X10,X9) & circle(X3,X4,X10,X16) & coll(X9,X6,X7) & para(X4,X5,X9,X8) & coll(X8,X5,X6) & coll(X8,X4,X7) & circle(X3,X4,X7,X15) & circle(X3,X4,X6,X14) & circle(X3,X4,X5,X13) & circle(X3,X4,X11,X12)))),
  introduced(choice_axiom,[])).
fof(f319,plain,(
  ? [X0,X1,X2] : (~cong(X2,X1,X1,X0) & ! [X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16] : (perp(X10,X3,X10,X9) & circle(X3,X4,X10,X16) & coll(X9,X6,X7) & para(X4,X5,X9,X8) & coll(X8,X5,X6) & coll(X8,X4,X7) & circle(X3,X4,X7,X15) & circle(X3,X4,X6,X14) & circle(X3,X4,X5,X13) & circle(X3,X4,X11,X12)))),
  inference(ennf_transformation,[],[f167])).
fof(f167,plain,(
  ~! [X0,X1,X2] : (! [X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16] : (perp(X10,X3,X10,X9) & circle(X3,X4,X10,X16) & coll(X9,X6,X7) & para(X4,X5,X9,X8) & coll(X8,X5,X6) & coll(X8,X4,X7) & circle(X3,X4,X7,X15) & circle(X3,X4,X6,X14) & circle(X3,X4,X5,X13) & circle(X3,X4,X11,X12)) => cong(X2,X1,X1,X0))),
  inference(rectify,[],[f96])).
fof(f96,negated_conjecture,(
  ~! [X13,X12,X5] : (! [X0,X1,X2,X3,X4,X5,X12,X13,X18,X19,X20,X21,X22,X23] : (perp(X13,X0,X13,X12) & circle(X0,X1,X13,X23) & coll(X12,X3,X4) & para(X1,X2,X12,X5) & coll(X5,X2,X3) & coll(X5,X1,X4) & circle(X0,X1,X4,X22) & circle(X0,X1,X3,X21) & circle(X0,X1,X2,X20) & circle(X0,X1,X18,X19)) => cong(X5,X12,X12,X13))),
  inference(negated_conjecture,[],[f95])).
fof(f95,conjecture,(
  ( ! [X12,X5,X13] : (! [X0,X1,X2,X3,X4,X5,X12,X13,X18,X19,X20,X21,X22,X23] : (perp(X13,X0,X13,X12) & circle(X0,X1,X13,X23) & coll(X12,X3,X4) & para(X1,X2,X12,X5) & coll(X5,X2,X3) & coll(X5,X1,X4) & circle(X0,X1,X4,X22) & circle(X0,X1,X3,X21) & circle(X0,X1,X2,X20) & circle(X0,X1,X18,X19)) => cong(X5,X12,X12,X13)) )),
  file('./7_Book/00EE/06/E051-29.p',exemplo7Book00EE06E05129)).
fof(f505,plain,(
  ( ! [X2,X0,X1] : (cong(X1,X0,X0,X2) | ~midp(X0,X1,X2)) )),
  inference(resolution,[],[f500,f381])).
fof(f381,plain,(
  ( ! [X2,X0,X3,X1] : (~cong(X0,X1,X2,X3) | cong(X2,X3,X0,X1)) )),
  inference(cnf_transformation,[],[f199])).
fof(f199,plain,(
  ! [X0,X1,X2,X3] : (cong(X2,X3,X0,X1) | ~cong(X0,X1,X2,X3))),
  inference(ennf_transformation,[],[f24])).
fof(f24,axiom,(
  ! [X0,X1,X2,X3] : (cong(X0,X1,X2,X3) => cong(X2,X3,X0,X1))),
  file('./7_Book/00EE/06/E051-29.p',ruleD24)).
fof(f500,plain,(
  ( ! [X4,X5,X3] : (cong(X3,X5,X4,X3) | ~midp(X3,X4,X5)) )),
  inference(resolution,[],[f497,f380])).
fof(f380,plain,(
  ( ! [X2,X0,X3,X1] : (~cong(X0,X1,X2,X3) | cong(X0,X1,X3,X2)) )),
  inference(cnf_transformation,[],[f198])).
fof(f198,plain,(
  ! [X0,X1,X2,X3] : (cong(X0,X1,X3,X2) | ~cong(X0,X1,X2,X3))),
  inference(ennf_transformation,[],[f23])).
fof(f23,axiom,(
  ! [X0,X1,X2,X3] : (cong(X0,X1,X2,X3) => cong(X0,X1,X3,X2))),
  file('./7_Book/00EE/06/E051-29.p',cong23)).
fof(f497,plain,(
  ( ! [X2,X0,X1] : (cong(X0,X1,X0,X2) | ~midp(X0,X2,X1)) )),
  inference(resolution,[],[f381,f426])).
fof(f426,plain,(
  ( ! [X2,X0,X1] : (cong(X0,X1,X0,X2) | ~midp(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f270])).
fof(f270,plain,(
  ! [X0,X1,X2] : (cong(X0,X1,X0,X2) | ~midp(X0,X1,X2))),
  inference(ennf_transformation,[],[f69])).
fof(f69,axiom,(
  ! [X0,X1,X2] : (midp(X0,X1,X2) => cong(X0,X1,X0,X2))),
  file('./7_Book/00EE/06/E051-29.p',ruleD68)).
fof(f1345,plain,(
  ( ! [X54,X52,X53] : (midp(X52,X53,X54)) )),
  inference(subsumption_resolution,[],[f1344,f479])).
fof(f479,plain,(
  ( ! [X4,X10,X3,X16] : (circle(X3,X4,X10,X16)) )),
  inference(cnf_transformation,[],[f357])).
fof(f1344,plain,(
  ( ! [X54,X52,X55,X53] : (midp(X52,X53,X54) | ~circle(X55,X52,X53,X54)) )),
  inference(subsumption_resolution,[],[f1232,f478])).
fof(f478,plain,(
  ( ! [X6,X7,X9] : (coll(X9,X6,X7)) )),
  inference(cnf_transformation,[],[f357])).
fof(f1232,plain,(
  ( ! [X54,X52,X55,X53] : (midp(X52,X53,X54) | ~coll(X52,X53,X54) | ~circle(X55,X52,X53,X54)) )),
  inference(resolution,[],[f409,f858])).
fof(f858,plain,(
  ( ! [X12,X10,X13,X11] : (eqangle(X10,X11,X10,X12,X13,X11,X13,X10)) )),
  inference(resolution,[],[f854,f376])).
fof(f376,plain,(
  ( ! [X6,X4,X2,X0,X7,X5,X3,X1] : (~eqangle(X0,X1,X2,X3,X4,X5,X6,X7) | eqangle(X2,X3,X0,X1,X6,X7,X4,X5)) )),
  inference(cnf_transformation,[],[f193])).
fof(f193,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6,X7] : (eqangle(X2,X3,X0,X1,X6,X7,X4,X5) | ~eqangle(X0,X1,X2,X3,X4,X5,X6,X7))),
  inference(ennf_transformation,[],[f101])).
fof(f101,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6,X7] : (eqangle(X0,X1,X2,X3,X4,X5,X6,X7) => eqangle(X2,X3,X0,X1,X6,X7,X4,X5))),
  inference(rectify,[],[f19])).
fof(f19,axiom,(
  ! [X0,X1,X2,X3,X8,X9,X10,X11] : (eqangle(X0,X1,X2,X3,X8,X9,X10,X11) => eqangle(X2,X3,X0,X1,X10,X11,X8,X9))),
  file('./7_Book/00EE/06/E051-29.p',ruleD19)).
fof(f854,plain,(
  ( ! [X4,X2,X3,X1] : (eqangle(X1,X2,X1,X3,X4,X1,X4,X3)) )),
  inference(subsumption_resolution,[],[f853,f550])).
fof(f550,plain,(
  ( ! [X2,X0,X3,X1] : (perp(X0,X1,X2,X3)) )),
  inference(subsumption_resolution,[],[f545,f477])).
fof(f477,plain,(
  ( ! [X4,X8,X5,X9] : (para(X4,X5,X9,X8)) )),
  inference(cnf_transformation,[],[f357])).
fof(f545,plain,(
  ( ! [X4,X2,X0,X3,X1] : (perp(X0,X1,X2,X3) | ~para(X0,X1,X2,X4)) )),
  inference(resolution,[],[f367,f480])).
fof(f480,plain,(
  ( ! [X10,X3,X9] : (perp(X10,X3,X10,X9)) )),
  inference(cnf_transformation,[],[f357])).
fof(f367,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (~perp(X2,X3,X4,X5) | perp(X0,X1,X4,X5) | ~para(X0,X1,X2,X3)) )),
  inference(cnf_transformation,[],[f181])).
fof(f181,plain,(
  ! [X0,X1,X2,X3,X4,X5] : (perp(X0,X1,X4,X5) | ~perp(X2,X3,X4,X5) | ~para(X0,X1,X2,X3))),
  inference(flattening,[],[f180])).
fof(f180,plain,(
  ! [X0,X1,X2,X3,X4,X5] : (perp(X0,X1,X4,X5) | (~perp(X2,X3,X4,X5) | ~para(X0,X1,X2,X3)))),
  inference(ennf_transformation,[],[f10])).
fof(f10,axiom,(
  ! [X0,X1,X2,X3,X4,X5] : ((perp(X2,X3,X4,X5) & para(X0,X1,X2,X3)) => perp(X0,X1,X4,X5))),
  file('./7_Book/00EE/06/E051-29.p',ruleD10)).
fof(f853,plain,(
  ( ! [X4,X2,X0,X3,X1] : (~perp(X0,X1,X1,X2) | eqangle(X1,X2,X1,X3,X4,X1,X4,X3)) )),
  inference(resolution,[],[f406,f479])).
fof(f406,plain,(
  ( ! [X4,X2,X0,X3,X1] : (~circle(X3,X0,X1,X2) | ~perp(X3,X0,X0,X4) | eqangle(X0,X4,X0,X1,X2,X0,X2,X1)) )),
  inference(cnf_transformation,[],[f235])).
fof(f235,plain,(
  ! [X0,X1,X2,X3,X4] : (eqangle(X0,X4,X0,X1,X2,X0,X2,X1) | ~perp(X3,X0,X0,X4) | ~circle(X3,X0,X1,X2))),
  inference(flattening,[],[f234])).
fof(f234,plain,(
  ! [X0,X1,X2,X3,X4] : (eqangle(X0,X4,X0,X1,X2,X0,X2,X1) | (~perp(X3,X0,X0,X4) | ~circle(X3,X0,X1,X2)))),
  inference(ennf_transformation,[],[f128])).
fof(f128,plain,(
  ! [X0,X1,X2,X3,X4] : ((perp(X3,X0,X0,X4) & circle(X3,X0,X1,X2)) => eqangle(X0,X4,X0,X1,X2,X0,X2,X1))),
  inference(rectify,[],[f49])).
fof(f49,axiom,(
  ! [X0,X1,X2,X7,X15] : ((perp(X7,X0,X0,X15) & circle(X7,X0,X1,X2)) => eqangle(X0,X15,X0,X1,X2,X0,X2,X1))),
  file('./7_Book/00EE/06/E051-29.p',ruleD48)).
fof(f409,plain,(
  ( ! [X4,X2,X0,X3,X1] : (~eqangle(X0,X1,X0,X2,X3,X1,X3,X4) | midp(X4,X1,X2) | ~coll(X4,X1,X2) | ~circle(X3,X0,X1,X2)) )),
  inference(cnf_transformation,[],[f241])).
fof(f241,plain,(
  ! [X0,X1,X2,X3,X4] : (midp(X4,X1,X2) | ~eqangle(X0,X1,X0,X2,X3,X1,X3,X4) | ~coll(X4,X1,X2) | ~circle(X3,X0,X1,X2))),
  inference(flattening,[],[f240])).
fof(f240,plain,(
  ! [X0,X1,X2,X3,X4] : (midp(X4,X1,X2) | (~eqangle(X0,X1,X0,X2,X3,X1,X3,X4) | ~coll(X4,X1,X2) | ~circle(X3,X0,X1,X2)))),
  inference(ennf_transformation,[],[f131])).
fof(f131,plain,(
  ! [X0,X1,X2,X3,X4] : ((eqangle(X0,X1,X0,X2,X3,X1,X3,X4) & coll(X4,X1,X2) & circle(X3,X0,X1,X2)) => midp(X4,X1,X2))),
  inference(rectify,[],[f52])).
fof(f52,axiom,(
  ! [X0,X1,X2,X7,X6] : ((eqangle(X0,X1,X0,X2,X7,X1,X7,X6) & coll(X6,X1,X2) & circle(X7,X0,X1,X2)) => midp(X6,X1,X2))),
  file('./7_Book/00EE/06/E051-29.p',ruleD51)).
% SZS output end Proof for E051-29
% ------------------------------
% Version: Vampire 4.4.0 (commit 7916d27 on 2019-08-23 08:50:16 +0100)
% Termination reason: Refutation

% Memory used [KB]: 5500
% Time elapsed: 0.037 s
% ------------------------------
% ------------------------------
