% Refutation found. Thanks to Tanya!
% SZS status Theorem for E051-32
% SZS output start Proof for E051-32
4. ! [X0,X1,X2,X3] : (para(X0,X1,X2,X3) => para(X0,X1,X3,X2)) [input]
5. ! [X0,X1,X2,X3] : (para(X0,X1,X2,X3) => para(X2,X3,X0,X1)) [input]
6. ! [X0,X1,X2,X3,X4,X5] : ((para(X2,X3,X4,X5) & para(X0,X1,X2,X3)) => para(X0,X1,X4,X5)) [input]
18. ! [X0,X1,X2,X3,X8,X9,X10,X11] : (eqangle(X0,X1,X2,X3,X8,X9,X10,X11) => eqangle(X1,X0,X2,X3,X8,X9,X10,X11)) [input]
19. ! [X0,X1,X2,X3,X8,X9,X10,X11] : (eqangle(X0,X1,X2,X3,X8,X9,X10,X11) => eqangle(X2,X3,X0,X1,X10,X11,X8,X9)) [input]
21. ! [X0,X1,X2,X3,X8,X9,X10,X11] : (eqangle(X0,X1,X2,X3,X8,X9,X10,X11) => eqangle(X0,X1,X8,X9,X2,X3,X10,X11)) [input]
22. ! [X0,X1,X2,X3,X8,X9,X10,X11,X4,X5,X12,X13] : ((eqangle(X8,X9,X10,X11,X4,X5,X12,X13) & eqangle(X0,X1,X2,X3,X8,X9,X10,X11)) => eqangle(X0,X1,X2,X3,X4,X5,X12,X13)) [input]
39. ! [X0,X1,X2,X3,X8,X9] : (eqangle(X0,X1,X8,X9,X2,X3,X8,X9) => para(X0,X1,X2,X3)) [input]
73. ! [X0,X1,X2,X3,X8,X9,X10,X11] : ((para(X8,X9,X10,X11) & eqangle(X0,X1,X2,X3,X8,X9,X10,X11)) => para(X0,X1,X2,X3)) [input]
94. ! [X0,X1,X2,X3,X4,X5] : (eqangle(X1,X2,X2,X5,X2,X5,X5,X1) & eqangle(X1,X2,X2,X5,X2,X1,X1,X5) & eqangle(X0,X2,X2,X4,X2,X4,X4,X0) & eqangle(X0,X2,X2,X4,X2,X0,X0,X4) & eqangle(X1,X0,X0,X3,X0,X3,X3,X1) & eqangle(X1,X0,X0,X3,X0,X1,X1,X3)) => para(X4,X3,X2,X5) [input]
95. ~! [X5,X2,X3,X4] : (! [X0,X1,X2,X3,X4,X5] : (eqangle(X1,X2,X2,X5,X2,X5,X5,X1) & eqangle(X1,X2,X2,X5,X2,X1,X1,X5) & eqangle(X0,X2,X2,X4,X2,X4,X4,X0) & eqangle(X0,X2,X2,X4,X2,X0,X0,X4) & eqangle(X1,X0,X0,X3,X0,X3,X3,X1) & eqangle(X1,X0,X0,X3,X0,X1,X1,X3)) => para(X4,X3,X2,X5)) [negated conjecture 94]
99. ! [X0,X1,X2,X3,X4,X5,X6,X7] : (eqangle(X0,X1,X2,X3,X4,X5,X6,X7) => eqangle(X1,X0,X2,X3,X4,X5,X6,X7)) [rectify 18]
100. ! [X0,X1,X2,X3,X4,X5,X6,X7] : (eqangle(X0,X1,X2,X3,X4,X5,X6,X7) => eqangle(X2,X3,X0,X1,X6,X7,X4,X5)) [rectify 19]
102. ! [X0,X1,X2,X3,X4,X5,X6,X7] : (eqangle(X0,X1,X2,X3,X4,X5,X6,X7) => eqangle(X0,X1,X4,X5,X2,X3,X6,X7)) [rectify 21]
103. ! [X0,X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11] : ((eqangle(X4,X5,X6,X7,X8,X9,X10,X11) & eqangle(X0,X1,X2,X3,X4,X5,X6,X7)) => eqangle(X0,X1,X2,X3,X8,X9,X10,X11)) [rectify 22]
117. ! [X0,X1,X2,X3,X4,X5] : (eqangle(X0,X1,X4,X5,X2,X3,X4,X5) => para(X0,X1,X2,X3)) [rectify 39]
144. ! [X0,X1,X2,X3,X4,X5,X6,X7] : ((para(X4,X5,X6,X7) & eqangle(X0,X1,X2,X3,X4,X5,X6,X7)) => para(X0,X1,X2,X3)) [rectify 73]
165. ~! [X0,X1,X2,X3] : (! [X4,X5,X6,X7,X8,X9] : (eqangle(X5,X6,X6,X9,X6,X9,X9,X5) & eqangle(X5,X6,X6,X9,X6,X5,X5,X9) & eqangle(X4,X6,X6,X8,X6,X8,X8,X4) & eqangle(X4,X6,X6,X8,X6,X4,X4,X8) & eqangle(X5,X4,X4,X7,X4,X7,X7,X5) & eqangle(X5,X4,X4,X7,X4,X5,X5,X7)) => para(X3,X2,X1,X0)) [rectify 95]
170. ! [X0,X1,X2,X3] : (para(X0,X1,X3,X2) | ~para(X0,X1,X2,X3)) [ennf transformation 4]
171. ! [X0,X1,X2,X3] : (para(X2,X3,X0,X1) | ~para(X0,X1,X2,X3)) [ennf transformation 5]
172. ! [X0,X1,X2,X3,X4,X5] : (para(X0,X1,X4,X5) | (~para(X2,X3,X4,X5) | ~para(X0,X1,X2,X3))) [ennf transformation 6]
173. ! [X0,X1,X2,X3,X4,X5] : (para(X0,X1,X4,X5) | ~para(X2,X3,X4,X5) | ~para(X0,X1,X2,X3)) [flattening 172]
190. ! [X0,X1,X2,X3,X4,X5,X6,X7] : (eqangle(X1,X0,X2,X3,X4,X5,X6,X7) | ~eqangle(X0,X1,X2,X3,X4,X5,X6,X7)) [ennf transformation 99]
191. ! [X0,X1,X2,X3,X4,X5,X6,X7] : (eqangle(X2,X3,X0,X1,X6,X7,X4,X5) | ~eqangle(X0,X1,X2,X3,X4,X5,X6,X7)) [ennf transformation 100]
193. ! [X0,X1,X2,X3,X4,X5,X6,X7] : (eqangle(X0,X1,X4,X5,X2,X3,X6,X7) | ~eqangle(X0,X1,X2,X3,X4,X5,X6,X7)) [ennf transformation 102]
194. ! [X0,X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11] : (eqangle(X0,X1,X2,X3,X8,X9,X10,X11) | (~eqangle(X4,X5,X6,X7,X8,X9,X10,X11) | ~eqangle(X0,X1,X2,X3,X4,X5,X6,X7))) [ennf transformation 103]
195. ! [X0,X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11] : (eqangle(X0,X1,X2,X3,X8,X9,X10,X11) | ~eqangle(X4,X5,X6,X7,X8,X9,X10,X11) | ~eqangle(X0,X1,X2,X3,X4,X5,X6,X7)) [flattening 194]
216. ! [X0,X1,X2,X3,X4,X5] : (para(X0,X1,X2,X3) | ~eqangle(X0,X1,X4,X5,X2,X3,X4,X5)) [ennf transformation 117]
274. ! [X0,X1,X2,X3,X4,X5,X6,X7] : (para(X0,X1,X2,X3) | (~para(X4,X5,X6,X7) | ~eqangle(X0,X1,X2,X3,X4,X5,X6,X7))) [ennf transformation 144]
275. ! [X0,X1,X2,X3,X4,X5,X6,X7] : (para(X0,X1,X2,X3) | ~para(X4,X5,X6,X7) | ~eqangle(X0,X1,X2,X3,X4,X5,X6,X7)) [flattening 274]
315. ? [X0,X1,X2,X3] : (~para(X3,X2,X1,X0) & ! [X4,X5,X6,X7,X8,X9] : (eqangle(X5,X6,X6,X9,X6,X9,X9,X5) & eqangle(X5,X6,X6,X9,X6,X5,X5,X9) & eqangle(X4,X6,X6,X8,X6,X8,X8,X4) & eqangle(X4,X6,X6,X8,X6,X4,X4,X8) & eqangle(X5,X4,X4,X7,X4,X7,X7,X5) & eqangle(X5,X4,X4,X7,X4,X5,X5,X7))) [ennf transformation 165]
352. ? [X0,X1,X2,X3] : (~para(X3,X2,X1,X0) & ! [X4,X5,X6,X7,X8,X9] : (eqangle(X5,X6,X6,X9,X6,X9,X9,X5) & eqangle(X5,X6,X6,X9,X6,X5,X5,X9) & eqangle(X4,X6,X6,X8,X6,X8,X8,X4) & eqangle(X4,X6,X6,X8,X6,X4,X4,X8) & eqangle(X5,X4,X4,X7,X4,X7,X7,X5) & eqangle(X5,X4,X4,X7,X4,X5,X5,X7))) => (~para(sK23,sK22,sK21,sK20) & ! [X4,X5,X6,X7,X8,X9] : (eqangle(X5,X6,X6,X9,X6,X9,X9,X5) & eqangle(X5,X6,X6,X9,X6,X5,X5,X9) & eqangle(X4,X6,X6,X8,X6,X8,X8,X4) & eqangle(X4,X6,X6,X8,X6,X4,X4,X8) & eqangle(X5,X4,X4,X7,X4,X7,X7,X5) & eqangle(X5,X4,X4,X7,X4,X5,X5,X7))) [choice axiom]
353. ~para(sK23,sK22,sK21,sK20) & ! [X4,X5,X6,X7,X8,X9] : (eqangle(X5,X6,X6,X9,X6,X9,X9,X5) & eqangle(X5,X6,X6,X9,X6,X5,X5,X9) & eqangle(X4,X6,X6,X8,X6,X8,X8,X4) & eqangle(X4,X6,X6,X8,X6,X4,X4,X8) & eqangle(X5,X4,X4,X7,X4,X7,X7,X5) & eqangle(X5,X4,X4,X7,X4,X5,X5,X7)) [skolemisation 315,352]
357. ~para(X0,X1,X2,X3) | para(X0,X1,X3,X2) [cnf transformation 170]
358. ~para(X0,X1,X2,X3) | para(X2,X3,X0,X1) [cnf transformation 171]
359. ~para(X2,X3,X4,X5) | para(X0,X1,X4,X5) | ~para(X0,X1,X2,X3) [cnf transformation 173]
371. ~eqangle(X0,X1,X2,X3,X4,X5,X6,X7) | eqangle(X1,X0,X2,X3,X4,X5,X6,X7) [cnf transformation 190]
372. ~eqangle(X0,X1,X2,X3,X4,X5,X6,X7) | eqangle(X2,X3,X0,X1,X6,X7,X4,X5) [cnf transformation 191]
374. ~eqangle(X0,X1,X2,X3,X4,X5,X6,X7) | eqangle(X0,X1,X4,X5,X2,X3,X6,X7) [cnf transformation 193]
375. ~eqangle(X4,X5,X6,X7,X8,X9,X10,X11) | eqangle(X0,X1,X2,X3,X8,X9,X10,X11) | ~eqangle(X0,X1,X2,X3,X4,X5,X6,X7) [cnf transformation 195]
392. ~eqangle(X0,X1,X4,X5,X2,X3,X4,X5) | para(X0,X1,X2,X3) [cnf transformation 216]
427. ~eqangle(X0,X1,X2,X3,X4,X5,X6,X7) | ~para(X4,X5,X6,X7) | para(X0,X1,X2,X3) [cnf transformation 275]
471. eqangle(X5,X6,X6,X9,X6,X5,X5,X9) [cnf transformation 353]
472. eqangle(X5,X6,X6,X9,X6,X9,X9,X5) [cnf transformation 353]
473. ~para(sK23,sK22,sK21,sK20) [cnf transformation 353]
616. eqangle(X0,X1,X0,X2,X0,X1,X1,X2) [resolution 371,471]
624. eqangle(X0,X1,X2,X0,X2,X1,X0,X2) [resolution 372,471]
630. eqangle(X25,X26,X25,X27,X27,X26,X25,X27) [resolution 372,616]
651. eqangle(X21,X22,X23,X22,X23,X21,X21,X23) [resolution 374,624]
796. para(X0,X1,X2,X1) [resolution 630,392]
813. para(X18,X19,X19,X20) [resolution 796,357]
824. para(X15,X16,X17,X15) [resolution 813,358]
834. para(X18,X19,X18,X20) [resolution 824,357]
1563. ~eqangle(X21,X22,X23,X24,X26,X25,X25,X27) | eqangle(X21,X22,X23,X24,X25,X26,X26,X27) [resolution 375,471]
1564. ~eqangle(X28,X29,X30,X31,X34,X32,X32,X33) | eqangle(X28,X29,X30,X31,X32,X33,X33,X34) [resolution 375,472]
3323. eqangle(X135,X136,X137,X136,X135,X137,X137,X137) [resolution 1563,651]
3771. eqangle(X27,X28,X29,X28,X29,X29,X29,X27) [resolution 3323,1564]
4796. eqangle(X11,X12,X13,X13,X13,X12,X13,X11) [resolution 3771,374]
7876. ~para(X41,X42,X41,X43) | para(X43,X42,X41,X41) [resolution 4796,427]
7896. para(X43,X42,X41,X41) [subsumption resolution 7876,834]
7911. para(X18,X18,X19,X20) [resolution 7896,358]
7928. para(X13,X14,X15,X16) | ~para(X13,X14,X17,X17) [resolution 7911,359]
7936. para(X13,X14,X15,X16) [subsumption resolution 7928,7896]
7953. $false [resolution 7936,473]
% SZS output end Proof for E051-32
% ------------------------------
% Version: Vampire 4.4.0 (commit 7916d27 on 2019-08-23 08:50:16 +0100)
% Termination reason: Refutation

% Memory used [KB]: 7164
% Time elapsed: 0.120 s
% ------------------------------
% ------------------------------
