% Refutation found. Thanks to Tanya!
% SZS status Theorem for fig208
% SZS output start Proof for fig208
9. ! [X0,X1,X2,X3,X4,X5] : ((perp(X2,X3,X4,X5) & perp(X0,X1,X2,X3)) => para(X0,X1,X4,X5)) [input]
53. ! [X0,X1,X2,X7] : ((coll(X7,X0,X2) & circle(X7,X0,X1,X2)) => perp(X0,X1,X1,X2)) [input]
94. ! [X0,X1,X2,X3,X4,X5,X12,X13,X18,X19,X20,X21,X22,X23] : (circle(X0,X2,X13,X23) & coll(X2,X5,X13) & circle(X0,X3,X12,X22) & coll(X3,X4,X12) & circle(X1,X2,X5,X21) & circle(X1,X2,X4,X20) & circle(X0,X2,X3,X19) & circle(X1,X2,X3,X18)) => para(X12,X13,X4,X5) [input]
95. ~! [X5,X4,X13,X12] : (! [X0,X1,X2,X3,X4,X5,X12,X13,X18,X19,X20,X21,X22,X23] : (circle(X0,X2,X13,X23) & coll(X2,X5,X13) & circle(X0,X3,X12,X22) & coll(X3,X4,X12) & circle(X1,X2,X5,X21) & circle(X1,X2,X4,X20) & circle(X0,X2,X3,X19) & circle(X1,X2,X3,X18)) => para(X12,X13,X4,X5)) [negated conjecture 94]
131. ! [X0,X1,X2,X3] : ((coll(X3,X0,X2) & circle(X3,X0,X1,X2)) => perp(X0,X1,X1,X2)) [rectify 53]
165. ~! [X0,X1,X2,X3] : (! [X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17] : (circle(X4,X6,X11,X17) & coll(X6,X9,X11) & circle(X4,X7,X10,X16) & coll(X7,X8,X10) & circle(X5,X6,X9,X15) & circle(X5,X6,X8,X14) & circle(X4,X6,X7,X13) & circle(X5,X6,X7,X12)) => para(X3,X2,X1,X0)) [rectify 95]
176. ! [X0,X1,X2,X3,X4,X5] : (para(X0,X1,X4,X5) | (~perp(X2,X3,X4,X5) | ~perp(X0,X1,X2,X3))) [ennf transformation 9]
177. ! [X0,X1,X2,X3,X4,X5] : (para(X0,X1,X4,X5) | ~perp(X2,X3,X4,X5) | ~perp(X0,X1,X2,X3)) [flattening 176]
240. ! [X0,X1,X2,X3] : (perp(X0,X1,X1,X2) | (~coll(X3,X0,X2) | ~circle(X3,X0,X1,X2))) [ennf transformation 131]
241. ! [X0,X1,X2,X3] : (perp(X0,X1,X1,X2) | ~coll(X3,X0,X2) | ~circle(X3,X0,X1,X2)) [flattening 240]
315. ? [X0,X1,X2,X3] : (~para(X3,X2,X1,X0) & ! [X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17] : (circle(X4,X6,X11,X17) & coll(X6,X9,X11) & circle(X4,X7,X10,X16) & coll(X7,X8,X10) & circle(X5,X6,X9,X15) & circle(X5,X6,X8,X14) & circle(X4,X6,X7,X13) & circle(X5,X6,X7,X12))) [ennf transformation 165]
352. ? [X0,X1,X2,X3] : (~para(X3,X2,X1,X0) & ! [X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17] : (circle(X4,X6,X11,X17) & coll(X6,X9,X11) & circle(X4,X7,X10,X16) & coll(X7,X8,X10) & circle(X5,X6,X9,X15) & circle(X5,X6,X8,X14) & circle(X4,X6,X7,X13) & circle(X5,X6,X7,X12))) => (~para(sK23,sK22,sK21,sK20) & ! [X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17] : (circle(X4,X6,X11,X17) & coll(X6,X9,X11) & circle(X4,X7,X10,X16) & coll(X7,X8,X10) & circle(X5,X6,X9,X15) & circle(X5,X6,X8,X14) & circle(X4,X6,X7,X13) & circle(X5,X6,X7,X12))) [choice axiom]
353. ~para(sK23,sK22,sK21,sK20) & ! [X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17] : (circle(X4,X6,X11,X17) & coll(X6,X9,X11) & circle(X4,X7,X10,X16) & coll(X7,X8,X10) & circle(X5,X6,X9,X15) & circle(X5,X6,X8,X14) & circle(X4,X6,X7,X13) & circle(X5,X6,X7,X12)) [skolemisation 315,352]
362. ~perp(X2,X3,X4,X5) | para(X0,X1,X4,X5) | ~perp(X0,X1,X2,X3) [cnf transformation 177]
407. ~circle(X3,X0,X1,X2) | ~coll(X3,X0,X2) | perp(X0,X1,X1,X2) [cnf transformation 241]
473. coll(X6,X9,X11) [cnf transformation 353]
474. circle(X4,X6,X11,X17) [cnf transformation 353]
475. ~para(sK23,sK22,sK21,sK20) [cnf transformation 353]
518. ~coll(X0,X1,X2) | perp(X1,X3,X3,X2) [resolution 407,474]
519. perp(X1,X3,X3,X2) [subsumption resolution 518,473]
541. ~perp(X0,X1,X4,X2) | para(X0,X1,X2,X3) [resolution 362,519]
631. para(X0,X1,X2,X3) [resolution 541,519]
641. $false [resolution 631,475]
% SZS output end Proof for fig208
% ------------------------------
% Version: Vampire 4.4.0 (commit 7916d27 on 2019-08-23 08:50:16 +0100)
% Termination reason: Refutation

% Memory used [KB]: 5117
% Time elapsed: 0.029 s
% ------------------------------
% ------------------------------
