% Refutation found. Thanks to Tanya!
% SZS status Theorem for fig208
% SZS output start Proof for fig208
fof(f645,plain,(
  $false),
  inference(resolution,[],[f635,f479])).
fof(f479,plain,(
  ~para(sK23,sK22,sK21,sK20)),
  inference(cnf_transformation,[],[f357])).
fof(f357,plain,(
  ~para(sK23,sK22,sK21,sK20) & ! [X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17] : (circle(X4,X6,X11,X17) & coll(X6,X9,X11) & circle(X4,X7,X10,X16) & coll(X7,X8,X10) & circle(X5,X6,X9,X15) & circle(X5,X6,X8,X14) & circle(X4,X6,X7,X13) & circle(X5,X6,X7,X12))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK20,sK21,sK22,sK23])],[f319,f356])).
fof(f356,plain,(
  ? [X0,X1,X2,X3] : (~para(X3,X2,X1,X0) & ! [X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17] : (circle(X4,X6,X11,X17) & coll(X6,X9,X11) & circle(X4,X7,X10,X16) & coll(X7,X8,X10) & circle(X5,X6,X9,X15) & circle(X5,X6,X8,X14) & circle(X4,X6,X7,X13) & circle(X5,X6,X7,X12))) => (~para(sK23,sK22,sK21,sK20) & ! [X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17] : (circle(X4,X6,X11,X17) & coll(X6,X9,X11) & circle(X4,X7,X10,X16) & coll(X7,X8,X10) & circle(X5,X6,X9,X15) & circle(X5,X6,X8,X14) & circle(X4,X6,X7,X13) & circle(X5,X6,X7,X12)))),
  introduced(choice_axiom,[])).
fof(f319,plain,(
  ? [X0,X1,X2,X3] : (~para(X3,X2,X1,X0) & ! [X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17] : (circle(X4,X6,X11,X17) & coll(X6,X9,X11) & circle(X4,X7,X10,X16) & coll(X7,X8,X10) & circle(X5,X6,X9,X15) & circle(X5,X6,X8,X14) & circle(X4,X6,X7,X13) & circle(X5,X6,X7,X12)))),
  inference(ennf_transformation,[],[f167])).
fof(f167,plain,(
  ~! [X0,X1,X2,X3] : (! [X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17] : (circle(X4,X6,X11,X17) & coll(X6,X9,X11) & circle(X4,X7,X10,X16) & coll(X7,X8,X10) & circle(X5,X6,X9,X15) & circle(X5,X6,X8,X14) & circle(X4,X6,X7,X13) & circle(X5,X6,X7,X12)) => para(X3,X2,X1,X0))),
  inference(rectify,[],[f96])).
fof(f96,negated_conjecture,(
  ~! [X5,X4,X13,X12] : (! [X0,X1,X2,X3,X4,X5,X12,X13,X18,X19,X20,X21,X22,X23] : (circle(X0,X2,X13,X23) & coll(X2,X5,X13) & circle(X0,X3,X12,X22) & coll(X3,X4,X12) & circle(X1,X2,X5,X21) & circle(X1,X2,X4,X20) & circle(X0,X2,X3,X19) & circle(X1,X2,X3,X18)) => para(X12,X13,X4,X5))),
  inference(negated_conjecture,[],[f95])).
fof(f95,conjecture,(
  ( ! [X4,X12,X5,X13] : (! [X0,X1,X2,X3,X4,X5,X12,X13,X18,X19,X20,X21,X22,X23] : (circle(X0,X2,X13,X23) & coll(X2,X5,X13) & circle(X0,X3,X12,X22) & coll(X3,X4,X12) & circle(X1,X2,X5,X21) & circle(X1,X2,X4,X20) & circle(X0,X2,X3,X19) & circle(X1,X2,X3,X18)) => para(X12,X13,X4,X5)) )),
  file('./3_JAR/JAR02new/fig208.p',exemplo3JARJAR02newfig208)).
fof(f635,plain,(
  ( ! [X2,X0,X3,X1] : (para(X0,X1,X2,X3)) )),
  inference(resolution,[],[f545,f523])).
fof(f523,plain,(
  ( ! [X2,X3,X1] : (perp(X1,X3,X3,X2)) )),
  inference(subsumption_resolution,[],[f522,f477])).
fof(f477,plain,(
  ( ! [X6,X11,X9] : (coll(X6,X9,X11)) )),
  inference(cnf_transformation,[],[f357])).
fof(f522,plain,(
  ( ! [X2,X0,X3,X1] : (~coll(X0,X1,X2) | perp(X1,X3,X3,X2)) )),
  inference(resolution,[],[f411,f478])).
fof(f478,plain,(
  ( ! [X6,X4,X17,X11] : (circle(X4,X6,X11,X17)) )),
  inference(cnf_transformation,[],[f357])).
fof(f411,plain,(
  ( ! [X2,X0,X3,X1] : (~circle(X3,X0,X1,X2) | ~coll(X3,X0,X2) | perp(X0,X1,X1,X2)) )),
  inference(cnf_transformation,[],[f245])).
fof(f245,plain,(
  ! [X0,X1,X2,X3] : (perp(X0,X1,X1,X2) | ~coll(X3,X0,X2) | ~circle(X3,X0,X1,X2))),
  inference(flattening,[],[f244])).
fof(f244,plain,(
  ! [X0,X1,X2,X3] : (perp(X0,X1,X1,X2) | (~coll(X3,X0,X2) | ~circle(X3,X0,X1,X2)))),
  inference(ennf_transformation,[],[f133])).
fof(f133,plain,(
  ! [X0,X1,X2,X3] : ((coll(X3,X0,X2) & circle(X3,X0,X1,X2)) => perp(X0,X1,X1,X2))),
  inference(rectify,[],[f54])).
fof(f54,axiom,(
  ! [X0,X1,X2,X7] : ((coll(X7,X0,X2) & circle(X7,X0,X1,X2)) => perp(X0,X1,X1,X2))),
  file('./3_JAR/JAR02new/fig208.p',ruleD53)).
fof(f545,plain,(
  ( ! [X4,X2,X0,X3,X1] : (~perp(X0,X1,X4,X2) | para(X0,X1,X2,X3)) )),
  inference(resolution,[],[f366,f523])).
fof(f366,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (~perp(X2,X3,X4,X5) | para(X0,X1,X4,X5) | ~perp(X0,X1,X2,X3)) )),
  inference(cnf_transformation,[],[f179])).
fof(f179,plain,(
  ! [X0,X1,X2,X3,X4,X5] : (para(X0,X1,X4,X5) | ~perp(X2,X3,X4,X5) | ~perp(X0,X1,X2,X3))),
  inference(flattening,[],[f178])).
fof(f178,plain,(
  ! [X0,X1,X2,X3,X4,X5] : (para(X0,X1,X4,X5) | (~perp(X2,X3,X4,X5) | ~perp(X0,X1,X2,X3)))),
  inference(ennf_transformation,[],[f9])).
fof(f9,axiom,(
  ! [X0,X1,X2,X3,X4,X5] : ((perp(X2,X3,X4,X5) & perp(X0,X1,X2,X3)) => para(X0,X1,X4,X5))),
  file('./3_JAR/JAR02new/fig208.p',ruleD9)).
% SZS output end Proof for fig208
% ------------------------------
% Version: Vampire 4.4.0 (commit 7916d27 on 2019-08-23 08:50:16 +0100)
% Termination reason: Refutation

% Memory used [KB]: 5117
% Time elapsed: 0.030 s
% ------------------------------
% ------------------------------
