% Refutation found. Thanks to Tanya!
% SZS status Theorem for 5cirdb
% SZS output start Proof for 5cirdb
fof(f647450,plain,(
  $false),
  inference(avatar_sat_refutation,[],[f3932,f7979,f647447])).
fof(f647447,plain,(
  ~spl65_464),
  inference(avatar_contradiction_clause,[],[f647446])).
fof(f647446,plain,(
  $false | ~spl65_464),
  inference(resolution,[],[f38792,f523])).
fof(f523,plain,(
  ~cyclic(sK50,sK51,sK52,sK54)),
  inference(cnf_transformation,[],[f359])).
fof(f359,plain,(
  ~cyclic(sK50,sK51,sK52,sK54) & circle(sK45,sK44,sK54,sK64) & circle(sK49,sK44,sK54,sK63) & circle(sK49,sK43,sK53,sK62) & circle(sK48,sK43,sK53,sK61) & circle(sK48,sK42,sK52,sK60) & circle(sK47,sK42,sK52,sK59) & circle(sK47,sK41,sK51,sK58) & circle(sK46,sK41,sK51,sK57) & circle(sK46,sK40,sK50,sK56) & circle(sK45,sK40,sK50,sK55) & perp(sK43,sK44,sK20,sK49) & midp(sK20,sK43,sK44) & perp(sK39,sK44,sK21,sK49) & midp(sK21,sK39,sK44) & perp(sK39,sK43,sK22,sK49) & midp(sK22,sK39,sK43) & perp(sK42,sK43,sK23,sK48) & midp(sK23,sK42,sK43) & perp(sK38,sK43,sK24,sK48) & midp(sK24,sK38,sK43) & perp(sK38,sK42,sK25,sK48) & midp(sK25,sK38,sK42) & perp(sK41,sK42,sK26,sK47) & midp(sK26,sK41,sK42) & perp(sK37,sK42,sK27,sK47) & midp(sK27,sK37,sK42) & perp(sK37,sK41,sK28,sK47) & midp(sK28,sK37,sK41) & perp(sK40,sK41,sK29,sK46) & midp(sK29,sK40,sK41) & perp(sK36,sK41,sK30,sK46) & midp(sK30,sK36,sK41) & perp(sK36,sK40,sK31,sK46) & midp(sK31,sK36,sK40) & perp(sK40,sK44,sK32,sK45) & midp(sK32,sK40,sK44) & perp(sK35,sK44,sK33,sK45) & midp(sK33,sK35,sK44) & perp(sK35,sK40,sK34,sK45) & midp(sK34,sK35,sK40) & coll(sK44,sK36,sK39) & coll(sK44,sK35,sK38) & coll(sK43,sK37,sK39) & coll(sK43,sK35,sK38) & coll(sK42,sK37,sK39) & coll(sK42,sK36,sK38) & coll(sK41,sK36,sK38) & coll(sK41,sK35,sK37) & coll(sK40,sK36,sK39) & coll(sK40,sK35,sK37)),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK20,sK21,sK22,sK23,sK24,sK25,sK26,sK27,sK28,sK29,sK30,sK31,sK32,sK33,sK34,sK35,sK36,sK37,sK38,sK39,sK40,sK41,sK42,sK43,sK44,sK45,sK46,sK47,sK48,sK49,sK50,sK51,sK52,sK53,sK54,sK55,sK56,sK57,sK58,sK59,sK60,sK61,sK62,sK63,sK64])],[f321,f358])).
fof(f358,plain,(
  ? [X0,X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26,X27,X28,X29,X30,X31,X32,X33,X34,X35,X36,X37,X38,X39,X40,X41,X42,X43,X44] : (~cyclic(X30,X31,X32,X34) & circle(X25,X24,X34,X44) & circle(X29,X24,X34,X43) & circle(X29,X23,X33,X42) & circle(X28,X23,X33,X41) & circle(X28,X22,X32,X40) & circle(X27,X22,X32,X39) & circle(X27,X21,X31,X38) & circle(X26,X21,X31,X37) & circle(X26,X20,X30,X36) & circle(X25,X20,X30,X35) & perp(X23,X24,X0,X29) & midp(X0,X23,X24) & perp(X19,X24,X1,X29) & midp(X1,X19,X24) & perp(X19,X23,X2,X29) & midp(X2,X19,X23) & perp(X22,X23,X3,X28) & midp(X3,X22,X23) & perp(X18,X23,X4,X28) & midp(X4,X18,X23) & perp(X18,X22,X5,X28) & midp(X5,X18,X22) & perp(X21,X22,X6,X27) & midp(X6,X21,X22) & perp(X17,X22,X7,X27) & midp(X7,X17,X22) & perp(X17,X21,X8,X27) & midp(X8,X17,X21) & perp(X20,X21,X9,X26) & midp(X9,X20,X21) & perp(X16,X21,X10,X26) & midp(X10,X16,X21) & perp(X16,X20,X11,X26) & midp(X11,X16,X20) & perp(X20,X24,X12,X25) & midp(X12,X20,X24) & perp(X15,X24,X13,X25) & midp(X13,X15,X24) & perp(X15,X20,X14,X25) & midp(X14,X15,X20) & coll(X24,X16,X19) & coll(X24,X15,X18) & coll(X23,X17,X19) & coll(X23,X15,X18) & coll(X22,X17,X19) & coll(X22,X16,X18) & coll(X21,X16,X18) & coll(X21,X15,X17) & coll(X20,X16,X19) & coll(X20,X15,X17)) => (~cyclic(sK50,sK51,sK52,sK54) & circle(sK45,sK44,sK54,sK64) & circle(sK49,sK44,sK54,sK63) & circle(sK49,sK43,sK53,sK62) & circle(sK48,sK43,sK53,sK61) & circle(sK48,sK42,sK52,sK60) & circle(sK47,sK42,sK52,sK59) & circle(sK47,sK41,sK51,sK58) & circle(sK46,sK41,sK51,sK57) & circle(sK46,sK40,sK50,sK56) & circle(sK45,sK40,sK50,sK55) & perp(sK43,sK44,sK20,sK49) & midp(sK20,sK43,sK44) & perp(sK39,sK44,sK21,sK49) & midp(sK21,sK39,sK44) & perp(sK39,sK43,sK22,sK49) & midp(sK22,sK39,sK43) & perp(sK42,sK43,sK23,sK48) & midp(sK23,sK42,sK43) & perp(sK38,sK43,sK24,sK48) & midp(sK24,sK38,sK43) & perp(sK38,sK42,sK25,sK48) & midp(sK25,sK38,sK42) & perp(sK41,sK42,sK26,sK47) & midp(sK26,sK41,sK42) & perp(sK37,sK42,sK27,sK47) & midp(sK27,sK37,sK42) & perp(sK37,sK41,sK28,sK47) & midp(sK28,sK37,sK41) & perp(sK40,sK41,sK29,sK46) & midp(sK29,sK40,sK41) & perp(sK36,sK41,sK30,sK46) & midp(sK30,sK36,sK41) & perp(sK36,sK40,sK31,sK46) & midp(sK31,sK36,sK40) & perp(sK40,sK44,sK32,sK45) & midp(sK32,sK40,sK44) & perp(sK35,sK44,sK33,sK45) & midp(sK33,sK35,sK44) & perp(sK35,sK40,sK34,sK45) & midp(sK34,sK35,sK40) & coll(sK44,sK36,sK39) & coll(sK44,sK35,sK38) & coll(sK43,sK37,sK39) & coll(sK43,sK35,sK38) & coll(sK42,sK37,sK39) & coll(sK42,sK36,sK38) & coll(sK41,sK36,sK38) & coll(sK41,sK35,sK37) & coll(sK40,sK36,sK39) & coll(sK40,sK35,sK37))),
  introduced(choice_axiom,[])).
fof(f321,plain,(
  ? [X0,X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26,X27,X28,X29,X30,X31,X32,X33,X34,X35,X36,X37,X38,X39,X40,X41,X42,X43,X44] : (~cyclic(X30,X31,X32,X34) & circle(X25,X24,X34,X44) & circle(X29,X24,X34,X43) & circle(X29,X23,X33,X42) & circle(X28,X23,X33,X41) & circle(X28,X22,X32,X40) & circle(X27,X22,X32,X39) & circle(X27,X21,X31,X38) & circle(X26,X21,X31,X37) & circle(X26,X20,X30,X36) & circle(X25,X20,X30,X35) & perp(X23,X24,X0,X29) & midp(X0,X23,X24) & perp(X19,X24,X1,X29) & midp(X1,X19,X24) & perp(X19,X23,X2,X29) & midp(X2,X19,X23) & perp(X22,X23,X3,X28) & midp(X3,X22,X23) & perp(X18,X23,X4,X28) & midp(X4,X18,X23) & perp(X18,X22,X5,X28) & midp(X5,X18,X22) & perp(X21,X22,X6,X27) & midp(X6,X21,X22) & perp(X17,X22,X7,X27) & midp(X7,X17,X22) & perp(X17,X21,X8,X27) & midp(X8,X17,X21) & perp(X20,X21,X9,X26) & midp(X9,X20,X21) & perp(X16,X21,X10,X26) & midp(X10,X16,X21) & perp(X16,X20,X11,X26) & midp(X11,X16,X20) & perp(X20,X24,X12,X25) & midp(X12,X20,X24) & perp(X15,X24,X13,X25) & midp(X13,X15,X24) & perp(X15,X20,X14,X25) & midp(X14,X15,X20) & coll(X24,X16,X19) & coll(X24,X15,X18) & coll(X23,X17,X19) & coll(X23,X15,X18) & coll(X22,X17,X19) & coll(X22,X16,X18) & coll(X21,X16,X18) & coll(X21,X15,X17) & coll(X20,X16,X19) & coll(X20,X15,X17))),
  inference(flattening,[],[f320])).
fof(f320,plain,(
  ? [X0,X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26,X27,X28,X29,X30,X31,X32,X33,X34,X35,X36,X37,X38,X39,X40,X41,X42,X43,X44] : (~cyclic(X30,X31,X32,X34) & (circle(X25,X24,X34,X44) & circle(X29,X24,X34,X43) & circle(X29,X23,X33,X42) & circle(X28,X23,X33,X41) & circle(X28,X22,X32,X40) & circle(X27,X22,X32,X39) & circle(X27,X21,X31,X38) & circle(X26,X21,X31,X37) & circle(X26,X20,X30,X36) & circle(X25,X20,X30,X35) & perp(X23,X24,X0,X29) & midp(X0,X23,X24) & perp(X19,X24,X1,X29) & midp(X1,X19,X24) & perp(X19,X23,X2,X29) & midp(X2,X19,X23) & perp(X22,X23,X3,X28) & midp(X3,X22,X23) & perp(X18,X23,X4,X28) & midp(X4,X18,X23) & perp(X18,X22,X5,X28) & midp(X5,X18,X22) & perp(X21,X22,X6,X27) & midp(X6,X21,X22) & perp(X17,X22,X7,X27) & midp(X7,X17,X22) & perp(X17,X21,X8,X27) & midp(X8,X17,X21) & perp(X20,X21,X9,X26) & midp(X9,X20,X21) & perp(X16,X21,X10,X26) & midp(X10,X16,X21) & perp(X16,X20,X11,X26) & midp(X11,X16,X20) & perp(X20,X24,X12,X25) & midp(X12,X20,X24) & perp(X15,X24,X13,X25) & midp(X13,X15,X24) & perp(X15,X20,X14,X25) & midp(X14,X15,X20) & coll(X24,X16,X19) & coll(X24,X15,X18) & coll(X23,X17,X19) & coll(X23,X15,X18) & coll(X22,X17,X19) & coll(X22,X16,X18) & coll(X21,X16,X18) & coll(X21,X15,X17) & coll(X20,X16,X19) & coll(X20,X15,X17)))),
  inference(ennf_transformation,[],[f168])).
fof(f168,plain,(
  ~! [X0,X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26,X27,X28,X29,X30,X31,X32,X33,X34,X35,X36,X37,X38,X39,X40,X41,X42,X43,X44] : ((circle(X25,X24,X34,X44) & circle(X29,X24,X34,X43) & circle(X29,X23,X33,X42) & circle(X28,X23,X33,X41) & circle(X28,X22,X32,X40) & circle(X27,X22,X32,X39) & circle(X27,X21,X31,X38) & circle(X26,X21,X31,X37) & circle(X26,X20,X30,X36) & circle(X25,X20,X30,X35) & perp(X23,X24,X0,X29) & midp(X0,X23,X24) & perp(X19,X24,X1,X29) & midp(X1,X19,X24) & perp(X19,X23,X2,X29) & midp(X2,X19,X23) & perp(X22,X23,X3,X28) & midp(X3,X22,X23) & perp(X18,X23,X4,X28) & midp(X4,X18,X23) & perp(X18,X22,X5,X28) & midp(X5,X18,X22) & perp(X21,X22,X6,X27) & midp(X6,X21,X22) & perp(X17,X22,X7,X27) & midp(X7,X17,X22) & perp(X17,X21,X8,X27) & midp(X8,X17,X21) & perp(X20,X21,X9,X26) & midp(X9,X20,X21) & perp(X16,X21,X10,X26) & midp(X10,X16,X21) & perp(X16,X20,X11,X26) & midp(X11,X16,X20) & perp(X20,X24,X12,X25) & midp(X12,X20,X24) & perp(X15,X24,X13,X25) & midp(X13,X15,X24) & perp(X15,X20,X14,X25) & midp(X14,X15,X20) & coll(X24,X16,X19) & coll(X24,X15,X18) & coll(X23,X17,X19) & coll(X23,X15,X18) & coll(X22,X17,X19) & coll(X22,X16,X18) & coll(X21,X16,X18) & coll(X21,X15,X17) & coll(X20,X16,X19) & coll(X20,X15,X17)) => cyclic(X30,X31,X32,X34))),
  inference(flattening,[],[f167])).
fof(f167,plain,(
  ~! [X0,X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14] : ! [X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25,X26,X27,X28,X29,X30,X31,X32,X33,X34,X35,X36,X37,X38,X39,X40,X41,X42,X43,X44] : ((circle(X25,X24,X34,X44) & circle(X29,X24,X34,X43) & circle(X29,X23,X33,X42) & circle(X28,X23,X33,X41) & circle(X28,X22,X32,X40) & circle(X27,X22,X32,X39) & circle(X27,X21,X31,X38) & circle(X26,X21,X31,X37) & circle(X26,X20,X30,X36) & circle(X25,X20,X30,X35) & perp(X23,X24,X0,X29) & midp(X0,X23,X24) & perp(X19,X24,X1,X29) & midp(X1,X19,X24) & perp(X19,X23,X2,X29) & midp(X2,X19,X23) & perp(X22,X23,X3,X28) & midp(X3,X22,X23) & perp(X18,X23,X4,X28) & midp(X4,X18,X23) & perp(X18,X22,X5,X28) & midp(X5,X18,X22) & perp(X21,X22,X6,X27) & midp(X6,X21,X22) & perp(X17,X22,X7,X27) & midp(X7,X17,X22) & perp(X17,X21,X8,X27) & midp(X8,X17,X21) & perp(X20,X21,X9,X26) & midp(X9,X20,X21) & perp(X16,X21,X10,X26) & midp(X10,X16,X21) & perp(X16,X20,X11,X26) & midp(X11,X16,X20) & perp(X20,X24,X12,X25) & midp(X12,X20,X24) & perp(X15,X24,X13,X25) & midp(X13,X15,X24) & perp(X15,X20,X14,X25) & midp(X14,X15,X20) & coll(X24,X16,X19) & coll(X24,X15,X18) & coll(X23,X17,X19) & coll(X23,X15,X18) & coll(X22,X17,X19) & coll(X22,X16,X18) & coll(X21,X16,X18) & coll(X21,X15,X17) & coll(X20,X16,X19) & coll(X20,X15,X17)) => cyclic(X30,X31,X32,X34))),
  inference(rectify,[],[f96])).
fof(f96,negated_conjecture,(
  ~! [X62,X61,X60,X59,X58,X57,X56,X55,X54,X53,X52,X51,X50,X49,X48] : ! [X18,X19,X20,X21,X22,X23,X24,X25,X26,X27,X28,X29,X30,X31,X32,X33,X34,X35,X36,X37,X38,X39,X40,X41,X42,X43,X44,X45,X46,X47] : ((circle(X28,X27,X37,X47) & circle(X32,X27,X37,X46) & circle(X32,X26,X36,X45) & circle(X31,X26,X36,X44) & circle(X31,X25,X35,X43) & circle(X30,X25,X35,X42) & circle(X30,X24,X34,X41) & circle(X29,X24,X34,X40) & circle(X29,X23,X33,X39) & circle(X28,X23,X33,X38) & perp(X26,X27,X62,X32) & midp(X62,X26,X27) & perp(X22,X27,X61,X32) & midp(X61,X22,X27) & perp(X22,X26,X60,X32) & midp(X60,X22,X26) & perp(X25,X26,X59,X31) & midp(X59,X25,X26) & perp(X21,X26,X58,X31) & midp(X58,X21,X26) & perp(X21,X25,X57,X31) & midp(X57,X21,X25) & perp(X24,X25,X56,X30) & midp(X56,X24,X25) & perp(X20,X25,X55,X30) & midp(X55,X20,X25) & perp(X20,X24,X54,X30) & midp(X54,X20,X24) & perp(X23,X24,X53,X29) & midp(X53,X23,X24) & perp(X19,X24,X52,X29) & midp(X52,X19,X24) & perp(X19,X23,X51,X29) & midp(X51,X19,X23) & perp(X23,X27,X50,X28) & midp(X50,X23,X27) & perp(X18,X27,X49,X28) & midp(X49,X18,X27) & perp(X18,X23,X48,X28) & midp(X48,X18,X23) & coll(X27,X19,X22) & coll(X27,X18,X21) & coll(X26,X20,X22) & coll(X26,X18,X21) & coll(X25,X20,X22) & coll(X25,X19,X21) & coll(X24,X19,X21) & coll(X24,X18,X20) & coll(X23,X19,X22) & coll(X23,X18,X20)) => cyclic(X33,X34,X35,X37))),
  inference(negated_conjecture,[],[f95])).
fof(f95,conjecture,(
  ( ! [X61,X59,X57,X54,X52,X50,X48,X62,X60,X58,X56,X55,X53,X51,X49] : (! [X18,X19,X20,X21,X22,X23,X24,X25,X26,X27,X28,X29,X30,X31,X32,X33,X34,X35,X36,X37,X38,X39,X40,X41,X42,X43,X44,X45,X46,X47] : ((circle(X28,X27,X37,X47) & circle(X32,X27,X37,X46) & circle(X32,X26,X36,X45) & circle(X31,X26,X36,X44) & circle(X31,X25,X35,X43) & circle(X30,X25,X35,X42) & circle(X30,X24,X34,X41) & circle(X29,X24,X34,X40) & circle(X29,X23,X33,X39) & circle(X28,X23,X33,X38) & perp(X26,X27,X62,X32) & midp(X62,X26,X27) & perp(X22,X27,X61,X32) & midp(X61,X22,X27) & perp(X22,X26,X60,X32) & midp(X60,X22,X26) & perp(X25,X26,X59,X31) & midp(X59,X25,X26) & perp(X21,X26,X58,X31) & midp(X58,X21,X26) & perp(X21,X25,X57,X31) & midp(X57,X21,X25) & perp(X24,X25,X56,X30) & midp(X56,X24,X25) & perp(X20,X25,X55,X30) & midp(X55,X20,X25) & perp(X20,X24,X54,X30) & midp(X54,X20,X24) & perp(X23,X24,X53,X29) & midp(X53,X23,X24) & perp(X19,X24,X52,X29) & midp(X52,X19,X24) & perp(X19,X23,X51,X29) & midp(X51,X19,X23) & perp(X23,X27,X50,X28) & midp(X50,X23,X27) & perp(X18,X27,X49,X28) & midp(X49,X18,X27) & perp(X18,X23,X48,X28) & midp(X48,X18,X23) & coll(X27,X19,X22) & coll(X27,X18,X21) & coll(X26,X20,X22) & coll(X26,X18,X21) & coll(X25,X20,X22) & coll(X25,X19,X21) & coll(X24,X19,X21) & coll(X24,X18,X20) & coll(X23,X19,X22) & coll(X23,X18,X20)) => cyclic(X33,X34,X35,X37))) )),
  file('./3_JAR/5cirdb.p',exemplo3JAR5cirdb)).
fof(f38792,plain,(
  ( ! [X12,X10,X11,X9] : (cyclic(X9,X10,X11,X12)) ) | ~spl65_464),
  inference(subsumption_resolution,[],[f38756,f20628])).
fof(f20628,plain,(
  ( ! [X2,X0,X1] : (cyclic(X0,X0,X1,X2)) ) | ~spl65_464),
  inference(resolution,[],[f7932,f3928])).
fof(f3928,plain,(
  ( ! [X2,X3] : (para(X2,X3,X2,X3)) ) | ~spl65_464),
  inference(avatar_component_clause,[],[f3927])).
fof(f3927,plain,(
  spl65_464 <=> ! [X3,X2] : para(X2,X3,X2,X3)),
  introduced(avatar_definition,[new_symbols(naming,[spl65_464])])).
fof(f7932,plain,(
  ( ! [X68,X66,X67] : (~para(X66,X67,X66,X67) | cyclic(X67,X67,X66,X68)) )),
  inference(subsumption_resolution,[],[f4087,f4088])).
fof(f4088,plain,(
  ( ! [X70,X71,X69] : (~para(X69,X70,X69,X70) | coll(X69,X71,X70) | cyclic(X70,X70,X69,X71)) )),
  inference(resolution,[],[f1217,f401])).
fof(f401,plain,(
  ( ! [X2,X0,X3,X1] : (~eqangle(X2,X0,X2,X1,X3,X0,X3,X1) | coll(X2,X3,X0) | cyclic(X0,X1,X2,X3)) )),
  inference(cnf_transformation,[],[f223])).
fof(f223,plain,(
  ! [X0,X1,X2,X3] : (cyclic(X0,X1,X2,X3) | coll(X2,X3,X0) | ~eqangle(X2,X0,X2,X1,X3,X0,X3,X1))),
  inference(flattening,[],[f222])).
fof(f222,plain,(
  ! [X0,X1,X2,X3] : (cyclic(X0,X1,X2,X3) | (coll(X2,X3,X0) | ~eqangle(X2,X0,X2,X1,X3,X0,X3,X1)))),
  inference(ennf_transformation,[],[f121])).
fof(f121,plain,(
  ! [X0,X1,X2,X3] : ((~coll(X2,X3,X0) & eqangle(X2,X0,X2,X1,X3,X0,X3,X1)) => cyclic(X0,X1,X2,X3))),
  inference(rectify,[],[f42])).
fof(f42,axiom,(
  ! [X0,X1,X8,X9] : ((~coll(X8,X9,X0) & eqangle(X8,X0,X8,X1,X9,X0,X9,X1)) => cyclic(X0,X1,X8,X9))),
  file('./3_JAR/5cirdb.p',ruleD42a)).
fof(f1217,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (eqangle(X0,X1,X2,X3,X4,X5,X4,X5) | ~para(X0,X1,X2,X3)) )),
  inference(resolution,[],[f380,f399])).
fof(f399,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (eqangle(X0,X1,X4,X5,X2,X3,X4,X5) | ~para(X0,X1,X2,X3)) )),
  inference(cnf_transformation,[],[f220])).
fof(f220,plain,(
  ! [X0,X1,X2,X3,X4,X5] : (eqangle(X0,X1,X4,X5,X2,X3,X4,X5) | ~para(X0,X1,X2,X3))),
  inference(ennf_transformation,[],[f119])).
fof(f119,plain,(
  ! [X0,X1,X2,X3,X4,X5] : (para(X0,X1,X2,X3) => eqangle(X0,X1,X4,X5,X2,X3,X4,X5))),
  inference(rectify,[],[f40])).
fof(f40,axiom,(
  ! [X0,X1,X2,X3,X8,X9] : (para(X0,X1,X2,X3) => eqangle(X0,X1,X8,X9,X2,X3,X8,X9))),
  file('./3_JAR/5cirdb.p',ruleD40)).
fof(f380,plain,(
  ( ! [X6,X4,X2,X0,X7,X5,X3,X1] : (~eqangle(X0,X1,X2,X3,X4,X5,X6,X7) | eqangle(X0,X1,X4,X5,X2,X3,X6,X7)) )),
  inference(cnf_transformation,[],[f196])).
fof(f196,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6,X7] : (eqangle(X0,X1,X4,X5,X2,X3,X6,X7) | ~eqangle(X0,X1,X2,X3,X4,X5,X6,X7))),
  inference(ennf_transformation,[],[f103])).
fof(f103,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6,X7] : (eqangle(X0,X1,X2,X3,X4,X5,X6,X7) => eqangle(X0,X1,X4,X5,X2,X3,X6,X7))),
  inference(rectify,[],[f21])).
fof(f21,axiom,(
  ! [X0,X1,X2,X3,X8,X9,X10,X11] : (eqangle(X0,X1,X2,X3,X8,X9,X10,X11) => eqangle(X0,X1,X8,X9,X2,X3,X10,X11))),
  file('./3_JAR/5cirdb.p',ruleD21)).
fof(f4087,plain,(
  ( ! [X68,X66,X67] : (~para(X66,X67,X66,X67) | ~coll(X66,X68,X67) | cyclic(X67,X67,X66,X68)) )),
  inference(resolution,[],[f1217,f402])).
fof(f402,plain,(
  ( ! [X2,X0,X3,X1] : (~eqangle(X2,X0,X2,X1,X3,X0,X3,X1) | ~coll(X2,X3,X1) | cyclic(X0,X1,X2,X3)) )),
  inference(cnf_transformation,[],[f225])).
fof(f225,plain,(
  ! [X0,X1,X2,X3] : (cyclic(X0,X1,X2,X3) | ~coll(X2,X3,X1) | ~eqangle(X2,X0,X2,X1,X3,X0,X3,X1))),
  inference(flattening,[],[f224])).
fof(f224,plain,(
  ! [X0,X1,X2,X3] : (cyclic(X0,X1,X2,X3) | (~coll(X2,X3,X1) | ~eqangle(X2,X0,X2,X1,X3,X0,X3,X1)))),
  inference(ennf_transformation,[],[f122])).
fof(f122,plain,(
  ! [X0,X1,X2,X3] : ((coll(X2,X3,X1) & eqangle(X2,X0,X2,X1,X3,X0,X3,X1)) => cyclic(X0,X1,X2,X3))),
  inference(rectify,[],[f43])).
fof(f43,axiom,(
  ! [X0,X1,X8,X9] : ((coll(X8,X9,X1) & eqangle(X8,X0,X8,X1,X9,X0,X9,X1)) => cyclic(X0,X1,X8,X9))),
  file('./3_JAR/5cirdb.p',ruleD42b)).
fof(f38756,plain,(
  ( ! [X12,X10,X11,X9] : (cyclic(X9,X10,X11,X12) | ~cyclic(X9,X9,X10,X11)) ) | ~spl65_464),
  inference(resolution,[],[f20628,f376])).
fof(f376,plain,(
  ( ! [X4,X2,X0,X3,X1] : (~cyclic(X0,X1,X2,X4) | cyclic(X1,X2,X3,X4) | ~cyclic(X0,X1,X2,X3)) )),
  inference(cnf_transformation,[],[f192])).
fof(f192,plain,(
  ! [X0,X1,X2,X3,X4] : (cyclic(X1,X2,X3,X4) | ~cyclic(X0,X1,X2,X4) | ~cyclic(X0,X1,X2,X3))),
  inference(flattening,[],[f191])).
fof(f191,plain,(
  ! [X0,X1,X2,X3,X4] : (cyclic(X1,X2,X3,X4) | (~cyclic(X0,X1,X2,X4) | ~cyclic(X0,X1,X2,X3)))),
  inference(ennf_transformation,[],[f17])).
fof(f17,axiom,(
  ! [X0,X1,X2,X3,X4] : ((cyclic(X0,X1,X2,X4) & cyclic(X0,X1,X2,X3)) => cyclic(X1,X2,X3,X4))),
  file('./3_JAR/5cirdb.p',ruleD17)).
fof(f7979,plain,(
  ~spl65_466),
  inference(avatar_contradiction_clause,[],[f7950])).
fof(f7950,plain,(
  $false | ~spl65_466),
  inference(resolution,[],[f3931,f2658])).
fof(f2658,plain,(
  para(sK20,sK49,sK20,sK49)),
  inference(resolution,[],[f1139,f621])).
fof(f621,plain,(
  perp(sK20,sK49,sK43,sK44)),
  inference(resolution,[],[f367,f512])).
fof(f512,plain,(
  perp(sK43,sK44,sK20,sK49)),
  inference(cnf_transformation,[],[f359])).
fof(f367,plain,(
  ( ! [X2,X0,X3,X1] : (~perp(X0,X1,X2,X3) | perp(X2,X3,X0,X1)) )),
  inference(cnf_transformation,[],[f178])).
fof(f178,plain,(
  ! [X0,X1,X2,X3] : (perp(X2,X3,X0,X1) | ~perp(X0,X1,X2,X3))),
  inference(ennf_transformation,[],[f8])).
fof(f8,axiom,(
  ! [X0,X1,X2,X3] : (perp(X0,X1,X2,X3) => perp(X2,X3,X0,X1))),
  file('./3_JAR/5cirdb.p',ruleD8)).
fof(f1139,plain,(
  ( ! [X35,X34] : (~perp(X34,X35,sK43,sK44) | para(X34,X35,sK20,sK49)) )),
  inference(resolution,[],[f368,f512])).
fof(f368,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (~perp(X2,X3,X4,X5) | para(X0,X1,X4,X5) | ~perp(X0,X1,X2,X3)) )),
  inference(cnf_transformation,[],[f180])).
fof(f180,plain,(
  ! [X0,X1,X2,X3,X4,X5] : (para(X0,X1,X4,X5) | ~perp(X2,X3,X4,X5) | ~perp(X0,X1,X2,X3))),
  inference(flattening,[],[f179])).
fof(f179,plain,(
  ! [X0,X1,X2,X3,X4,X5] : (para(X0,X1,X4,X5) | (~perp(X2,X3,X4,X5) | ~perp(X0,X1,X2,X3)))),
  inference(ennf_transformation,[],[f9])).
fof(f9,axiom,(
  ! [X0,X1,X2,X3,X4,X5] : ((perp(X2,X3,X4,X5) & perp(X0,X1,X2,X3)) => para(X0,X1,X4,X5))),
  file('./3_JAR/5cirdb.p',ruleD9)).
fof(f3931,plain,(
  ( ! [X0,X1] : (~para(X0,X1,X0,X1)) ) | ~spl65_466),
  inference(avatar_component_clause,[],[f3930])).
fof(f3930,plain,(
  spl65_466 <=> ! [X1,X0] : ~para(X0,X1,X0,X1)),
  introduced(avatar_definition,[new_symbols(naming,[spl65_466])])).
fof(f3932,plain,(
  spl65_464 | spl65_466),
  inference(avatar_split_clause,[],[f3903,f3930,f3927])).
fof(f3903,plain,(
  ( ! [X2,X0,X3,X1] : (~para(X0,X1,X0,X1) | para(X2,X3,X2,X3)) )),
  inference(resolution,[],[f1203,f398])).
fof(f398,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (~eqangle(X0,X1,X4,X5,X2,X3,X4,X5) | para(X0,X1,X2,X3)) )),
  inference(cnf_transformation,[],[f219])).
fof(f219,plain,(
  ! [X0,X1,X2,X3,X4,X5] : (para(X0,X1,X2,X3) | ~eqangle(X0,X1,X4,X5,X2,X3,X4,X5))),
  inference(ennf_transformation,[],[f118])).
fof(f118,plain,(
  ! [X0,X1,X2,X3,X4,X5] : (eqangle(X0,X1,X4,X5,X2,X3,X4,X5) => para(X0,X1,X2,X3))),
  inference(rectify,[],[f39])).
fof(f39,axiom,(
  ! [X0,X1,X2,X3,X8,X9] : (eqangle(X0,X1,X8,X9,X2,X3,X8,X9) => para(X0,X1,X2,X3))),
  file('./3_JAR/5cirdb.p',ruleD39)).
fof(f1203,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (eqangle(X0,X1,X2,X3,X0,X1,X4,X5) | ~para(X2,X3,X4,X5)) )),
  inference(resolution,[],[f378,f399])).
fof(f378,plain,(
  ( ! [X6,X4,X2,X0,X7,X5,X3,X1] : (~eqangle(X0,X1,X2,X3,X4,X5,X6,X7) | eqangle(X2,X3,X0,X1,X6,X7,X4,X5)) )),
  inference(cnf_transformation,[],[f194])).
fof(f194,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6,X7] : (eqangle(X2,X3,X0,X1,X6,X7,X4,X5) | ~eqangle(X0,X1,X2,X3,X4,X5,X6,X7))),
  inference(ennf_transformation,[],[f101])).
fof(f101,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6,X7] : (eqangle(X0,X1,X2,X3,X4,X5,X6,X7) => eqangle(X2,X3,X0,X1,X6,X7,X4,X5))),
  inference(rectify,[],[f19])).
fof(f19,axiom,(
  ! [X0,X1,X2,X3,X8,X9,X10,X11] : (eqangle(X0,X1,X2,X3,X8,X9,X10,X11) => eqangle(X2,X3,X0,X1,X10,X11,X8,X9))),
  file('./3_JAR/5cirdb.p',ruleD19)).
% SZS output end Proof for 5cirdb
% ------------------------------
% Version: Vampire 4.4.0 (commit 7916d27 on 2019-08-23 08:50:16 +0100)
% Termination reason: Refutation

% Memory used [KB]: 286434
% Time elapsed: 40.564 s
% ------------------------------
% ------------------------------
