% Refutation found. Thanks to Tanya!
% SZS status Theorem for 5cirdb
% SZS output start Proof for 5cirdb
fof(f1085,plain,(
  $false),
  inference(resolution,[],[f1073,f521])).
fof(f521,plain,(
  ~cyclic(sK23,sK22,sK21,sK20)),
  inference(cnf_transformation,[],[f357])).
fof(f357,plain,(
  ~cyclic(sK23,sK22,sK21,sK20) & ! [X19,X20,X21,X22,X23,X24,X25,X26,X27,X28,X29,X30,X31,X32,X33,X34,X35,X36,X37,X38,X39,X40,X41,X42,X43,X44,X45,X46,X47,X48] : (circle(X29,X28,X38,X48) & circle(X33,X28,X38,X47) & circle(X33,X27,X37,X46) & circle(X32,X27,X37,X45) & circle(X32,X26,X36,X44) & circle(X31,X26,X36,X43) & circle(X31,X25,X35,X42) & circle(X30,X25,X35,X41) & circle(X30,X24,X34,X40) & circle(X29,X24,X34,X39) & perp(X27,X28,sK24,X33) & midp(sK24,X27,X28) & perp(X23,X28,sK25,X33) & midp(sK25,X23,X28) & perp(X23,X27,sK26,X33) & midp(sK26,X23,X27) & perp(X26,X27,sK27,X32) & midp(sK27,X26,X27) & perp(X22,X27,sK28,X32) & midp(sK28,X22,X27) & perp(X22,X26,sK29,X32) & midp(sK29,X22,X26) & perp(X25,X26,sK30,X31) & midp(sK30,X25,X26) & perp(X21,X26,sK31,X31) & midp(sK31,X21,X26) & perp(X21,X25,sK32,X31) & midp(sK32,X21,X25) & perp(X24,X25,sK33,X30) & midp(sK33,X24,X25) & perp(X20,X25,sK34,X30) & midp(sK34,X20,X25) & perp(X20,X24,sK35,X30) & midp(sK35,X20,X24) & perp(X24,X28,sK36,X29) & midp(sK36,X24,X28) & perp(X19,X28,sK37,X29) & midp(sK37,X19,X28) & perp(X19,X24,sK38,X29) & midp(sK38,X19,X24) & coll(X28,X20,X23) & coll(X28,X19,X22) & coll(X27,X21,X23) & coll(X27,X19,X22) & coll(X26,X21,X23) & coll(X26,X20,X22) & coll(X25,X20,X22) & coll(X25,X19,X21) & coll(X24,X20,X23) & coll(X24,X19,X21))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK20,sK21,sK22,sK23,sK24,sK25,sK26,sK27,sK28,sK29,sK30,sK31,sK32,sK33,sK34,sK35,sK36,sK37,sK38])],[f319,f356])).
fof(f356,plain,(
  ? [X0,X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18] : (~cyclic(X3,X2,X1,X0) & ! [X19,X20,X21,X22,X23,X24,X25,X26,X27,X28,X29,X30,X31,X32,X33,X34,X35,X36,X37,X38,X39,X40,X41,X42,X43,X44,X45,X46,X47,X48] : (circle(X29,X28,X38,X48) & circle(X33,X28,X38,X47) & circle(X33,X27,X37,X46) & circle(X32,X27,X37,X45) & circle(X32,X26,X36,X44) & circle(X31,X26,X36,X43) & circle(X31,X25,X35,X42) & circle(X30,X25,X35,X41) & circle(X30,X24,X34,X40) & circle(X29,X24,X34,X39) & perp(X27,X28,X4,X33) & midp(X4,X27,X28) & perp(X23,X28,X5,X33) & midp(X5,X23,X28) & perp(X23,X27,X6,X33) & midp(X6,X23,X27) & perp(X26,X27,X7,X32) & midp(X7,X26,X27) & perp(X22,X27,X8,X32) & midp(X8,X22,X27) & perp(X22,X26,X9,X32) & midp(X9,X22,X26) & perp(X25,X26,X10,X31) & midp(X10,X25,X26) & perp(X21,X26,X11,X31) & midp(X11,X21,X26) & perp(X21,X25,X12,X31) & midp(X12,X21,X25) & perp(X24,X25,X13,X30) & midp(X13,X24,X25) & perp(X20,X25,X14,X30) & midp(X14,X20,X25) & perp(X20,X24,X15,X30) & midp(X15,X20,X24) & perp(X24,X28,X16,X29) & midp(X16,X24,X28) & perp(X19,X28,X17,X29) & midp(X17,X19,X28) & perp(X19,X24,X18,X29) & midp(X18,X19,X24) & coll(X28,X20,X23) & coll(X28,X19,X22) & coll(X27,X21,X23) & coll(X27,X19,X22) & coll(X26,X21,X23) & coll(X26,X20,X22) & coll(X25,X20,X22) & coll(X25,X19,X21) & coll(X24,X20,X23) & coll(X24,X19,X21))) => (~cyclic(sK23,sK22,sK21,sK20) & ! [X48,X47,X46,X45,X44,X43,X42,X41,X40,X39,X38,X37,X36,X35,X34,X33,X32,X31,X30,X29,X28,X27,X26,X25,X24,X23,X22,X21,X20,X19] : (circle(X29,X28,X38,X48) & circle(X33,X28,X38,X47) & circle(X33,X27,X37,X46) & circle(X32,X27,X37,X45) & circle(X32,X26,X36,X44) & circle(X31,X26,X36,X43) & circle(X31,X25,X35,X42) & circle(X30,X25,X35,X41) & circle(X30,X24,X34,X40) & circle(X29,X24,X34,X39) & perp(X27,X28,sK24,X33) & midp(sK24,X27,X28) & perp(X23,X28,sK25,X33) & midp(sK25,X23,X28) & perp(X23,X27,sK26,X33) & midp(sK26,X23,X27) & perp(X26,X27,sK27,X32) & midp(sK27,X26,X27) & perp(X22,X27,sK28,X32) & midp(sK28,X22,X27) & perp(X22,X26,sK29,X32) & midp(sK29,X22,X26) & perp(X25,X26,sK30,X31) & midp(sK30,X25,X26) & perp(X21,X26,sK31,X31) & midp(sK31,X21,X26) & perp(X21,X25,sK32,X31) & midp(sK32,X21,X25) & perp(X24,X25,sK33,X30) & midp(sK33,X24,X25) & perp(X20,X25,sK34,X30) & midp(sK34,X20,X25) & perp(X20,X24,sK35,X30) & midp(sK35,X20,X24) & perp(X24,X28,sK36,X29) & midp(sK36,X24,X28) & perp(X19,X28,sK37,X29) & midp(sK37,X19,X28) & perp(X19,X24,sK38,X29) & midp(sK38,X19,X24) & coll(X28,X20,X23) & coll(X28,X19,X22) & coll(X27,X21,X23) & coll(X27,X19,X22) & coll(X26,X21,X23) & coll(X26,X20,X22) & coll(X25,X20,X22) & coll(X25,X19,X21) & coll(X24,X20,X23) & coll(X24,X19,X21)))),
  introduced(choice_axiom,[])).
fof(f319,plain,(
  ? [X0,X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18] : (~cyclic(X3,X2,X1,X0) & ! [X19,X20,X21,X22,X23,X24,X25,X26,X27,X28,X29,X30,X31,X32,X33,X34,X35,X36,X37,X38,X39,X40,X41,X42,X43,X44,X45,X46,X47,X48] : (circle(X29,X28,X38,X48) & circle(X33,X28,X38,X47) & circle(X33,X27,X37,X46) & circle(X32,X27,X37,X45) & circle(X32,X26,X36,X44) & circle(X31,X26,X36,X43) & circle(X31,X25,X35,X42) & circle(X30,X25,X35,X41) & circle(X30,X24,X34,X40) & circle(X29,X24,X34,X39) & perp(X27,X28,X4,X33) & midp(X4,X27,X28) & perp(X23,X28,X5,X33) & midp(X5,X23,X28) & perp(X23,X27,X6,X33) & midp(X6,X23,X27) & perp(X26,X27,X7,X32) & midp(X7,X26,X27) & perp(X22,X27,X8,X32) & midp(X8,X22,X27) & perp(X22,X26,X9,X32) & midp(X9,X22,X26) & perp(X25,X26,X10,X31) & midp(X10,X25,X26) & perp(X21,X26,X11,X31) & midp(X11,X21,X26) & perp(X21,X25,X12,X31) & midp(X12,X21,X25) & perp(X24,X25,X13,X30) & midp(X13,X24,X25) & perp(X20,X25,X14,X30) & midp(X14,X20,X25) & perp(X20,X24,X15,X30) & midp(X15,X20,X24) & perp(X24,X28,X16,X29) & midp(X16,X24,X28) & perp(X19,X28,X17,X29) & midp(X17,X19,X28) & perp(X19,X24,X18,X29) & midp(X18,X19,X24) & coll(X28,X20,X23) & coll(X28,X19,X22) & coll(X27,X21,X23) & coll(X27,X19,X22) & coll(X26,X21,X23) & coll(X26,X20,X22) & coll(X25,X20,X22) & coll(X25,X19,X21) & coll(X24,X20,X23) & coll(X24,X19,X21)))),
  inference(ennf_transformation,[],[f167])).
fof(f167,plain,(
  ~! [X0,X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18] : (! [X19,X20,X21,X22,X23,X24,X25,X26,X27,X28,X29,X30,X31,X32,X33,X34,X35,X36,X37,X38,X39,X40,X41,X42,X43,X44,X45,X46,X47,X48] : (circle(X29,X28,X38,X48) & circle(X33,X28,X38,X47) & circle(X33,X27,X37,X46) & circle(X32,X27,X37,X45) & circle(X32,X26,X36,X44) & circle(X31,X26,X36,X43) & circle(X31,X25,X35,X42) & circle(X30,X25,X35,X41) & circle(X30,X24,X34,X40) & circle(X29,X24,X34,X39) & perp(X27,X28,X4,X33) & midp(X4,X27,X28) & perp(X23,X28,X5,X33) & midp(X5,X23,X28) & perp(X23,X27,X6,X33) & midp(X6,X23,X27) & perp(X26,X27,X7,X32) & midp(X7,X26,X27) & perp(X22,X27,X8,X32) & midp(X8,X22,X27) & perp(X22,X26,X9,X32) & midp(X9,X22,X26) & perp(X25,X26,X10,X31) & midp(X10,X25,X26) & perp(X21,X26,X11,X31) & midp(X11,X21,X26) & perp(X21,X25,X12,X31) & midp(X12,X21,X25) & perp(X24,X25,X13,X30) & midp(X13,X24,X25) & perp(X20,X25,X14,X30) & midp(X14,X20,X25) & perp(X20,X24,X15,X30) & midp(X15,X20,X24) & perp(X24,X28,X16,X29) & midp(X16,X24,X28) & perp(X19,X28,X17,X29) & midp(X17,X19,X28) & perp(X19,X24,X18,X29) & midp(X18,X19,X24) & coll(X28,X20,X23) & coll(X28,X19,X22) & coll(X27,X21,X23) & coll(X27,X19,X22) & coll(X26,X21,X23) & coll(X26,X20,X22) & coll(X25,X20,X22) & coll(X25,X19,X21) & coll(X24,X20,X23) & coll(X24,X19,X21)) => cyclic(X3,X2,X1,X0))),
  inference(rectify,[],[f96])).
fof(f96,negated_conjecture,(
  ~! [X37,X35,X34,X33,X62,X61,X60,X59,X58,X57,X56,X55,X54,X53,X52,X51,X50,X49,X48] : (! [X18,X19,X20,X21,X22,X23,X24,X25,X26,X27,X28,X29,X30,X31,X32,X33,X34,X35,X36,X37,X38,X39,X40,X41,X42,X43,X44,X45,X46,X47] : (circle(X28,X27,X37,X47) & circle(X32,X27,X37,X46) & circle(X32,X26,X36,X45) & circle(X31,X26,X36,X44) & circle(X31,X25,X35,X43) & circle(X30,X25,X35,X42) & circle(X30,X24,X34,X41) & circle(X29,X24,X34,X40) & circle(X29,X23,X33,X39) & circle(X28,X23,X33,X38) & perp(X26,X27,X62,X32) & midp(X62,X26,X27) & perp(X22,X27,X61,X32) & midp(X61,X22,X27) & perp(X22,X26,X60,X32) & midp(X60,X22,X26) & perp(X25,X26,X59,X31) & midp(X59,X25,X26) & perp(X21,X26,X58,X31) & midp(X58,X21,X26) & perp(X21,X25,X57,X31) & midp(X57,X21,X25) & perp(X24,X25,X56,X30) & midp(X56,X24,X25) & perp(X20,X25,X55,X30) & midp(X55,X20,X25) & perp(X20,X24,X54,X30) & midp(X54,X20,X24) & perp(X23,X24,X53,X29) & midp(X53,X23,X24) & perp(X19,X24,X52,X29) & midp(X52,X19,X24) & perp(X19,X23,X51,X29) & midp(X51,X19,X23) & perp(X23,X27,X50,X28) & midp(X50,X23,X27) & perp(X18,X27,X49,X28) & midp(X49,X18,X27) & perp(X18,X23,X48,X28) & midp(X48,X18,X23) & coll(X27,X19,X22) & coll(X27,X18,X21) & coll(X26,X20,X22) & coll(X26,X18,X21) & coll(X25,X20,X22) & coll(X25,X19,X21) & coll(X24,X19,X21) & coll(X24,X18,X20) & coll(X23,X19,X22) & coll(X23,X18,X20)) => cyclic(X33,X34,X35,X37))),
  inference(negated_conjecture,[],[f95])).
fof(f95,conjecture,(
  ( ! [X61,X59,X57,X37,X35,X33,X54,X52,X50,X48,X62,X60,X58,X56,X34,X55,X53,X51,X49] : (! [X18,X19,X20,X21,X22,X23,X24,X25,X26,X27,X28,X29,X30,X31,X32,X33,X34,X35,X36,X37,X38,X39,X40,X41,X42,X43,X44,X45,X46,X47] : (circle(X28,X27,X37,X47) & circle(X32,X27,X37,X46) & circle(X32,X26,X36,X45) & circle(X31,X26,X36,X44) & circle(X31,X25,X35,X43) & circle(X30,X25,X35,X42) & circle(X30,X24,X34,X41) & circle(X29,X24,X34,X40) & circle(X29,X23,X33,X39) & circle(X28,X23,X33,X38) & perp(X26,X27,X62,X32) & midp(X62,X26,X27) & perp(X22,X27,X61,X32) & midp(X61,X22,X27) & perp(X22,X26,X60,X32) & midp(X60,X22,X26) & perp(X25,X26,X59,X31) & midp(X59,X25,X26) & perp(X21,X26,X58,X31) & midp(X58,X21,X26) & perp(X21,X25,X57,X31) & midp(X57,X21,X25) & perp(X24,X25,X56,X30) & midp(X56,X24,X25) & perp(X20,X25,X55,X30) & midp(X55,X20,X25) & perp(X20,X24,X54,X30) & midp(X54,X20,X24) & perp(X23,X24,X53,X29) & midp(X53,X23,X24) & perp(X19,X24,X52,X29) & midp(X52,X19,X24) & perp(X19,X23,X51,X29) & midp(X51,X19,X23) & perp(X23,X27,X50,X28) & midp(X50,X23,X27) & perp(X18,X27,X49,X28) & midp(X49,X18,X27) & perp(X18,X23,X48,X28) & midp(X48,X18,X23) & coll(X27,X19,X22) & coll(X27,X18,X21) & coll(X26,X20,X22) & coll(X26,X18,X21) & coll(X25,X20,X22) & coll(X25,X19,X21) & coll(X24,X19,X21) & coll(X24,X18,X20) & coll(X23,X19,X22) & coll(X23,X18,X20)) => cyclic(X33,X34,X35,X37)) )),
  file('./3_JAR/5cirdb.p',exemplo3JAR5cirdb)).
fof(f1073,plain,(
  ( ! [X2,X0,X3,X1] : (cyclic(X0,X1,X2,X3)) )),
  inference(subsumption_resolution,[],[f1069,f1059])).
fof(f1059,plain,(
  ( ! [X12,X10,X11] : (cyclic(X10,X11,X10,X12)) )),
  inference(resolution,[],[f1051,f371])).
fof(f371,plain,(
  ( ! [X2,X0,X3,X1] : (~cyclic(X0,X1,X2,X3) | cyclic(X0,X1,X3,X2)) )),
  inference(cnf_transformation,[],[f187])).
fof(f187,plain,(
  ! [X0,X1,X2,X3] : (cyclic(X0,X1,X3,X2) | ~cyclic(X0,X1,X2,X3))),
  inference(ennf_transformation,[],[f14])).
fof(f14,axiom,(
  ! [X0,X1,X2,X3] : (cyclic(X0,X1,X2,X3) => cyclic(X0,X1,X3,X2))),
  file('./3_JAR/5cirdb.p',ruleD14)).
fof(f1051,plain,(
  ( ! [X6,X4,X5] : (cyclic(X4,X5,X6,X4)) )),
  inference(resolution,[],[f1024,f373])).
fof(f373,plain,(
  ( ! [X2,X0,X3,X1] : (~cyclic(X0,X1,X2,X3) | cyclic(X1,X0,X2,X3)) )),
  inference(cnf_transformation,[],[f189])).
fof(f189,plain,(
  ! [X0,X1,X2,X3] : (cyclic(X1,X0,X2,X3) | ~cyclic(X0,X1,X2,X3))),
  inference(ennf_transformation,[],[f16])).
fof(f16,axiom,(
  ! [X0,X1,X2,X3] : (cyclic(X0,X1,X2,X3) => cyclic(X1,X0,X2,X3))),
  file('./3_JAR/5cirdb.p',ruleD16)).
fof(f1024,plain,(
  ( ! [X8,X7,X9] : (cyclic(X7,X8,X9,X8)) )),
  inference(resolution,[],[f1021,f372])).
fof(f372,plain,(
  ( ! [X2,X0,X3,X1] : (~cyclic(X0,X1,X2,X3) | cyclic(X0,X2,X1,X3)) )),
  inference(cnf_transformation,[],[f188])).
fof(f188,plain,(
  ! [X0,X1,X2,X3] : (cyclic(X0,X2,X1,X3) | ~cyclic(X0,X1,X2,X3))),
  inference(ennf_transformation,[],[f15])).
fof(f15,axiom,(
  ! [X0,X1,X2,X3] : (cyclic(X0,X1,X2,X3) => cyclic(X0,X2,X1,X3))),
  file('./3_JAR/5cirdb.p',ruleD15)).
fof(f1021,plain,(
  ( ! [X2,X0,X1] : (cyclic(X2,X1,X0,X0)) )),
  inference(subsumption_resolution,[],[f1020,f637])).
fof(f637,plain,(
  ( ! [X2,X0,X3,X1] : (para(X0,X1,X2,X3)) )),
  inference(subsumption_resolution,[],[f622,f509])).
fof(f509,plain,(
  ( ! [X28,X27] : (midp(sK24,X27,X28)) )),
  inference(cnf_transformation,[],[f357])).
fof(f622,plain,(
  ( ! [X2,X0,X3,X1] : (para(X0,X1,X2,X3) | ~midp(sK24,X0,X2)) )),
  inference(resolution,[],[f421,f509])).
fof(f421,plain,(
  ( ! [X4,X2,X0,X3,X1] : (~midp(X4,X2,X3) | para(X0,X2,X1,X3) | ~midp(X4,X0,X1)) )),
  inference(cnf_transformation,[],[f262])).
fof(f262,plain,(
  ! [X0,X1,X2,X3,X4] : (para(X0,X2,X1,X3) | ~midp(X4,X2,X3) | ~midp(X4,X0,X1))),
  inference(flattening,[],[f261])).
fof(f261,plain,(
  ! [X0,X1,X2,X3,X4] : (para(X0,X2,X1,X3) | (~midp(X4,X2,X3) | ~midp(X4,X0,X1)))),
  inference(ennf_transformation,[],[f142])).
fof(f142,plain,(
  ! [X0,X1,X2,X3,X4] : ((midp(X4,X2,X3) & midp(X4,X0,X1)) => para(X0,X2,X1,X3))),
  inference(rectify,[],[f64])).
fof(f64,axiom,(
  ! [X0,X1,X2,X3,X6] : ((midp(X6,X2,X3) & midp(X6,X0,X1)) => para(X0,X2,X1,X3))),
  file('./3_JAR/5cirdb.p',ruleD63)).
fof(f1020,plain,(
  ( ! [X2,X0,X1] : (cyclic(X2,X1,X0,X0) | ~para(X0,X2,X0,X2)) )),
  inference(subsumption_resolution,[],[f1016,f480])).
fof(f480,plain,(
  ( ! [X28,X23,X20] : (coll(X28,X20,X23)) )),
  inference(cnf_transformation,[],[f357])).
fof(f1016,plain,(
  ( ! [X2,X0,X1] : (~coll(X0,X0,X1) | cyclic(X2,X1,X0,X0) | ~para(X0,X2,X0,X2)) )),
  inference(resolution,[],[f400,f397])).
fof(f397,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (eqangle(X0,X1,X4,X5,X2,X3,X4,X5) | ~para(X0,X1,X2,X3)) )),
  inference(cnf_transformation,[],[f219])).
fof(f219,plain,(
  ! [X0,X1,X2,X3,X4,X5] : (eqangle(X0,X1,X4,X5,X2,X3,X4,X5) | ~para(X0,X1,X2,X3))),
  inference(ennf_transformation,[],[f119])).
fof(f119,plain,(
  ! [X0,X1,X2,X3,X4,X5] : (para(X0,X1,X2,X3) => eqangle(X0,X1,X4,X5,X2,X3,X4,X5))),
  inference(rectify,[],[f40])).
fof(f40,axiom,(
  ! [X0,X1,X2,X3,X8,X9] : (para(X0,X1,X2,X3) => eqangle(X0,X1,X8,X9,X2,X3,X8,X9))),
  file('./3_JAR/5cirdb.p',ruleD40)).
fof(f400,plain,(
  ( ! [X2,X0,X3,X1] : (~eqangle(X2,X0,X2,X1,X3,X0,X3,X1) | ~coll(X2,X3,X1) | cyclic(X0,X1,X2,X3)) )),
  inference(cnf_transformation,[],[f224])).
fof(f224,plain,(
  ! [X0,X1,X2,X3] : (cyclic(X0,X1,X2,X3) | ~coll(X2,X3,X1) | ~eqangle(X2,X0,X2,X1,X3,X0,X3,X1))),
  inference(flattening,[],[f223])).
fof(f223,plain,(
  ! [X0,X1,X2,X3] : (cyclic(X0,X1,X2,X3) | (~coll(X2,X3,X1) | ~eqangle(X2,X0,X2,X1,X3,X0,X3,X1)))),
  inference(ennf_transformation,[],[f122])).
fof(f122,plain,(
  ! [X0,X1,X2,X3] : ((coll(X2,X3,X1) & eqangle(X2,X0,X2,X1,X3,X0,X3,X1)) => cyclic(X0,X1,X2,X3))),
  inference(rectify,[],[f43])).
fof(f43,axiom,(
  ! [X0,X1,X8,X9] : ((coll(X8,X9,X1) & eqangle(X8,X0,X8,X1,X9,X0,X9,X1)) => cyclic(X0,X1,X8,X9))),
  file('./3_JAR/5cirdb.p',ruleD42b)).
fof(f1069,plain,(
  ( ! [X2,X0,X3,X1] : (cyclic(X0,X1,X2,X3) | ~cyclic(X1,X0,X1,X2)) )),
  inference(resolution,[],[f1059,f374])).
fof(f374,plain,(
  ( ! [X4,X2,X0,X3,X1] : (~cyclic(X0,X1,X2,X4) | cyclic(X1,X2,X3,X4) | ~cyclic(X0,X1,X2,X3)) )),
  inference(cnf_transformation,[],[f191])).
fof(f191,plain,(
  ! [X0,X1,X2,X3,X4] : (cyclic(X1,X2,X3,X4) | ~cyclic(X0,X1,X2,X4) | ~cyclic(X0,X1,X2,X3))),
  inference(flattening,[],[f190])).
fof(f190,plain,(
  ! [X0,X1,X2,X3,X4] : (cyclic(X1,X2,X3,X4) | (~cyclic(X0,X1,X2,X4) | ~cyclic(X0,X1,X2,X3)))),
  inference(ennf_transformation,[],[f17])).
fof(f17,axiom,(
  ! [X0,X1,X2,X3,X4] : ((cyclic(X0,X1,X2,X4) & cyclic(X0,X1,X2,X3)) => cyclic(X1,X2,X3,X4))),
  file('./3_JAR/5cirdb.p',ruleD17)).
% SZS output end Proof for 5cirdb
% ------------------------------
% Version: Vampire 4.4.0 (commit 7916d27 on 2019-08-23 08:50:16 +0100)
% Termination reason: Refutation

% Memory used [KB]: 5373
% Time elapsed: 0.033 s
% ------------------------------
% ------------------------------
