% Refutation found. Thanks to Tanya!
% SZS status Theorem for 2angle
% SZS output start Proof for 2angle
18. ! [X0,X1,X2,X3,X8,X9,X10,X11] : (eqangle(X0,X1,X2,X3,X8,X9,X10,X11) => eqangle(X1,X0,X2,X3,X8,X9,X10,X11)) [input]
19. ! [X0,X1,X2,X3,X8,X9,X10,X11] : (eqangle(X0,X1,X2,X3,X8,X9,X10,X11) => eqangle(X2,X3,X0,X1,X10,X11,X8,X9)) [input]
21. ! [X0,X1,X2,X3,X8,X9,X10,X11] : (eqangle(X0,X1,X2,X3,X8,X9,X10,X11) => eqangle(X0,X1,X8,X9,X2,X3,X10,X11)) [input]
22. ! [X0,X1,X2,X3,X8,X9,X10,X11,X4,X5,X12,X13] : ((eqangle(X8,X9,X10,X11,X4,X5,X12,X13) & eqangle(X0,X1,X2,X3,X8,X9,X10,X11)) => eqangle(X0,X1,X2,X3,X4,X5,X12,X13)) [input]
40. ! [X0,X1,X2,X3,X8,X9] : (para(X0,X1,X2,X3) => eqangle(X0,X1,X8,X9,X2,X3,X8,X9)) [input]
44. ! [X0,X1,X2,X4,X5] : ((midp(X5,X0,X2) & midp(X4,X0,X1)) => para(X4,X5,X1,X2)) [input]
94. ! [X0,X1,X2,X3,X4,X5,X12,X13] : (coll(X13,X1,X2) & coll(X13,X1,X2) & coll(X13,X4,X5) & coll(X12,X4,X5) & coll(X12,X0,X3) & midp(X5,X1,X0) & midp(X4,X2,X3) & cong(X0,X3,X2,X1)) => eqangle(X0,X12,X12,X5,X5,X13,X13,X1) [input]
95. ~! [X1,X13,X5,X12,X0] : (! [X0,X1,X2,X3,X4,X5,X12,X13] : (coll(X13,X1,X2) & coll(X13,X1,X2) & coll(X13,X4,X5) & coll(X12,X4,X5) & coll(X12,X0,X3) & midp(X5,X1,X0) & midp(X4,X2,X3) & cong(X0,X3,X2,X1)) => eqangle(X0,X12,X12,X5,X5,X13,X13,X1)) [negated conjecture 94]
99. ! [X0,X1,X2,X3,X4,X5,X6,X7] : (eqangle(X0,X1,X2,X3,X4,X5,X6,X7) => eqangle(X1,X0,X2,X3,X4,X5,X6,X7)) [rectify 18]
100. ! [X0,X1,X2,X3,X4,X5,X6,X7] : (eqangle(X0,X1,X2,X3,X4,X5,X6,X7) => eqangle(X2,X3,X0,X1,X6,X7,X4,X5)) [rectify 19]
102. ! [X0,X1,X2,X3,X4,X5,X6,X7] : (eqangle(X0,X1,X2,X3,X4,X5,X6,X7) => eqangle(X0,X1,X4,X5,X2,X3,X6,X7)) [rectify 21]
103. ! [X0,X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11] : ((eqangle(X4,X5,X6,X7,X8,X9,X10,X11) & eqangle(X0,X1,X2,X3,X4,X5,X6,X7)) => eqangle(X0,X1,X2,X3,X8,X9,X10,X11)) [rectify 22]
118. ! [X0,X1,X2,X3,X4,X5] : (para(X0,X1,X2,X3) => eqangle(X0,X1,X4,X5,X2,X3,X4,X5)) [rectify 40]
122. ! [X0,X1,X2,X3,X4] : ((midp(X4,X0,X2) & midp(X3,X0,X1)) => para(X3,X4,X1,X2)) [rectify 44]
165. ~! [X0,X1,X2,X3,X4] : (! [X5,X6,X7,X8,X9,X10,X11,X12] : (coll(X12,X6,X7) & coll(X12,X6,X7) & coll(X12,X9,X10) & coll(X11,X9,X10) & coll(X11,X5,X8) & midp(X10,X6,X5) & midp(X9,X7,X8) & cong(X5,X8,X7,X6)) => eqangle(X4,X3,X3,X2,X2,X1,X1,X0)) [rectify 95]
190. ! [X0,X1,X2,X3,X4,X5,X6,X7] : (eqangle(X1,X0,X2,X3,X4,X5,X6,X7) | ~eqangle(X0,X1,X2,X3,X4,X5,X6,X7)) [ennf transformation 99]
191. ! [X0,X1,X2,X3,X4,X5,X6,X7] : (eqangle(X2,X3,X0,X1,X6,X7,X4,X5) | ~eqangle(X0,X1,X2,X3,X4,X5,X6,X7)) [ennf transformation 100]
193. ! [X0,X1,X2,X3,X4,X5,X6,X7] : (eqangle(X0,X1,X4,X5,X2,X3,X6,X7) | ~eqangle(X0,X1,X2,X3,X4,X5,X6,X7)) [ennf transformation 102]
194. ! [X0,X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11] : (eqangle(X0,X1,X2,X3,X8,X9,X10,X11) | (~eqangle(X4,X5,X6,X7,X8,X9,X10,X11) | ~eqangle(X0,X1,X2,X3,X4,X5,X6,X7))) [ennf transformation 103]
195. ! [X0,X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11] : (eqangle(X0,X1,X2,X3,X8,X9,X10,X11) | ~eqangle(X4,X5,X6,X7,X8,X9,X10,X11) | ~eqangle(X0,X1,X2,X3,X4,X5,X6,X7)) [flattening 194]
217. ! [X0,X1,X2,X3,X4,X5] : (eqangle(X0,X1,X4,X5,X2,X3,X4,X5) | ~para(X0,X1,X2,X3)) [ennf transformation 118]
223. ! [X0,X1,X2,X3,X4] : (para(X3,X4,X1,X2) | (~midp(X4,X0,X2) | ~midp(X3,X0,X1))) [ennf transformation 122]
224. ! [X0,X1,X2,X3,X4] : (para(X3,X4,X1,X2) | ~midp(X4,X0,X2) | ~midp(X3,X0,X1)) [flattening 223]
315. ? [X0,X1,X2,X3,X4] : (~eqangle(X4,X3,X3,X2,X2,X1,X1,X0) & ! [X5,X6,X7,X8,X9,X10,X11,X12] : (coll(X12,X6,X7) & coll(X12,X6,X7) & coll(X12,X9,X10) & coll(X11,X9,X10) & coll(X11,X5,X8) & midp(X10,X6,X5) & midp(X9,X7,X8) & cong(X5,X8,X7,X6))) [ennf transformation 165]
352. ? [X0,X1,X2,X3,X4] : (~eqangle(X4,X3,X3,X2,X2,X1,X1,X0) & ! [X5,X6,X7,X8,X9,X10,X11,X12] : (coll(X12,X6,X7) & coll(X12,X6,X7) & coll(X12,X9,X10) & coll(X11,X9,X10) & coll(X11,X5,X8) & midp(X10,X6,X5) & midp(X9,X7,X8) & cong(X5,X8,X7,X6))) => (~eqangle(sK24,sK23,sK23,sK22,sK22,sK21,sK21,sK20) & ! [X5,X6,X7,X8,X9,X10,X11,X12] : (coll(X12,X6,X7) & coll(X12,X6,X7) & coll(X12,X9,X10) & coll(X11,X9,X10) & coll(X11,X5,X8) & midp(X10,X6,X5) & midp(X9,X7,X8) & cong(X5,X8,X7,X6))) [choice axiom]
353. ~eqangle(sK24,sK23,sK23,sK22,sK22,sK21,sK21,sK20) & ! [X5,X6,X7,X8,X9,X10,X11,X12] : (coll(X12,X6,X7) & coll(X12,X6,X7) & coll(X12,X9,X10) & coll(X11,X9,X10) & coll(X11,X5,X8) & midp(X10,X6,X5) & midp(X9,X7,X8) & cong(X5,X8,X7,X6)) [skolemisation 315,352]
371. ~eqangle(X0,X1,X2,X3,X4,X5,X6,X7) | eqangle(X1,X0,X2,X3,X4,X5,X6,X7) [cnf transformation 190]
372. ~eqangle(X0,X1,X2,X3,X4,X5,X6,X7) | eqangle(X2,X3,X0,X1,X6,X7,X4,X5) [cnf transformation 191]
374. ~eqangle(X0,X1,X2,X3,X4,X5,X6,X7) | eqangle(X0,X1,X4,X5,X2,X3,X6,X7) [cnf transformation 193]
375. ~eqangle(X4,X5,X6,X7,X8,X9,X10,X11) | eqangle(X0,X1,X2,X3,X8,X9,X10,X11) | ~eqangle(X0,X1,X2,X3,X4,X5,X6,X7) [cnf transformation 195]
393. eqangle(X0,X1,X4,X5,X2,X3,X4,X5) | ~para(X0,X1,X2,X3) [cnf transformation 217]
398. ~midp(X4,X0,X2) | para(X3,X4,X1,X2) | ~midp(X3,X0,X1) [cnf transformation 224]
469. midp(X10,X6,X5) [cnf transformation 353]
475. ~eqangle(sK24,sK23,sK23,sK22,sK22,sK21,sK21,sK20) [cnf transformation 353]
485. para(X0,X1,X2,X3) | ~midp(X0,X4,X2) [resolution 398,469]
486. para(X0,X1,X2,X3) [subsumption resolution 485,469]
556. eqangle(X0,X1,X2,X3,X4,X5,X2,X3) | ~para(X1,X0,X4,X5) [resolution 371,393]
560. eqangle(X0,X1,X2,X3,X4,X5,X2,X3) [subsumption resolution 556,486]
564. eqangle(X0,X1,X2,X3,X0,X1,X4,X5) [resolution 372,560]
584. eqangle(X0,X1,X2,X3,X4,X5,X4,X5) [resolution 374,560]
613. eqangle(X2,X3,X2,X3,X4,X5,X6,X7) [resolution 564,374]
1591. eqangle(X258,X259,X260,X261,X262,X263,X264,X265) | ~eqangle(X258,X259,X260,X261,X266,X267,X266,X267) [resolution 375,613]
1603. eqangle(X258,X259,X260,X261,X262,X263,X264,X265) [subsumption resolution 1591,584]
2730. $false [resolution 1603,475]
% SZS output end Proof for 2angle
% ------------------------------
% Version: Vampire 4.4.0 (commit 7916d27 on 2019-08-23 08:50:16 +0100)
% Termination reason: Refutation

% Memory used [KB]: 5628
% Time elapsed: 0.043 s
% ------------------------------
% ------------------------------
