% Refutation found. Thanks to Tanya!
% SZS status Theorem for think
% SZS output start Proof for think
23. ! [X0,X1,X2,X3] : (cong(X0,X1,X2,X3) => cong(X0,X1,X3,X2)) [input]
24. ! [X0,X1,X2,X3] : (cong(X0,X1,X2,X3) => cong(X2,X3,X0,X1)) [input]
25. ! [X0,X1,X2,X3,X4,X5] : ((cong(X2,X3,X4,X5) & cong(X0,X1,X2,X3)) => cong(X0,X1,X4,X5)) [input]
67. ! [X0,X1,X2] : ((coll(X0,X1,X2) & cong(X0,X1,X0,X2)) => midp(X0,X1,X2)) [input]
68. ! [X0,X1,X2] : (midp(X0,X1,X2) => cong(X0,X1,X0,X2)) [input]
94. ! [X0,X1,X2,X3,X4,X5] : (cong(X3,X5,X4,X5) & coll(X5,X3,X4) & coll(X5,X0,X1) & coll(X4,X1,X2) & coll(X3,X0,X2) & (cong(X2,X0,X2,X1) | cong(X2,X0,X0,X1))) => cong(X3,X0,X1,X4) [input]
95. ~! [X4,X1,X0,X3] : (! [X0,X1,X2,X3,X4,X5] : (cong(X3,X5,X4,X5) & coll(X5,X3,X4) & coll(X5,X0,X1) & coll(X4,X1,X2) & coll(X3,X0,X2) & (cong(X2,X0,X2,X1) | cong(X2,X0,X0,X1))) => cong(X3,X0,X1,X4)) [negated conjecture 94]
165. ~! [X0,X1,X2,X3] : (! [X4,X5,X6,X7,X8,X9] : (cong(X7,X9,X8,X9) & coll(X9,X7,X8) & coll(X9,X4,X5) & coll(X8,X5,X6) & coll(X7,X4,X6) & (cong(X6,X4,X6,X5) | cong(X6,X4,X4,X5))) => cong(X3,X2,X1,X0)) [rectify 95]
196. ! [X0,X1,X2,X3] : (cong(X0,X1,X3,X2) | ~cong(X0,X1,X2,X3)) [ennf transformation 23]
197. ! [X0,X1,X2,X3] : (cong(X2,X3,X0,X1) | ~cong(X0,X1,X2,X3)) [ennf transformation 24]
198. ! [X0,X1,X2,X3,X4,X5] : (cong(X0,X1,X4,X5) | (~cong(X2,X3,X4,X5) | ~cong(X0,X1,X2,X3))) [ennf transformation 25]
199. ! [X0,X1,X2,X3,X4,X5] : (cong(X0,X1,X4,X5) | ~cong(X2,X3,X4,X5) | ~cong(X0,X1,X2,X3)) [flattening 198]
264. ! [X0,X1,X2] : (midp(X0,X1,X2) | (~coll(X0,X1,X2) | ~cong(X0,X1,X0,X2))) [ennf transformation 67]
265. ! [X0,X1,X2] : (midp(X0,X1,X2) | ~coll(X0,X1,X2) | ~cong(X0,X1,X0,X2)) [flattening 264]
266. ! [X0,X1,X2] : (cong(X0,X1,X0,X2) | ~midp(X0,X1,X2)) [ennf transformation 68]
315. ? [X0,X1,X2,X3] : (~cong(X3,X2,X1,X0) & ! [X4,X5,X6,X7,X8,X9] : (cong(X7,X9,X8,X9) & coll(X9,X7,X8) & coll(X9,X4,X5) & coll(X8,X5,X6) & coll(X7,X4,X6) & (cong(X6,X4,X6,X5) | cong(X6,X4,X4,X5)))) [ennf transformation 165]
352. ? [X0,X1,X2,X3] : (~cong(X3,X2,X1,X0) & ! [X4,X5,X6,X7,X8,X9] : (cong(X7,X9,X8,X9) & coll(X9,X7,X8) & coll(X9,X4,X5) & coll(X8,X5,X6) & coll(X7,X4,X6) & (cong(X6,X4,X6,X5) | cong(X6,X4,X4,X5)))) => (~cong(sK23,sK22,sK21,sK20) & ! [X4,X5,X6,X7,X8,X9] : (cong(X7,X9,X8,X9) & coll(X9,X7,X8) & coll(X9,X4,X5) & coll(X8,X5,X6) & coll(X7,X4,X6) & (cong(X6,X4,X6,X5) | cong(X6,X4,X4,X5)))) [choice axiom]
353. ~cong(sK23,sK22,sK21,sK20) & ! [X4,X5,X6,X7,X8,X9] : (cong(X7,X9,X8,X9) & coll(X9,X7,X8) & coll(X9,X4,X5) & coll(X8,X5,X6) & coll(X7,X4,X6) & (cong(X6,X4,X6,X5) | cong(X6,X4,X4,X5))) [skolemisation 315,352]
376. ~cong(X0,X1,X2,X3) | cong(X0,X1,X3,X2) [cnf transformation 196]
377. ~cong(X0,X1,X2,X3) | cong(X2,X3,X0,X1) [cnf transformation 197]
378. ~cong(X2,X3,X4,X5) | cong(X0,X1,X4,X5) | ~cong(X0,X1,X2,X3) [cnf transformation 199]
421. ~cong(X0,X1,X0,X2) | ~coll(X0,X1,X2) | midp(X0,X1,X2) [cnf transformation 265]
422. cong(X0,X1,X0,X2) | ~midp(X0,X1,X2) [cnf transformation 266]
471. coll(X9,X7,X8) [cnf transformation 353]
472. cong(X7,X9,X8,X9) [cnf transformation 353]
473. ~cong(sK23,sK22,sK21,sK20) [cnf transformation 353]
477. cong(X0,X1,X1,X2) [resolution 376,472]
484. cong(X6,X7,X6,X8) | ~midp(X6,X8,X7) [resolution 377,422]
487. cong(X0,X1,X2,X0) [resolution 477,377]
493. cong(X3,X4,X3,X5) [resolution 487,376]
501. ~coll(X6,X7,X8) | midp(X6,X7,X8) [resolution 421,493]
510. midp(X6,X7,X8) [subsumption resolution 501,471]
588. ~cong(X0,X1,X4,X3) | cong(X0,X1,X2,X3) [resolution 378,472]
629. cong(X18,X19,X20,X21) | ~midp(X18,X21,X19) [resolution 588,484]
637. cong(X18,X19,X20,X21) [subsumption resolution 629,510]
639. $false [resolution 637,473]
% SZS output end Proof for think
% ------------------------------
% Version: Vampire 4.4.0 (commit 7916d27 on 2019-08-23 08:50:16 +0100)
% Termination reason: Refutation

% Memory used [KB]: 5117
% Time elapsed: 0.029 s
% ------------------------------
% ------------------------------
