% Refutation found. Thanks to Tanya!
% SZS status Theorem for E022-13
% SZS output start Proof for E022-13
fof(f651,plain,(
  $false),
  inference(resolution,[],[f641,f485])).
fof(f485,plain,(
  ~para(sK23,sK22,sK21,sK20)),
  inference(cnf_transformation,[],[f357])).
fof(f357,plain,(
  ~para(sK23,sK22,sK21,sK20) & ! [X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25] : (circle(X6,X5,X13,X25) & coll(X5,X10,X13) & circle(X6,X7,X12,X24) & coll(X7,X11,X12) & circle(X4,X5,X11,X23) & coll(X5,X9,X11) & circle(X4,X7,X10,X22) & coll(X7,X8,X10) & circle(X6,X5,X9,X21) & circle(X6,X5,X8,X20) & circle(X6,X5,X7,X19) & circle(X4,X5,X7,X18) & circle(X6,X5,X16,X17) & circle(X4,X5,X14,X15))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK20,sK21,sK22,sK23])],[f319,f356])).
fof(f356,plain,(
  ? [X0,X1,X2,X3] : (~para(X3,X2,X1,X0) & ! [X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25] : (circle(X6,X5,X13,X25) & coll(X5,X10,X13) & circle(X6,X7,X12,X24) & coll(X7,X11,X12) & circle(X4,X5,X11,X23) & coll(X5,X9,X11) & circle(X4,X7,X10,X22) & coll(X7,X8,X10) & circle(X6,X5,X9,X21) & circle(X6,X5,X8,X20) & circle(X6,X5,X7,X19) & circle(X4,X5,X7,X18) & circle(X6,X5,X16,X17) & circle(X4,X5,X14,X15))) => (~para(sK23,sK22,sK21,sK20) & ! [X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25] : (circle(X6,X5,X13,X25) & coll(X5,X10,X13) & circle(X6,X7,X12,X24) & coll(X7,X11,X12) & circle(X4,X5,X11,X23) & coll(X5,X9,X11) & circle(X4,X7,X10,X22) & coll(X7,X8,X10) & circle(X6,X5,X9,X21) & circle(X6,X5,X8,X20) & circle(X6,X5,X7,X19) & circle(X4,X5,X7,X18) & circle(X6,X5,X16,X17) & circle(X4,X5,X14,X15)))),
  introduced(choice_axiom,[])).
fof(f319,plain,(
  ? [X0,X1,X2,X3] : (~para(X3,X2,X1,X0) & ! [X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25] : (circle(X6,X5,X13,X25) & coll(X5,X10,X13) & circle(X6,X7,X12,X24) & coll(X7,X11,X12) & circle(X4,X5,X11,X23) & coll(X5,X9,X11) & circle(X4,X7,X10,X22) & coll(X7,X8,X10) & circle(X6,X5,X9,X21) & circle(X6,X5,X8,X20) & circle(X6,X5,X7,X19) & circle(X4,X5,X7,X18) & circle(X6,X5,X16,X17) & circle(X4,X5,X14,X15)))),
  inference(ennf_transformation,[],[f167])).
fof(f167,plain,(
  ~! [X0,X1,X2,X3] : (! [X4,X5,X6,X7,X8,X9,X10,X11,X12,X13,X14,X15,X16,X17,X18,X19,X20,X21,X22,X23,X24,X25] : (circle(X6,X5,X13,X25) & coll(X5,X10,X13) & circle(X6,X7,X12,X24) & coll(X7,X11,X12) & circle(X4,X5,X11,X23) & coll(X5,X9,X11) & circle(X4,X7,X10,X22) & coll(X7,X8,X10) & circle(X6,X5,X9,X21) & circle(X6,X5,X8,X20) & circle(X6,X5,X7,X19) & circle(X4,X5,X7,X18) & circle(X6,X5,X16,X17) & circle(X4,X5,X14,X15)) => para(X3,X2,X1,X0))),
  inference(rectify,[],[f96])).
fof(f96,negated_conjecture,(
  ~! [X19,X18,X5,X4] : (! [X0,X1,X2,X3,X4,X5,X12,X13,X18,X19,X20,X21,X22,X23,X24,X25,X26,X27,X28,X29,X30,X31] : (circle(X2,X1,X19,X31) & coll(X1,X12,X19) & circle(X2,X3,X18,X30) & coll(X3,X13,X18) & circle(X0,X1,X13,X29) & coll(X1,X5,X13) & circle(X0,X3,X12,X28) & coll(X3,X4,X12) & circle(X2,X1,X5,X27) & circle(X2,X1,X4,X26) & circle(X2,X1,X3,X25) & circle(X0,X1,X3,X24) & circle(X2,X1,X22,X23) & circle(X0,X1,X20,X21)) => para(X4,X5,X18,X19))),
  inference(negated_conjecture,[],[f95])).
fof(f95,conjecture,(
  ( ! [X4,X19,X5,X18] : (! [X0,X1,X2,X3,X4,X5,X12,X13,X18,X19,X20,X21,X22,X23,X24,X25,X26,X27,X28,X29,X30,X31] : (circle(X2,X1,X19,X31) & coll(X1,X12,X19) & circle(X2,X3,X18,X30) & coll(X3,X13,X18) & circle(X0,X1,X13,X29) & coll(X1,X5,X13) & circle(X0,X3,X12,X28) & coll(X3,X4,X12) & circle(X2,X1,X5,X27) & circle(X2,X1,X4,X26) & circle(X2,X1,X3,X25) & circle(X0,X1,X3,X24) & circle(X2,X1,X22,X23) & circle(X0,X1,X20,X21)) => para(X4,X5,X18,X19)) )),
  file('./6_GDD_FULL/more/E022-13.p',exemplo6GDDFULLmoreE02213)).
fof(f641,plain,(
  ( ! [X2,X0,X3,X1] : (para(X0,X1,X2,X3)) )),
  inference(resolution,[],[f551,f529])).
fof(f529,plain,(
  ( ! [X2,X3,X1] : (perp(X1,X3,X3,X2)) )),
  inference(subsumption_resolution,[],[f528,f483])).
fof(f483,plain,(
  ( ! [X10,X5,X13] : (coll(X5,X10,X13)) )),
  inference(cnf_transformation,[],[f357])).
fof(f528,plain,(
  ( ! [X2,X0,X3,X1] : (~coll(X0,X1,X2) | perp(X1,X3,X3,X2)) )),
  inference(resolution,[],[f411,f484])).
fof(f484,plain,(
  ( ! [X6,X25,X5,X13] : (circle(X6,X5,X13,X25)) )),
  inference(cnf_transformation,[],[f357])).
fof(f411,plain,(
  ( ! [X2,X0,X3,X1] : (~circle(X3,X0,X1,X2) | ~coll(X3,X0,X2) | perp(X0,X1,X1,X2)) )),
  inference(cnf_transformation,[],[f245])).
fof(f245,plain,(
  ! [X0,X1,X2,X3] : (perp(X0,X1,X1,X2) | ~coll(X3,X0,X2) | ~circle(X3,X0,X1,X2))),
  inference(flattening,[],[f244])).
fof(f244,plain,(
  ! [X0,X1,X2,X3] : (perp(X0,X1,X1,X2) | (~coll(X3,X0,X2) | ~circle(X3,X0,X1,X2)))),
  inference(ennf_transformation,[],[f133])).
fof(f133,plain,(
  ! [X0,X1,X2,X3] : ((coll(X3,X0,X2) & circle(X3,X0,X1,X2)) => perp(X0,X1,X1,X2))),
  inference(rectify,[],[f54])).
fof(f54,axiom,(
  ! [X0,X1,X2,X7] : ((coll(X7,X0,X2) & circle(X7,X0,X1,X2)) => perp(X0,X1,X1,X2))),
  file('./6_GDD_FULL/more/E022-13.p',ruleD53)).
fof(f551,plain,(
  ( ! [X4,X2,X0,X3,X1] : (~perp(X0,X1,X4,X2) | para(X0,X1,X2,X3)) )),
  inference(resolution,[],[f366,f529])).
fof(f366,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (~perp(X2,X3,X4,X5) | para(X0,X1,X4,X5) | ~perp(X0,X1,X2,X3)) )),
  inference(cnf_transformation,[],[f179])).
fof(f179,plain,(
  ! [X0,X1,X2,X3,X4,X5] : (para(X0,X1,X4,X5) | ~perp(X2,X3,X4,X5) | ~perp(X0,X1,X2,X3))),
  inference(flattening,[],[f178])).
fof(f178,plain,(
  ! [X0,X1,X2,X3,X4,X5] : (para(X0,X1,X4,X5) | (~perp(X2,X3,X4,X5) | ~perp(X0,X1,X2,X3)))),
  inference(ennf_transformation,[],[f9])).
fof(f9,axiom,(
  ! [X0,X1,X2,X3,X4,X5] : ((perp(X2,X3,X4,X5) & perp(X0,X1,X2,X3)) => para(X0,X1,X4,X5))),
  file('./6_GDD_FULL/more/E022-13.p',ruleD9)).
% SZS output end Proof for E022-13
% ------------------------------
% Version: Vampire 4.4.0 (commit 7916d27 on 2019-08-23 08:50:16 +0100)
% Termination reason: Refutation

% Memory used [KB]: 5117
% Time elapsed: 0.031 s
% ------------------------------
% ------------------------------
