% Refutation found. Thanks to Tanya!
% SZS status Theorem for E023-17
% SZS output start Proof for E023-17
fof(f24222,plain,(
  $false),
  inference(avatar_sat_refutation,[],[f1023,f1106,f1237,f2255,f2415,f2831,f2850,f2958,f3455,f5906,f10255,f14378,f14442,f14492,f19588,f23236])).
fof(f23236,plain,(
  ~spl33_76 | ~spl33_204 | ~spl33_210 | ~spl33_320 | ~spl33_1244),
  inference(avatar_contradiction_clause,[],[f23235])).
fof(f23235,plain,(
  $false | (~spl33_76 | ~spl33_204 | ~spl33_210 | ~spl33_320 | ~spl33_1244)),
  inference(subsumption_resolution,[],[f487,f20006])).
fof(f20006,plain,(
  ( ! [X6,X4,X5,X3] : (para(X3,X4,X5,X6)) ) | (~spl33_76 | ~spl33_204 | ~spl33_210 | ~spl33_320 | ~spl33_1244)),
  inference(subsumption_resolution,[],[f19943,f19895])).
fof(f19895,plain,(
  ( ! [X10,X9] : (midp(X9,sK24,X10)) ) | (~spl33_76 | ~spl33_204 | ~spl33_210 | ~spl33_320 | ~spl33_1244)),
  inference(subsumption_resolution,[],[f19894,f19547])).
fof(f19547,plain,(
  ( ! [X4] : (midp(X4,sK28,sK25)) ) | (~spl33_204 | ~spl33_320)),
  inference(subsumption_resolution,[],[f19546,f9043])).
fof(f9043,plain,(
  ( ! [X21,X22] : (midp(X21,X22,X22)) ) | ~spl33_204),
  inference(subsumption_resolution,[],[f9010,f3575])).
fof(f3575,plain,(
  ( ! [X0,X1] : (coll(X0,X1,X1)) ) | ~spl33_204),
  inference(resolution,[],[f2251,f426])).
fof(f426,plain,(
  ( ! [X2,X0,X1] : (~para(X0,X1,X0,X2) | coll(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f268])).
fof(f268,plain,(
  ! [X0,X1,X2] : (coll(X0,X1,X2) | ~para(X0,X1,X0,X2))),
  inference(ennf_transformation,[],[f67])).
fof(f67,axiom,(
  ! [X0,X1,X2] : (para(X0,X1,X0,X2) => coll(X0,X1,X2))),
  file('./6_GDD_FULL/more/E023-17.p',ruleD66)).
fof(f2251,plain,(
  ( ! [X2,X3] : (para(X2,X3,X2,X3)) ) | ~spl33_204),
  inference(avatar_component_clause,[],[f2250])).
fof(f2250,plain,(
  spl33_204 <=> ! [X3,X2] : para(X2,X3,X2,X3)),
  introduced(avatar_definition,[new_symbols(naming,[spl33_204])])).
fof(f9010,plain,(
  ( ! [X21,X22] : (~coll(X21,X22,X22) | midp(X21,X22,X22)) ) | ~spl33_204),
  inference(resolution,[],[f8818,f427])).
fof(f427,plain,(
  ( ! [X2,X0,X1] : (~cong(X0,X1,X0,X2) | ~coll(X0,X1,X2) | midp(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f270])).
fof(f270,plain,(
  ! [X0,X1,X2] : (midp(X0,X1,X2) | ~coll(X0,X1,X2) | ~cong(X0,X1,X0,X2))),
  inference(flattening,[],[f269])).
fof(f269,plain,(
  ! [X0,X1,X2] : (midp(X0,X1,X2) | (~coll(X0,X1,X2) | ~cong(X0,X1,X0,X2)))),
  inference(ennf_transformation,[],[f68])).
fof(f68,axiom,(
  ! [X0,X1,X2] : ((coll(X0,X1,X2) & cong(X0,X1,X0,X2)) => midp(X0,X1,X2))),
  file('./6_GDD_FULL/more/E023-17.p',ruleD67)).
fof(f8818,plain,(
  ( ! [X35,X36] : (cong(X35,X36,X35,X36)) ) | ~spl33_204),
  inference(subsumption_resolution,[],[f8817,f8757])).
fof(f8757,plain,(
  ( ! [X4,X5,X3] : (cyclic(X3,X4,X3,X5)) ) | ~spl33_204),
  inference(resolution,[],[f8737,f374])).
fof(f374,plain,(
  ( ! [X2,X0,X3,X1] : (~cyclic(X0,X1,X2,X3) | cyclic(X0,X2,X1,X3)) )),
  inference(cnf_transformation,[],[f189])).
fof(f189,plain,(
  ! [X0,X1,X2,X3] : (cyclic(X0,X2,X1,X3) | ~cyclic(X0,X1,X2,X3))),
  inference(ennf_transformation,[],[f15])).
fof(f15,axiom,(
  ! [X0,X1,X2,X3] : (cyclic(X0,X1,X2,X3) => cyclic(X0,X2,X1,X3))),
  file('./6_GDD_FULL/more/E023-17.p',ruleD15)).
fof(f8737,plain,(
  ( ! [X2,X0,X1] : (cyclic(X0,X0,X1,X2)) ) | ~spl33_204),
  inference(resolution,[],[f2837,f2251])).
fof(f2837,plain,(
  ( ! [X68,X66,X67] : (~para(X66,X67,X66,X67) | cyclic(X67,X67,X66,X68)) )),
  inference(subsumption_resolution,[],[f2398,f2399])).
fof(f2399,plain,(
  ( ! [X70,X71,X69] : (~para(X69,X70,X69,X70) | coll(X69,X71,X70) | cyclic(X70,X70,X69,X71)) )),
  inference(resolution,[],[f805,f401])).
fof(f401,plain,(
  ( ! [X2,X0,X3,X1] : (~eqangle(X2,X0,X2,X1,X3,X0,X3,X1) | coll(X2,X3,X0) | cyclic(X0,X1,X2,X3)) )),
  inference(cnf_transformation,[],[f223])).
fof(f223,plain,(
  ! [X0,X1,X2,X3] : (cyclic(X0,X1,X2,X3) | coll(X2,X3,X0) | ~eqangle(X2,X0,X2,X1,X3,X0,X3,X1))),
  inference(flattening,[],[f222])).
fof(f222,plain,(
  ! [X0,X1,X2,X3] : (cyclic(X0,X1,X2,X3) | (coll(X2,X3,X0) | ~eqangle(X2,X0,X2,X1,X3,X0,X3,X1)))),
  inference(ennf_transformation,[],[f121])).
fof(f121,plain,(
  ! [X0,X1,X2,X3] : ((~coll(X2,X3,X0) & eqangle(X2,X0,X2,X1,X3,X0,X3,X1)) => cyclic(X0,X1,X2,X3))),
  inference(rectify,[],[f42])).
fof(f42,axiom,(
  ! [X0,X1,X8,X9] : ((~coll(X8,X9,X0) & eqangle(X8,X0,X8,X1,X9,X0,X9,X1)) => cyclic(X0,X1,X8,X9))),
  file('./6_GDD_FULL/more/E023-17.p',ruleD42a)).
fof(f805,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (eqangle(X0,X1,X2,X3,X4,X5,X4,X5) | ~para(X0,X1,X2,X3)) )),
  inference(resolution,[],[f380,f399])).
fof(f399,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (eqangle(X0,X1,X4,X5,X2,X3,X4,X5) | ~para(X0,X1,X2,X3)) )),
  inference(cnf_transformation,[],[f220])).
fof(f220,plain,(
  ! [X0,X1,X2,X3,X4,X5] : (eqangle(X0,X1,X4,X5,X2,X3,X4,X5) | ~para(X0,X1,X2,X3))),
  inference(ennf_transformation,[],[f119])).
fof(f119,plain,(
  ! [X0,X1,X2,X3,X4,X5] : (para(X0,X1,X2,X3) => eqangle(X0,X1,X4,X5,X2,X3,X4,X5))),
  inference(rectify,[],[f40])).
fof(f40,axiom,(
  ! [X0,X1,X2,X3,X8,X9] : (para(X0,X1,X2,X3) => eqangle(X0,X1,X8,X9,X2,X3,X8,X9))),
  file('./6_GDD_FULL/more/E023-17.p',ruleD40)).
fof(f380,plain,(
  ( ! [X6,X4,X2,X0,X7,X5,X3,X1] : (~eqangle(X0,X1,X2,X3,X4,X5,X6,X7) | eqangle(X0,X1,X4,X5,X2,X3,X6,X7)) )),
  inference(cnf_transformation,[],[f196])).
fof(f196,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6,X7] : (eqangle(X0,X1,X4,X5,X2,X3,X6,X7) | ~eqangle(X0,X1,X2,X3,X4,X5,X6,X7))),
  inference(ennf_transformation,[],[f103])).
fof(f103,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6,X7] : (eqangle(X0,X1,X2,X3,X4,X5,X6,X7) => eqangle(X0,X1,X4,X5,X2,X3,X6,X7))),
  inference(rectify,[],[f21])).
fof(f21,axiom,(
  ! [X0,X1,X2,X3,X8,X9,X10,X11] : (eqangle(X0,X1,X2,X3,X8,X9,X10,X11) => eqangle(X0,X1,X8,X9,X2,X3,X10,X11))),
  file('./6_GDD_FULL/more/E023-17.p',ruleD21)).
fof(f2398,plain,(
  ( ! [X68,X66,X67] : (~para(X66,X67,X66,X67) | ~coll(X66,X68,X67) | cyclic(X67,X67,X66,X68)) )),
  inference(resolution,[],[f805,f402])).
fof(f402,plain,(
  ( ! [X2,X0,X3,X1] : (~eqangle(X2,X0,X2,X1,X3,X0,X3,X1) | ~coll(X2,X3,X1) | cyclic(X0,X1,X2,X3)) )),
  inference(cnf_transformation,[],[f225])).
fof(f225,plain,(
  ! [X0,X1,X2,X3] : (cyclic(X0,X1,X2,X3) | ~coll(X2,X3,X1) | ~eqangle(X2,X0,X2,X1,X3,X0,X3,X1))),
  inference(flattening,[],[f224])).
fof(f224,plain,(
  ! [X0,X1,X2,X3] : (cyclic(X0,X1,X2,X3) | (~coll(X2,X3,X1) | ~eqangle(X2,X0,X2,X1,X3,X0,X3,X1)))),
  inference(ennf_transformation,[],[f122])).
fof(f122,plain,(
  ! [X0,X1,X2,X3] : ((coll(X2,X3,X1) & eqangle(X2,X0,X2,X1,X3,X0,X3,X1)) => cyclic(X0,X1,X2,X3))),
  inference(rectify,[],[f43])).
fof(f43,axiom,(
  ! [X0,X1,X8,X9] : ((coll(X8,X9,X1) & eqangle(X8,X0,X8,X1,X9,X0,X9,X1)) => cyclic(X0,X1,X8,X9))),
  file('./6_GDD_FULL/more/E023-17.p',ruleD42b)).
fof(f8817,plain,(
  ( ! [X35,X36] : (cong(X35,X36,X35,X36) | ~cyclic(X35,X36,X35,X35)) ) | ~spl33_204),
  inference(subsumption_resolution,[],[f8803,f8757])).
fof(f8803,plain,(
  ( ! [X35,X36] : (cong(X35,X36,X35,X36) | ~cyclic(X35,X36,X35,X36) | ~cyclic(X35,X36,X35,X35)) ) | ~spl33_204),
  inference(resolution,[],[f8757,f1522])).
fof(f1522,plain,(
  ( ! [X10,X8,X7,X9] : (~cyclic(X7,X8,X9,X10) | cong(X7,X8,X7,X8) | ~cyclic(X7,X8,X9,X8) | ~cyclic(X7,X8,X9,X7)) )),
  inference(duplicate_literal_removal,[],[f1516])).
fof(f1516,plain,(
  ( ! [X10,X8,X7,X9] : (cong(X7,X8,X7,X8) | ~cyclic(X7,X8,X9,X10) | ~cyclic(X7,X8,X9,X8) | ~cyclic(X7,X8,X9,X7) | ~cyclic(X7,X8,X9,X10)) )),
  inference(resolution,[],[f403,f400])).
fof(f400,plain,(
  ( ! [X2,X0,X3,X1] : (eqangle(X2,X0,X2,X1,X3,X0,X3,X1) | ~cyclic(X0,X1,X2,X3)) )),
  inference(cnf_transformation,[],[f221])).
fof(f221,plain,(
  ! [X0,X1,X2,X3] : (eqangle(X2,X0,X2,X1,X3,X0,X3,X1) | ~cyclic(X0,X1,X2,X3))),
  inference(ennf_transformation,[],[f120])).
fof(f120,plain,(
  ! [X0,X1,X2,X3] : (cyclic(X0,X1,X2,X3) => eqangle(X2,X0,X2,X1,X3,X0,X3,X1))),
  inference(rectify,[],[f41])).
fof(f41,axiom,(
  ! [X0,X1,X8,X9] : (cyclic(X0,X1,X8,X9) => eqangle(X8,X0,X8,X1,X9,X0,X9,X1))),
  file('./6_GDD_FULL/more/E023-17.p',ruleD41)).
fof(f403,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (~eqangle(X2,X0,X2,X1,X5,X3,X5,X4) | cong(X0,X1,X3,X4) | ~cyclic(X0,X1,X2,X5) | ~cyclic(X0,X1,X2,X4) | ~cyclic(X0,X1,X2,X3)) )),
  inference(cnf_transformation,[],[f227])).
fof(f227,plain,(
  ! [X0,X1,X2,X3,X4,X5] : (cong(X0,X1,X3,X4) | ~eqangle(X2,X0,X2,X1,X5,X3,X5,X4) | ~cyclic(X0,X1,X2,X5) | ~cyclic(X0,X1,X2,X4) | ~cyclic(X0,X1,X2,X3))),
  inference(flattening,[],[f226])).
fof(f226,plain,(
  ! [X0,X1,X2,X3,X4,X5] : (cong(X0,X1,X3,X4) | (~eqangle(X2,X0,X2,X1,X5,X3,X5,X4) | ~cyclic(X0,X1,X2,X5) | ~cyclic(X0,X1,X2,X4) | ~cyclic(X0,X1,X2,X3)))),
  inference(ennf_transformation,[],[f123])).
fof(f123,plain,(
  ! [X0,X1,X2,X3,X4,X5] : ((eqangle(X2,X0,X2,X1,X5,X3,X5,X4) & cyclic(X0,X1,X2,X5) & cyclic(X0,X1,X2,X4) & cyclic(X0,X1,X2,X3)) => cong(X0,X1,X3,X4))),
  inference(rectify,[],[f44])).
fof(f44,axiom,(
  ! [X0,X1,X2,X8,X9,X14] : ((eqangle(X2,X0,X2,X1,X14,X8,X14,X9) & cyclic(X0,X1,X2,X14) & cyclic(X0,X1,X2,X9) & cyclic(X0,X1,X2,X8)) => cong(X0,X1,X8,X9))),
  file('./6_GDD_FULL/more/E023-17.p',ruleD43)).
fof(f19546,plain,(
  ( ! [X4] : (midp(X4,sK28,sK25) | ~midp(X4,sK20,sK20)) ) | ~spl33_320),
  inference(subsumption_resolution,[],[f19536,f19533])).
fof(f19533,plain,(
  para(sK20,sK28,sK20,sK25) | ~spl33_320),
  inference(resolution,[],[f15891,f364])).
fof(f364,plain,(
  ( ! [X2,X0,X3,X1] : (~para(X0,X1,X2,X3) | para(X2,X3,X0,X1)) )),
  inference(cnf_transformation,[],[f174])).
fof(f174,plain,(
  ! [X0,X1,X2,X3] : (para(X2,X3,X0,X1) | ~para(X0,X1,X2,X3))),
  inference(ennf_transformation,[],[f5])).
fof(f5,axiom,(
  ! [X0,X1,X2,X3] : (para(X0,X1,X2,X3) => para(X2,X3,X0,X1))),
  file('./6_GDD_FULL/more/E023-17.p',ruleD5)).
fof(f15891,plain,(
  para(sK20,sK25,sK20,sK28) | ~spl33_320),
  inference(resolution,[],[f3454,f730])).
fof(f730,plain,(
  ( ! [X8,X9] : (~perp(X8,X9,sK25,sK24) | para(X8,X9,sK20,sK28)) )),
  inference(resolution,[],[f368,f482])).
fof(f482,plain,(
  perp(sK25,sK24,sK20,sK28)),
  inference(cnf_transformation,[],[f359])).
fof(f359,plain,(
  ~para(sK26,sK27,sK29,sK30) & circle(sK28,sK23,sK30,sK32) & coll(sK23,sK26,sK30) & circle(sK28,sK24,sK29,sK31) & coll(sK24,sK27,sK29) & perp(sK25,sK24,sK20,sK28) & midp(sK20,sK25,sK24) & perp(sK23,sK24,sK21,sK28) & midp(sK21,sK23,sK24) & perp(sK23,sK25,sK22,sK28) & midp(sK22,sK23,sK25) & coll(sK27,sK23,sK25) & perp(sK27,sK24,sK23,sK25) & coll(sK26,sK24,sK25) & perp(sK26,sK23,sK24,sK25)),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK20,sK21,sK22,sK23,sK24,sK25,sK26,sK27,sK28,sK29,sK30,sK31,sK32])],[f321,f358])).
fof(f358,plain,(
  ? [X0,X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12] : (~para(X6,X7,X9,X10) & circle(X8,X3,X10,X12) & coll(X3,X6,X10) & circle(X8,X4,X9,X11) & coll(X4,X7,X9) & perp(X5,X4,X0,X8) & midp(X0,X5,X4) & perp(X3,X4,X1,X8) & midp(X1,X3,X4) & perp(X3,X5,X2,X8) & midp(X2,X3,X5) & coll(X7,X3,X5) & perp(X7,X4,X3,X5) & coll(X6,X4,X5) & perp(X6,X3,X4,X5)) => (~para(sK26,sK27,sK29,sK30) & circle(sK28,sK23,sK30,sK32) & coll(sK23,sK26,sK30) & circle(sK28,sK24,sK29,sK31) & coll(sK24,sK27,sK29) & perp(sK25,sK24,sK20,sK28) & midp(sK20,sK25,sK24) & perp(sK23,sK24,sK21,sK28) & midp(sK21,sK23,sK24) & perp(sK23,sK25,sK22,sK28) & midp(sK22,sK23,sK25) & coll(sK27,sK23,sK25) & perp(sK27,sK24,sK23,sK25) & coll(sK26,sK24,sK25) & perp(sK26,sK23,sK24,sK25))),
  introduced(choice_axiom,[])).
fof(f321,plain,(
  ? [X0,X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12] : (~para(X6,X7,X9,X10) & circle(X8,X3,X10,X12) & coll(X3,X6,X10) & circle(X8,X4,X9,X11) & coll(X4,X7,X9) & perp(X5,X4,X0,X8) & midp(X0,X5,X4) & perp(X3,X4,X1,X8) & midp(X1,X3,X4) & perp(X3,X5,X2,X8) & midp(X2,X3,X5) & coll(X7,X3,X5) & perp(X7,X4,X3,X5) & coll(X6,X4,X5) & perp(X6,X3,X4,X5))),
  inference(flattening,[],[f320])).
fof(f320,plain,(
  ? [X0,X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12] : (~para(X6,X7,X9,X10) & (circle(X8,X3,X10,X12) & coll(X3,X6,X10) & circle(X8,X4,X9,X11) & coll(X4,X7,X9) & perp(X5,X4,X0,X8) & midp(X0,X5,X4) & perp(X3,X4,X1,X8) & midp(X1,X3,X4) & perp(X3,X5,X2,X8) & midp(X2,X3,X5) & coll(X7,X3,X5) & perp(X7,X4,X3,X5) & coll(X6,X4,X5) & perp(X6,X3,X4,X5)))),
  inference(ennf_transformation,[],[f168])).
fof(f168,plain,(
  ~! [X0,X1,X2,X3,X4,X5,X6,X7,X8,X9,X10,X11,X12] : ((circle(X8,X3,X10,X12) & coll(X3,X6,X10) & circle(X8,X4,X9,X11) & coll(X4,X7,X9) & perp(X5,X4,X0,X8) & midp(X0,X5,X4) & perp(X3,X4,X1,X8) & midp(X1,X3,X4) & perp(X3,X5,X2,X8) & midp(X2,X3,X5) & coll(X7,X3,X5) & perp(X7,X4,X3,X5) & coll(X6,X4,X5) & perp(X6,X3,X4,X5)) => para(X6,X7,X9,X10))),
  inference(flattening,[],[f167])).
fof(f167,plain,(
  ~! [X0,X1,X2] : ! [X3,X4,X5,X6,X7,X8,X9,X10,X11,X12] : ((circle(X8,X3,X10,X12) & coll(X3,X6,X10) & circle(X8,X4,X9,X11) & coll(X4,X7,X9) & perp(X5,X4,X0,X8) & midp(X0,X5,X4) & perp(X3,X4,X1,X8) & midp(X1,X3,X4) & perp(X3,X5,X2,X8) & midp(X2,X3,X5) & coll(X7,X3,X5) & perp(X7,X4,X3,X5) & coll(X6,X4,X5) & perp(X6,X3,X4,X5)) => para(X6,X7,X9,X10))),
  inference(rectify,[],[f96])).
fof(f96,negated_conjecture,(
  ~! [X22,X21,X20] : ! [X0,X1,X2,X3,X4,X5,X12,X13,X18,X19] : ((circle(X5,X0,X13,X19) & coll(X0,X3,X13) & circle(X5,X1,X12,X18) & coll(X1,X4,X12) & perp(X2,X1,X22,X5) & midp(X22,X2,X1) & perp(X0,X1,X21,X5) & midp(X21,X0,X1) & perp(X0,X2,X20,X5) & midp(X20,X0,X2) & coll(X4,X0,X2) & perp(X4,X1,X0,X2) & coll(X3,X1,X2) & perp(X3,X0,X1,X2)) => para(X3,X4,X12,X13))),
  inference(negated_conjecture,[],[f95])).
fof(f95,conjecture,(
  ( ! [X21,X22,X20] : (! [X0,X1,X2,X3,X4,X5,X12,X13,X18,X19] : ((circle(X5,X0,X13,X19) & coll(X0,X3,X13) & circle(X5,X1,X12,X18) & coll(X1,X4,X12) & perp(X2,X1,X22,X5) & midp(X22,X2,X1) & perp(X0,X1,X21,X5) & midp(X21,X0,X1) & perp(X0,X2,X20,X5) & midp(X20,X0,X2) & coll(X4,X0,X2) & perp(X4,X1,X0,X2) & coll(X3,X1,X2) & perp(X3,X0,X1,X2)) => para(X3,X4,X12,X13))) )),
  file('./6_GDD_FULL/more/E023-17.p',exemplo6GDDFULLmoreE02317)).
fof(f368,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (~perp(X2,X3,X4,X5) | para(X0,X1,X4,X5) | ~perp(X0,X1,X2,X3)) )),
  inference(cnf_transformation,[],[f180])).
fof(f180,plain,(
  ! [X0,X1,X2,X3,X4,X5] : (para(X0,X1,X4,X5) | ~perp(X2,X3,X4,X5) | ~perp(X0,X1,X2,X3))),
  inference(flattening,[],[f179])).
fof(f179,plain,(
  ! [X0,X1,X2,X3,X4,X5] : (para(X0,X1,X4,X5) | (~perp(X2,X3,X4,X5) | ~perp(X0,X1,X2,X3)))),
  inference(ennf_transformation,[],[f9])).
fof(f9,axiom,(
  ! [X0,X1,X2,X3,X4,X5] : ((perp(X2,X3,X4,X5) & perp(X0,X1,X2,X3)) => para(X0,X1,X4,X5))),
  file('./6_GDD_FULL/more/E023-17.p',ruleD9)).
fof(f3454,plain,(
  perp(sK20,sK25,sK25,sK24) | ~spl33_320),
  inference(avatar_component_clause,[],[f3453])).
fof(f3453,plain,(
  spl33_320 <=> perp(sK20,sK25,sK25,sK24)),
  introduced(avatar_definition,[new_symbols(naming,[spl33_320])])).
fof(f19536,plain,(
  ( ! [X4] : (midp(X4,sK28,sK25) | ~para(sK20,sK28,sK20,sK25) | ~midp(X4,sK20,sK20)) ) | ~spl33_320),
  inference(resolution,[],[f15891,f424])).
fof(f424,plain,(
  ( ! [X4,X2,X0,X3,X1] : (~para(X0,X3,X1,X2) | midp(X4,X2,X3) | ~para(X0,X2,X1,X3) | ~midp(X4,X0,X1)) )),
  inference(cnf_transformation,[],[f265])).
fof(f265,plain,(
  ! [X0,X1,X2,X3,X4] : (midp(X4,X2,X3) | ~para(X0,X3,X1,X2) | ~para(X0,X2,X1,X3) | ~midp(X4,X0,X1))),
  inference(flattening,[],[f264])).
fof(f264,plain,(
  ! [X0,X1,X2,X3,X4] : (midp(X4,X2,X3) | (~para(X0,X3,X1,X2) | ~para(X0,X2,X1,X3) | ~midp(X4,X0,X1)))),
  inference(ennf_transformation,[],[f143])).
fof(f143,plain,(
  ! [X0,X1,X2,X3,X4] : ((para(X0,X3,X1,X2) & para(X0,X2,X1,X3) & midp(X4,X0,X1)) => midp(X4,X2,X3))),
  inference(rectify,[],[f65])).
fof(f65,axiom,(
  ! [X0,X1,X2,X3,X6] : ((para(X0,X3,X1,X2) & para(X0,X2,X1,X3) & midp(X6,X0,X1)) => midp(X6,X2,X3))),
  file('./6_GDD_FULL/more/E023-17.p',ruleD64)).
fof(f19894,plain,(
  ( ! [X10,X9] : (midp(X9,sK24,X10) | ~midp(X9,sK28,sK25)) ) | (~spl33_76 | ~spl33_204 | ~spl33_210 | ~spl33_1244)),
  inference(subsumption_resolution,[],[f19883,f19587])).
fof(f19587,plain,(
  ( ! [X48,X49] : (para(sK28,X49,sK25,X48)) ) | ~spl33_1244),
  inference(avatar_component_clause,[],[f19586])).
fof(f19586,plain,(
  spl33_1244 <=> ! [X49,X48] : para(sK28,X49,sK25,X48)),
  introduced(avatar_definition,[new_symbols(naming,[spl33_1244])])).
fof(f19883,plain,(
  ( ! [X10,X9] : (midp(X9,sK24,X10) | ~para(sK28,sK24,sK25,X10) | ~midp(X9,sK28,sK25)) ) | (~spl33_76 | ~spl33_204 | ~spl33_210)),
  inference(resolution,[],[f16405,f424])).
fof(f16405,plain,(
  ( ! [X12] : (para(sK28,X12,sK25,sK24)) ) | (~spl33_76 | ~spl33_204 | ~spl33_210)),
  inference(resolution,[],[f14511,f2150])).
fof(f2150,plain,(
  ( ! [X2,X0,X3,X1] : (~para(X0,X1,X2,X3) | para(X1,X0,X2,X3)) )),
  inference(resolution,[],[f769,f398])).
fof(f398,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (~eqangle(X0,X1,X4,X5,X2,X3,X4,X5) | para(X0,X1,X2,X3)) )),
  inference(cnf_transformation,[],[f219])).
fof(f219,plain,(
  ! [X0,X1,X2,X3,X4,X5] : (para(X0,X1,X2,X3) | ~eqangle(X0,X1,X4,X5,X2,X3,X4,X5))),
  inference(ennf_transformation,[],[f118])).
fof(f118,plain,(
  ! [X0,X1,X2,X3,X4,X5] : (eqangle(X0,X1,X4,X5,X2,X3,X4,X5) => para(X0,X1,X2,X3))),
  inference(rectify,[],[f39])).
fof(f39,axiom,(
  ! [X0,X1,X2,X3,X8,X9] : (eqangle(X0,X1,X8,X9,X2,X3,X8,X9) => para(X0,X1,X2,X3))),
  file('./6_GDD_FULL/more/E023-17.p',ruleD39)).
fof(f769,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (eqangle(X0,X1,X2,X3,X4,X5,X2,X3) | ~para(X1,X0,X4,X5)) )),
  inference(resolution,[],[f377,f399])).
fof(f377,plain,(
  ( ! [X6,X4,X2,X0,X7,X5,X3,X1] : (~eqangle(X0,X1,X2,X3,X4,X5,X6,X7) | eqangle(X1,X0,X2,X3,X4,X5,X6,X7)) )),
  inference(cnf_transformation,[],[f193])).
fof(f193,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6,X7] : (eqangle(X1,X0,X2,X3,X4,X5,X6,X7) | ~eqangle(X0,X1,X2,X3,X4,X5,X6,X7))),
  inference(ennf_transformation,[],[f100])).
fof(f100,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6,X7] : (eqangle(X0,X1,X2,X3,X4,X5,X6,X7) => eqangle(X1,X0,X2,X3,X4,X5,X6,X7))),
  inference(rectify,[],[f18])).
fof(f18,axiom,(
  ! [X0,X1,X2,X3,X8,X9,X10,X11] : (eqangle(X0,X1,X2,X3,X8,X9,X10,X11) => eqangle(X1,X0,X2,X3,X8,X9,X10,X11))),
  file('./6_GDD_FULL/more/E023-17.p',ruleD18)).
fof(f14511,plain,(
  ( ! [X70] : (para(X70,sK28,sK25,sK24)) ) | (~spl33_76 | ~spl33_204 | ~spl33_210)),
  inference(resolution,[],[f14491,f1007])).
fof(f1007,plain,(
  ( ! [X2,X3] : (~perp(X2,X3,sK20,sK28) | para(X2,X3,sK25,sK24)) )),
  inference(resolution,[],[f529,f368])).
fof(f529,plain,(
  perp(sK20,sK28,sK25,sK24)),
  inference(resolution,[],[f367,f482])).
fof(f367,plain,(
  ( ! [X2,X0,X3,X1] : (~perp(X0,X1,X2,X3) | perp(X2,X3,X0,X1)) )),
  inference(cnf_transformation,[],[f178])).
fof(f178,plain,(
  ! [X0,X1,X2,X3] : (perp(X2,X3,X0,X1) | ~perp(X0,X1,X2,X3))),
  inference(ennf_transformation,[],[f8])).
fof(f8,axiom,(
  ! [X0,X1,X2,X3] : (perp(X0,X1,X2,X3) => perp(X2,X3,X0,X1))),
  file('./6_GDD_FULL/more/E023-17.p',ruleD8)).
fof(f14491,plain,(
  ( ! [X85,X86,X84] : (perp(X84,X85,X86,X85)) ) | (~spl33_76 | ~spl33_204 | ~spl33_210)),
  inference(subsumption_resolution,[],[f14490,f12864])).
fof(f12864,plain,(
  ( ! [X19,X17,X18] : (cong(X18,X17,X19,X17)) ) | (~spl33_204 | ~spl33_210)),
  inference(subsumption_resolution,[],[f12832,f9043])).
fof(f12832,plain,(
  ( ! [X19,X17,X18] : (~midp(X17,X18,X18) | cong(X18,X17,X19,X17)) ) | (~spl33_204 | ~spl33_210)),
  inference(resolution,[],[f10289,f412])).
fof(f412,plain,(
  ( ! [X2,X0,X3,X1] : (~perp(X0,X1,X1,X2) | ~midp(X3,X0,X2) | cong(X0,X3,X1,X3)) )),
  inference(cnf_transformation,[],[f244])).
fof(f244,plain,(
  ! [X0,X1,X2,X3] : (cong(X0,X3,X1,X3) | ~midp(X3,X0,X2) | ~perp(X0,X1,X1,X2))),
  inference(flattening,[],[f243])).
fof(f243,plain,(
  ! [X0,X1,X2,X3] : (cong(X0,X3,X1,X3) | (~midp(X3,X0,X2) | ~perp(X0,X1,X1,X2)))),
  inference(ennf_transformation,[],[f132])).
fof(f132,plain,(
  ! [X0,X1,X2,X3] : ((midp(X3,X0,X2) & perp(X0,X1,X1,X2)) => cong(X0,X3,X1,X3))),
  inference(rectify,[],[f53])).
fof(f53,axiom,(
  ! [X0,X1,X2,X6] : ((midp(X6,X0,X2) & perp(X0,X1,X1,X2)) => cong(X0,X6,X1,X6))),
  file('./6_GDD_FULL/more/E023-17.p',ruleD52)).
fof(f10289,plain,(
  ( ! [X8,X9] : (perp(X8,X9,X9,X8)) ) | (~spl33_204 | ~spl33_210)),
  inference(resolution,[],[f2414,f3582])).
fof(f3582,plain,(
  ( ! [X21,X20] : (para(X20,X21,X21,X20)) ) | ~spl33_204),
  inference(resolution,[],[f2251,f363])).
fof(f363,plain,(
  ( ! [X2,X0,X3,X1] : (~para(X0,X1,X2,X3) | para(X0,X1,X3,X2)) )),
  inference(cnf_transformation,[],[f173])).
fof(f173,plain,(
  ! [X0,X1,X2,X3] : (para(X0,X1,X3,X2) | ~para(X0,X1,X2,X3))),
  inference(ennf_transformation,[],[f4])).
fof(f4,axiom,(
  ! [X0,X1,X2,X3] : (para(X0,X1,X2,X3) => para(X0,X1,X3,X2))),
  file('./6_GDD_FULL/more/E023-17.p',ruleD4)).
fof(f2414,plain,(
  ( ! [X47,X45,X46,X44] : (~para(X44,X45,X46,X47) | perp(X44,X45,X46,X47)) ) | ~spl33_210),
  inference(avatar_component_clause,[],[f2413])).
fof(f2413,plain,(
  spl33_210 <=> ! [X44,X46,X45,X47] : (~para(X44,X45,X46,X47) | perp(X44,X45,X46,X47))),
  introduced(avatar_definition,[new_symbols(naming,[spl33_210])])).
fof(f14490,plain,(
  ( ! [X85,X86,X84] : (perp(X84,X85,X86,X85) | ~cong(X84,X86,X85,X86)) ) | (~spl33_76 | ~spl33_204)),
  inference(subsumption_resolution,[],[f14473,f9043])).
fof(f14473,plain,(
  ( ! [X85,X86,X84] : (~midp(X84,X85,X85) | perp(X84,X85,X86,X85) | ~cong(X84,X86,X85,X86)) ) | ~spl33_76),
  inference(resolution,[],[f1102,f416])).
fof(f416,plain,(
  ( ! [X2,X0,X3,X1] : (~cong(X0,X3,X1,X3) | perp(X0,X1,X2,X3) | ~cong(X0,X2,X1,X2)) )),
  inference(cnf_transformation,[],[f252])).
fof(f252,plain,(
  ! [X0,X1,X2,X3] : (perp(X0,X1,X2,X3) | ~cong(X0,X3,X1,X3) | ~cong(X0,X2,X1,X2))),
  inference(flattening,[],[f251])).
fof(f251,plain,(
  ! [X0,X1,X2,X3] : (perp(X0,X1,X2,X3) | (~cong(X0,X3,X1,X3) | ~cong(X0,X2,X1,X2)))),
  inference(ennf_transformation,[],[f135])).
fof(f135,plain,(
  ! [X0,X1,X2,X3] : ((cong(X0,X3,X1,X3) & cong(X0,X2,X1,X2)) => perp(X0,X1,X2,X3))),
  inference(rectify,[],[f57])).
fof(f57,axiom,(
  ! [X0,X1,X8,X9] : ((cong(X0,X9,X1,X9) & cong(X0,X8,X1,X8)) => perp(X0,X1,X8,X9))),
  file('./6_GDD_FULL/more/E023-17.p',ruleD56)).
fof(f1102,plain,(
  ( ! [X10,X11,X9] : (cong(X9,X10,X10,X11) | ~midp(X9,X10,X11)) ) | ~spl33_76),
  inference(avatar_component_clause,[],[f1101])).
fof(f1101,plain,(
  spl33_76 <=> ! [X9,X11,X10] : (cong(X9,X10,X10,X11) | ~midp(X9,X10,X11))),
  introduced(avatar_definition,[new_symbols(naming,[spl33_76])])).
fof(f19943,plain,(
  ( ! [X6,X4,X5,X3] : (para(X3,X4,X5,X6) | ~midp(X3,sK24,X5)) ) | (~spl33_76 | ~spl33_204 | ~spl33_210 | ~spl33_320 | ~spl33_1244)),
  inference(resolution,[],[f19895,f404])).
fof(f404,plain,(
  ( ! [X4,X2,X0,X3,X1] : (~midp(X4,X0,X2) | para(X3,X4,X1,X2) | ~midp(X3,X0,X1)) )),
  inference(cnf_transformation,[],[f229])).
fof(f229,plain,(
  ! [X0,X1,X2,X3,X4] : (para(X3,X4,X1,X2) | ~midp(X4,X0,X2) | ~midp(X3,X0,X1))),
  inference(flattening,[],[f228])).
fof(f228,plain,(
  ! [X0,X1,X2,X3,X4] : (para(X3,X4,X1,X2) | (~midp(X4,X0,X2) | ~midp(X3,X0,X1)))),
  inference(ennf_transformation,[],[f124])).
fof(f124,plain,(
  ! [X0,X1,X2,X3,X4] : ((midp(X4,X0,X2) & midp(X3,X0,X1)) => para(X3,X4,X1,X2))),
  inference(rectify,[],[f45])).
fof(f45,axiom,(
  ! [X0,X1,X2,X4,X5] : ((midp(X5,X0,X2) & midp(X4,X0,X1)) => para(X4,X5,X1,X2))),
  file('./6_GDD_FULL/more/E023-17.p',ruleD44)).
fof(f487,plain,(
  ~para(sK26,sK27,sK29,sK30)),
  inference(cnf_transformation,[],[f359])).
fof(f19588,plain,(
  spl33_272 | spl33_1244 | ~spl33_204 | ~spl33_210 | ~spl33_274 | ~spl33_320),
  inference(avatar_split_clause,[],[f19584,f3453,f2975,f2413,f2250,f19586,f2947])).
fof(f2947,plain,(
  spl33_272 <=> ! [X11,X10,X12] : ~midp(X10,X11,X12)),
  introduced(avatar_definition,[new_symbols(naming,[spl33_272])])).
fof(f2975,plain,(
  spl33_274 <=> ! [X13,X15,X14] : (~midp(X13,X14,X15) | cong(X14,X15,X13,X14))),
  introduced(avatar_definition,[new_symbols(naming,[spl33_274])])).
fof(f19584,plain,(
  ( ! [X52,X50,X48,X51,X49] : (para(sK28,X49,sK25,X48) | ~midp(X52,X51,X50)) ) | (~spl33_204 | ~spl33_210 | ~spl33_274 | ~spl33_320)),
  inference(subsumption_resolution,[],[f19583,f18120])).
fof(f18120,plain,(
  ( ! [X35,X33,X34] : (perp(X33,X34,X33,X35)) ) | (~spl33_204 | ~spl33_210 | ~spl33_274)),
  inference(resolution,[],[f14590,f366])).
fof(f366,plain,(
  ( ! [X2,X0,X3,X1] : (~perp(X0,X1,X2,X3) | perp(X0,X1,X3,X2)) )),
  inference(cnf_transformation,[],[f177])).
fof(f177,plain,(
  ! [X0,X1,X2,X3] : (perp(X0,X1,X3,X2) | ~perp(X0,X1,X2,X3))),
  inference(ennf_transformation,[],[f7])).
fof(f7,axiom,(
  ! [X0,X1,X2,X3] : (perp(X0,X1,X2,X3) => perp(X0,X1,X3,X2))),
  file('./6_GDD_FULL/more/E023-17.p',ruleD7)).
fof(f14590,plain,(
  ( ! [X85,X83,X84] : (perp(X84,X83,X85,X84)) ) | (~spl33_204 | ~spl33_210 | ~spl33_274)),
  inference(subsumption_resolution,[],[f14589,f12864])).
fof(f14589,plain,(
  ( ! [X85,X83,X84] : (perp(X84,X83,X85,X84) | ~cong(X84,X85,X83,X85)) ) | (~spl33_204 | ~spl33_274)),
  inference(subsumption_resolution,[],[f14566,f9043])).
fof(f14566,plain,(
  ( ! [X85,X83,X84] : (~midp(X83,X84,X84) | perp(X84,X83,X85,X84) | ~cong(X84,X85,X83,X85)) ) | ~spl33_274),
  inference(resolution,[],[f2976,f416])).
fof(f2976,plain,(
  ( ! [X14,X15,X13] : (cong(X14,X15,X13,X14) | ~midp(X13,X14,X15)) ) | ~spl33_274),
  inference(avatar_component_clause,[],[f2975])).
fof(f19583,plain,(
  ( ! [X52,X50,X48,X51,X49] : (~perp(X49,X51,X49,X48) | para(sK28,X49,sK25,X48) | ~midp(X52,X51,X50)) ) | (~spl33_204 | ~spl33_210 | ~spl33_274 | ~spl33_320)),
  inference(subsumption_resolution,[],[f19567,f18120])).
fof(f19567,plain,(
  ( ! [X52,X50,X48,X51,X49] : (~perp(X48,X49,X48,X50) | ~perp(X49,X51,X49,X48) | para(sK28,X49,sK25,X48) | ~midp(X52,X51,X50)) ) | (~spl33_204 | ~spl33_320)),
  inference(resolution,[],[f19547,f1374])).
fof(f1374,plain,(
  ( ! [X6,X4,X2,X0,X5,X3,X1] : (~midp(sK17(X4,X1,X3,X2,X0),X5,X6) | ~perp(X3,X4,X3,X2) | ~perp(X4,X1,X4,X3) | para(X5,X4,X6,X3) | ~midp(X0,X1,X2)) )),
  inference(resolution,[],[f468,f423])).
fof(f423,plain,(
  ( ! [X4,X2,X0,X3,X1] : (~midp(X4,X2,X3) | para(X0,X2,X1,X3) | ~midp(X4,X0,X1)) )),
  inference(cnf_transformation,[],[f263])).
fof(f263,plain,(
  ! [X0,X1,X2,X3,X4] : (para(X0,X2,X1,X3) | ~midp(X4,X2,X3) | ~midp(X4,X0,X1))),
  inference(flattening,[],[f262])).
fof(f262,plain,(
  ! [X0,X1,X2,X3,X4] : (para(X0,X2,X1,X3) | (~midp(X4,X2,X3) | ~midp(X4,X0,X1)))),
  inference(ennf_transformation,[],[f142])).
fof(f142,plain,(
  ! [X0,X1,X2,X3,X4] : ((midp(X4,X2,X3) & midp(X4,X0,X1)) => para(X0,X2,X1,X3))),
  inference(rectify,[],[f64])).
fof(f64,axiom,(
  ! [X0,X1,X2,X3,X6] : ((midp(X6,X2,X3) & midp(X6,X0,X1)) => para(X0,X2,X1,X3))),
  file('./6_GDD_FULL/more/E023-17.p',ruleD63)).
fof(f468,plain,(
  ( ! [X4,X2,X0,X3,X1] : (midp(sK17(X0,X1,X2,X3,X4),X0,X2) | ~midp(X4,X1,X3) | ~perp(X2,X0,X2,X3) | ~perp(X0,X1,X0,X2)) )),
  inference(cnf_transformation,[],[f353])).
fof(f353,plain,(
  ! [X0,X1,X2,X3,X4] : (midp(sK17(X0,X1,X2,X3,X4),X0,X2) | ~midp(X4,X1,X3) | ~perp(X2,X0,X2,X3) | ~perp(X0,X1,X0,X2))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK17])],[f315,f352])).
fof(f352,plain,(
  ! [X4,X3,X2,X1,X0] : (? [X5] : (midp(X5,X0,X2) | ~midp(X4,X1,X3) | ~perp(X2,X0,X2,X3) | ~perp(X0,X1,X0,X2)) => (midp(sK17(X0,X1,X2,X3,X4),X0,X2) | ~midp(X4,X1,X3) | ~perp(X2,X0,X2,X3) | ~perp(X0,X1,X0,X2)))),
  introduced(choice_axiom,[])).
fof(f315,plain,(
  ! [X0,X1,X2,X3,X4] : ? [X5] : (midp(X5,X0,X2) | ~midp(X4,X1,X3) | ~perp(X2,X0,X2,X3) | ~perp(X0,X1,X0,X2))),
  inference(flattening,[],[f314])).
fof(f314,plain,(
  ! [X0,X1,X2,X3,X4] : ? [X5] : (midp(X5,X0,X2) | (~midp(X4,X1,X3) | ~perp(X2,X0,X2,X3) | ~perp(X0,X1,X0,X2)))),
  inference(ennf_transformation,[],[f164])).
fof(f164,plain,(
  ! [X0,X1,X2,X3,X4] : ? [X5] : ((midp(X4,X1,X3) & perp(X2,X0,X2,X3) & perp(X0,X1,X0,X2)) => midp(X5,X0,X2))),
  inference(rectify,[],[f92])).
fof(f92,axiom,(
  ! [X0,X1,X2,X3,X6] : ? [X8] : ((midp(X6,X1,X3) & perp(X2,X0,X2,X3) & perp(X0,X1,X0,X2)) => midp(X8,X0,X2))),
  file('./6_GDD_FULL/more/E023-17.p',ruleX16)).
fof(f14492,plain,(
  spl33_274 | ~spl33_76),
  inference(avatar_split_clause,[],[f14467,f1101,f2975])).
fof(f14467,plain,(
  ( ! [X64,X62,X63] : (~midp(X62,X63,X64) | cong(X63,X64,X62,X63)) ) | ~spl33_76),
  inference(resolution,[],[f1102,f383])).
fof(f383,plain,(
  ( ! [X2,X0,X3,X1] : (~cong(X0,X1,X2,X3) | cong(X2,X3,X0,X1)) )),
  inference(cnf_transformation,[],[f200])).
fof(f200,plain,(
  ! [X0,X1,X2,X3] : (cong(X2,X3,X0,X1) | ~cong(X0,X1,X2,X3))),
  inference(ennf_transformation,[],[f24])).
fof(f24,axiom,(
  ! [X0,X1,X2,X3] : (cong(X0,X1,X2,X3) => cong(X2,X3,X0,X1))),
  file('./6_GDD_FULL/more/E023-17.p',ruleD24)).
fof(f14442,plain,(
  ~spl33_78 | ~spl33_204),
  inference(avatar_contradiction_clause,[],[f14441])).
fof(f14441,plain,(
  $false | (~spl33_78 | ~spl33_204)),
  inference(subsumption_resolution,[],[f14399,f8891])).
fof(f8891,plain,(
  ( ! [X14] : (midp(X14,X14,X14)) ) | ~spl33_204),
  inference(subsumption_resolution,[],[f8870,f3575])).
fof(f8870,plain,(
  ( ! [X14] : (~coll(X14,X14,X14) | midp(X14,X14,X14)) ) | ~spl33_204),
  inference(resolution,[],[f8784,f427])).
fof(f8784,plain,(
  ( ! [X35] : (cong(X35,X35,X35,X35)) ) | ~spl33_204),
  inference(subsumption_resolution,[],[f8780,f8737])).
fof(f8780,plain,(
  ( ! [X35,X36] : (cong(X35,X35,X35,X35) | ~cyclic(X35,X35,X36,X35)) ) | ~spl33_204),
  inference(duplicate_literal_removal,[],[f8767])).
fof(f8767,plain,(
  ( ! [X35,X36] : (cong(X35,X35,X35,X35) | ~cyclic(X35,X35,X36,X35) | ~cyclic(X35,X35,X36,X35)) ) | ~spl33_204),
  inference(resolution,[],[f8737,f1522])).
fof(f14399,plain,(
  ( ! [X0] : (~midp(X0,X0,X0)) ) | (~spl33_78 | ~spl33_204)),
  inference(resolution,[],[f1105,f8818])).
fof(f1105,plain,(
  ( ! [X6,X8,X7] : (~cong(X6,X7,X7,X8) | ~midp(X6,X7,X8)) ) | ~spl33_78),
  inference(avatar_component_clause,[],[f1104])).
fof(f1104,plain,(
  spl33_78 <=> ! [X7,X8,X6] : (~cong(X6,X7,X7,X8) | ~midp(X6,X7,X8))),
  introduced(avatar_definition,[new_symbols(naming,[spl33_78])])).
fof(f14378,plain,(
  ~spl33_76 | ~spl33_98 | ~spl33_204 | spl33_319),
  inference(avatar_contradiction_clause,[],[f14377])).
fof(f14377,plain,(
  $false | (~spl33_76 | ~spl33_98 | ~spl33_204 | ~spl33_319)),
  inference(subsumption_resolution,[],[f14376,f9043])).
fof(f14376,plain,(
  ~midp(sK25,sK24,sK24) | (~spl33_76 | ~spl33_98 | ~spl33_319)),
  inference(subsumption_resolution,[],[f14339,f3448])).
fof(f3448,plain,(
  ~perp(sK25,sK24,sK20,sK24) | ~spl33_319),
  inference(avatar_component_clause,[],[f3447])).
fof(f3447,plain,(
  spl33_319 <=> ~perp(sK25,sK24,sK20,sK24)),
  introduced(avatar_definition,[new_symbols(naming,[spl33_319])])).
fof(f14339,plain,(
  perp(sK25,sK24,sK20,sK24) | ~midp(sK25,sK24,sK24) | (~spl33_76 | ~spl33_98)),
  inference(resolution,[],[f8026,f6101])).
fof(f6101,plain,(
  cong(sK25,sK20,sK24,sK20) | ~spl33_98),
  inference(resolution,[],[f4882,f382])).
fof(f382,plain,(
  ( ! [X2,X0,X3,X1] : (~cong(X0,X1,X2,X3) | cong(X0,X1,X3,X2)) )),
  inference(cnf_transformation,[],[f199])).
fof(f199,plain,(
  ! [X0,X1,X2,X3] : (cong(X0,X1,X3,X2) | ~cong(X0,X1,X2,X3))),
  inference(ennf_transformation,[],[f23])).
fof(f23,axiom,(
  ! [X0,X1,X2,X3] : (cong(X0,X1,X2,X3) => cong(X0,X1,X3,X2))),
  file('./6_GDD_FULL/more/E023-17.p',ruleD23)).
fof(f4882,plain,(
  cong(sK25,sK20,sK20,sK24) | ~spl33_98),
  inference(resolution,[],[f3843,f383])).
fof(f3843,plain,(
  cong(sK20,sK24,sK25,sK20) | ~spl33_98),
  inference(resolution,[],[f1236,f382])).
fof(f1236,plain,(
  cong(sK20,sK24,sK20,sK25) | ~spl33_98),
  inference(avatar_component_clause,[],[f1235])).
fof(f1235,plain,(
  spl33_98 <=> cong(sK20,sK24,sK20,sK25)),
  introduced(avatar_definition,[new_symbols(naming,[spl33_98])])).
fof(f8026,plain,(
  ( ! [X17,X15,X16] : (~cong(X15,X17,X16,X17) | perp(X15,X16,X17,X16) | ~midp(X15,X16,X16)) ) | ~spl33_76),
  inference(resolution,[],[f1102,f416])).
fof(f10255,plain,(
  ~spl33_204 | ~spl33_208),
  inference(avatar_contradiction_clause,[],[f10235])).
fof(f10235,plain,(
  $false | (~spl33_204 | ~spl33_208)),
  inference(resolution,[],[f9751,f2411])).
fof(f2411,plain,(
  ( ! [X48,X49] : (~perp(X48,X49,X48,X49)) ) | ~spl33_208),
  inference(avatar_component_clause,[],[f2410])).
fof(f2410,plain,(
  spl33_208 <=> ! [X49,X48] : ~perp(X48,X49,X48,X49)),
  introduced(avatar_definition,[new_symbols(naming,[spl33_208])])).
fof(f9751,plain,(
  ( ! [X56,X55] : (perp(X55,X56,X56,X56)) ) | ~spl33_204),
  inference(subsumption_resolution,[],[f9750,f8784])).
fof(f9750,plain,(
  ( ! [X56,X55] : (perp(X55,X56,X56,X56) | ~cong(X56,X56,X56,X56)) ) | ~spl33_204),
  inference(subsumption_resolution,[],[f9746,f9043])).
fof(f9746,plain,(
  ( ! [X56,X55] : (~midp(X55,X56,X56) | perp(X55,X56,X56,X56) | ~cong(X56,X56,X56,X56)) )),
  inference(duplicate_literal_removal,[],[f9740])).
fof(f9740,plain,(
  ( ! [X56,X55] : (~midp(X55,X56,X56) | ~midp(X55,X56,X56) | perp(X55,X56,X56,X56) | ~cong(X56,X56,X56,X56)) )),
  inference(resolution,[],[f2112,f1027])).
fof(f1027,plain,(
  ( ! [X12,X13] : (~circle(X12,X13,X13,X13) | perp(X12,X13,X13,X13) | ~cong(X13,X13,X13,X13)) )),
  inference(resolution,[],[f409,f406])).
fof(f406,plain,(
  ( ! [X2,X0,X1] : (eqangle(X2,X0,X0,X1,X0,X1,X2,X1) | ~cong(X2,X0,X2,X1)) )),
  inference(cnf_transformation,[],[f232])).
fof(f232,plain,(
  ! [X0,X1,X2] : (eqangle(X2,X0,X0,X1,X0,X1,X2,X1) | ~cong(X2,X0,X2,X1))),
  inference(ennf_transformation,[],[f126])).
fof(f126,plain,(
  ! [X0,X1,X2] : (cong(X2,X0,X2,X1) => eqangle(X2,X0,X0,X1,X0,X1,X2,X1))),
  inference(rectify,[],[f47])).
fof(f47,axiom,(
  ! [X0,X1,X7] : (cong(X7,X0,X7,X1) => eqangle(X7,X0,X0,X1,X0,X1,X7,X1))),
  file('./6_GDD_FULL/more/E023-17.p',ruleD46)).
fof(f409,plain,(
  ( ! [X4,X2,X0,X3,X1] : (~eqangle(X0,X4,X0,X1,X2,X0,X2,X1) | perp(X3,X0,X0,X4) | ~circle(X3,X0,X1,X2)) )),
  inference(cnf_transformation,[],[f238])).
fof(f238,plain,(
  ! [X0,X1,X2,X3,X4] : (perp(X3,X0,X0,X4) | ~eqangle(X0,X4,X0,X1,X2,X0,X2,X1) | ~circle(X3,X0,X1,X2))),
  inference(flattening,[],[f237])).
fof(f237,plain,(
  ! [X0,X1,X2,X3,X4] : (perp(X3,X0,X0,X4) | (~eqangle(X0,X4,X0,X1,X2,X0,X2,X1) | ~circle(X3,X0,X1,X2)))),
  inference(ennf_transformation,[],[f129])).
fof(f129,plain,(
  ! [X0,X1,X2,X3,X4] : ((eqangle(X0,X4,X0,X1,X2,X0,X2,X1) & circle(X3,X0,X1,X2)) => perp(X3,X0,X0,X4))),
  inference(rectify,[],[f50])).
fof(f50,axiom,(
  ! [X0,X1,X2,X7,X15] : ((eqangle(X0,X15,X0,X1,X2,X0,X2,X1) & circle(X7,X0,X1,X2)) => perp(X7,X0,X0,X15))),
  file('./6_GDD_FULL/more/E023-17.p',ruleD49)).
fof(f2112,plain,(
  ( ! [X2,X0,X3,X1] : (circle(X0,X1,X2,X3) | ~midp(X0,X1,X3) | ~midp(X0,X2,X1)) )),
  inference(resolution,[],[f738,f541])).
fof(f541,plain,(
  ( ! [X2,X0,X1] : (cong(X0,X1,X0,X2) | ~midp(X0,X2,X1)) )),
  inference(resolution,[],[f383,f428])).
fof(f428,plain,(
  ( ! [X2,X0,X1] : (cong(X0,X1,X0,X2) | ~midp(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f271])).
fof(f271,plain,(
  ! [X0,X1,X2] : (cong(X0,X1,X0,X2) | ~midp(X0,X1,X2))),
  inference(ennf_transformation,[],[f69])).
fof(f69,axiom,(
  ! [X0,X1,X2] : (midp(X0,X1,X2) => cong(X0,X1,X0,X2))),
  file('./6_GDD_FULL/more/E023-17.p',ruleD68)).
fof(f738,plain,(
  ( ! [X2,X0,X3,X1] : (~cong(X0,X1,X0,X2) | circle(X0,X1,X2,X3) | ~midp(X0,X1,X3)) )),
  inference(resolution,[],[f371,f428])).
fof(f371,plain,(
  ( ! [X2,X0,X3,X1] : (~cong(X3,X0,X3,X2) | circle(X3,X0,X1,X2) | ~cong(X3,X0,X3,X1)) )),
  inference(cnf_transformation,[],[f185])).
fof(f185,plain,(
  ! [X0,X1,X2,X3] : (circle(X3,X0,X1,X2) | ~cong(X3,X0,X3,X2) | ~cong(X3,X0,X3,X1))),
  inference(flattening,[],[f184])).
fof(f184,plain,(
  ! [X0,X1,X2,X3] : (circle(X3,X0,X1,X2) | (~cong(X3,X0,X3,X2) | ~cong(X3,X0,X3,X1)))),
  inference(ennf_transformation,[],[f98])).
fof(f98,plain,(
  ! [X0,X1,X2,X3] : ((cong(X3,X0,X3,X2) & cong(X3,X0,X3,X1)) => circle(X3,X0,X1,X2))),
  inference(rectify,[],[f12])).
fof(f12,axiom,(
  ! [X0,X1,X2,X7] : ((cong(X7,X0,X7,X2) & cong(X7,X0,X7,X1)) => circle(X7,X0,X1,X2))),
  file('./6_GDD_FULL/more/E023-17.p',ruleD12)).
fof(f5906,plain,(
  ~spl33_272),
  inference(avatar_contradiction_clause,[],[f5883])).
fof(f5883,plain,(
  $false | ~spl33_272),
  inference(resolution,[],[f2948,f477])).
fof(f477,plain,(
  midp(sK22,sK23,sK25)),
  inference(cnf_transformation,[],[f359])).
fof(f2948,plain,(
  ( ! [X12,X10,X11] : (~midp(X10,X11,X12)) ) | ~spl33_272),
  inference(avatar_component_clause,[],[f2947])).
fof(f3455,plain,(
  ~spl33_319 | spl33_320 | ~spl33_66),
  inference(avatar_split_clause,[],[f3428,f1021,f3453,f3447])).
fof(f1021,plain,(
  spl33_66 <=> cong(sK20,sK25,sK20,sK24)),
  introduced(avatar_definition,[new_symbols(naming,[spl33_66])])).
fof(f3428,plain,(
  perp(sK20,sK25,sK25,sK24) | ~perp(sK25,sK24,sK20,sK24) | ~spl33_66),
  inference(resolution,[],[f1022,f1090])).
fof(f1090,plain,(
  ( ! [X21,X19,X20] : (~cong(X21,X19,X21,X20) | perp(X21,X19,X19,X20) | ~perp(X19,X20,X21,X20)) )),
  inference(resolution,[],[f434,f406])).
fof(f434,plain,(
  ( ! [X6,X4,X2,X0,X7,X5,X3,X1] : (~eqangle(X0,X1,X2,X3,X4,X5,X6,X7) | ~perp(X4,X5,X6,X7) | perp(X0,X1,X2,X3)) )),
  inference(cnf_transformation,[],[f282])).
fof(f282,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6,X7] : (perp(X0,X1,X2,X3) | ~perp(X4,X5,X6,X7) | ~eqangle(X0,X1,X2,X3,X4,X5,X6,X7))),
  inference(flattening,[],[f281])).
fof(f281,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6,X7] : (perp(X0,X1,X2,X3) | (~perp(X4,X5,X6,X7) | ~eqangle(X0,X1,X2,X3,X4,X5,X6,X7)))),
  inference(ennf_transformation,[],[f147])).
fof(f147,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6,X7] : ((perp(X4,X5,X6,X7) & eqangle(X0,X1,X2,X3,X4,X5,X6,X7)) => perp(X0,X1,X2,X3))),
  inference(rectify,[],[f75])).
fof(f75,axiom,(
  ! [X0,X1,X2,X3,X8,X9,X10,X11] : ((perp(X8,X9,X10,X11) & eqangle(X0,X1,X2,X3,X8,X9,X10,X11)) => perp(X0,X1,X2,X3))),
  file('./6_GDD_FULL/more/E023-17.p',ruleD74)).
fof(f1022,plain,(
  cong(sK20,sK25,sK20,sK24) | ~spl33_66),
  inference(avatar_component_clause,[],[f1021])).
fof(f2958,plain,(
  spl33_96 | ~spl33_64),
  inference(avatar_split_clause,[],[f2954,f1012,f1226])).
fof(f1226,plain,(
  spl33_96 <=> midp(sK28,sK24,sK25)),
  introduced(avatar_definition,[new_symbols(naming,[spl33_96])])).
fof(f1012,plain,(
  spl33_64 <=> midp(sK28,sK25,sK24)),
  introduced(avatar_definition,[new_symbols(naming,[spl33_64])])).
fof(f2954,plain,(
  midp(sK28,sK24,sK25) | ~spl33_64),
  inference(resolution,[],[f1013,f370])).
fof(f370,plain,(
  ( ! [X2,X0,X1] : (~midp(X2,X1,X0) | midp(X2,X0,X1)) )),
  inference(cnf_transformation,[],[f183])).
fof(f183,plain,(
  ! [X0,X1,X2] : (midp(X2,X0,X1) | ~midp(X2,X1,X0))),
  inference(ennf_transformation,[],[f97])).
fof(f97,plain,(
  ! [X0,X1,X2] : (midp(X2,X1,X0) => midp(X2,X0,X1))),
  inference(rectify,[],[f11])).
fof(f11,axiom,(
  ! [X0,X1,X6] : (midp(X6,X1,X0) => midp(X6,X0,X1))),
  file('./6_GDD_FULL/more/E023-17.p',ruleD11)).
fof(f1013,plain,(
  midp(sK28,sK25,sK24) | ~spl33_64),
  inference(avatar_component_clause,[],[f1012])).
fof(f2850,plain,(
  ~spl33_206),
  inference(avatar_contradiction_clause,[],[f2841])).
fof(f2841,plain,(
  $false | ~spl33_206),
  inference(resolution,[],[f2254,f1759])).
fof(f1759,plain,(
  para(sK20,sK28,sK20,sK28)),
  inference(resolution,[],[f730,f529])).
fof(f2254,plain,(
  ( ! [X0,X1] : (~para(X0,X1,X0,X1)) ) | ~spl33_206),
  inference(avatar_component_clause,[],[f2253])).
fof(f2253,plain,(
  spl33_206 <=> ! [X1,X0] : ~para(X0,X1,X0,X1)),
  introduced(avatar_definition,[new_symbols(naming,[spl33_206])])).
fof(f2831,plain,(
  spl33_65 | ~spl33_204),
  inference(avatar_contradiction_clause,[],[f2830])).
fof(f2830,plain,(
  $false | (~spl33_65 | ~spl33_204)),
  inference(resolution,[],[f2811,f1388])).
fof(f1388,plain,(
  ~coll(sK28,sK25,sK24) | ~spl33_65),
  inference(subsumption_resolution,[],[f1383,f1016])).
fof(f1016,plain,(
  ~midp(sK28,sK25,sK24) | ~spl33_65),
  inference(avatar_component_clause,[],[f1015])).
fof(f1015,plain,(
  spl33_65 <=> ~midp(sK28,sK25,sK24)),
  introduced(avatar_definition,[new_symbols(naming,[spl33_65])])).
fof(f1383,plain,(
  ~coll(sK28,sK25,sK24) | midp(sK28,sK25,sK24)),
  inference(resolution,[],[f1123,f427])).
fof(f1123,plain,(
  cong(sK28,sK25,sK28,sK24)),
  inference(subsumption_resolution,[],[f1120,f481])).
fof(f481,plain,(
  midp(sK20,sK25,sK24)),
  inference(cnf_transformation,[],[f359])).
fof(f1120,plain,(
  cong(sK28,sK25,sK28,sK24) | ~midp(sK20,sK25,sK24)),
  inference(resolution,[],[f887,f415])).
fof(f415,plain,(
  ( ! [X2,X0,X3,X1] : (~perp(X3,X2,X0,X1) | cong(X3,X0,X3,X1) | ~midp(X2,X0,X1)) )),
  inference(cnf_transformation,[],[f250])).
fof(f250,plain,(
  ! [X0,X1,X2,X3] : (cong(X3,X0,X3,X1) | ~perp(X3,X2,X0,X1) | ~midp(X2,X0,X1))),
  inference(flattening,[],[f249])).
fof(f249,plain,(
  ! [X0,X1,X2,X3] : (cong(X3,X0,X3,X1) | (~perp(X3,X2,X0,X1) | ~midp(X2,X0,X1)))),
  inference(ennf_transformation,[],[f134])).
fof(f134,plain,(
  ! [X0,X1,X2,X3] : ((perp(X3,X2,X0,X1) & midp(X2,X0,X1)) => cong(X3,X0,X3,X1))),
  inference(rectify,[],[f56])).
fof(f56,axiom,(
  ! [X0,X1,X6,X7] : ((perp(X7,X6,X0,X1) & midp(X6,X0,X1)) => cong(X7,X0,X7,X1))),
  file('./6_GDD_FULL/more/E023-17.p',ruleD55)).
fof(f887,plain,(
  perp(sK28,sK20,sK25,sK24)),
  inference(resolution,[],[f522,f367])).
fof(f522,plain,(
  perp(sK25,sK24,sK28,sK20)),
  inference(resolution,[],[f366,f482])).
fof(f2811,plain,(
  ( ! [X2,X0,X1] : (coll(X0,X1,X2)) ) | ~spl33_204),
  inference(subsumption_resolution,[],[f2808,f2803])).
fof(f2803,plain,(
  ( ! [X6,X5] : (coll(X5,X5,X6)) ) | ~spl33_204),
  inference(resolution,[],[f2795,f360])).
fof(f360,plain,(
  ( ! [X2,X0,X1] : (~coll(X0,X1,X2) | coll(X0,X2,X1)) )),
  inference(cnf_transformation,[],[f169])).
fof(f169,plain,(
  ! [X0,X1,X2] : (coll(X0,X2,X1) | ~coll(X0,X1,X2))),
  inference(ennf_transformation,[],[f1])).
fof(f1,axiom,(
  ! [X0,X1,X2] : (coll(X0,X1,X2) => coll(X0,X2,X1))),
  file('./6_GDD_FULL/more/E023-17.p',ruleD1)).
fof(f2795,plain,(
  ( ! [X4,X3] : (coll(X3,X4,X3)) ) | ~spl33_204),
  inference(resolution,[],[f2780,f361])).
fof(f361,plain,(
  ( ! [X2,X0,X1] : (~coll(X0,X1,X2) | coll(X1,X0,X2)) )),
  inference(cnf_transformation,[],[f170])).
fof(f170,plain,(
  ! [X0,X1,X2] : (coll(X1,X0,X2) | ~coll(X0,X1,X2))),
  inference(ennf_transformation,[],[f2])).
fof(f2,axiom,(
  ! [X0,X1,X2] : (coll(X0,X1,X2) => coll(X1,X0,X2))),
  file('./6_GDD_FULL/more/E023-17.p',ruleD2)).
fof(f2780,plain,(
  ( ! [X0,X1] : (coll(X0,X1,X1)) ) | ~spl33_204),
  inference(resolution,[],[f2251,f426])).
fof(f2808,plain,(
  ( ! [X2,X0,X1] : (coll(X0,X1,X2) | ~coll(X2,X2,X0)) ) | ~spl33_204),
  inference(resolution,[],[f2803,f362])).
fof(f362,plain,(
  ( ! [X2,X0,X3,X1] : (~coll(X0,X1,X3) | coll(X2,X3,X0) | ~coll(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f172])).
fof(f172,plain,(
  ! [X0,X1,X2,X3] : (coll(X2,X3,X0) | ~coll(X0,X1,X3) | ~coll(X0,X1,X2))),
  inference(flattening,[],[f171])).
fof(f171,plain,(
  ! [X0,X1,X2,X3] : (coll(X2,X3,X0) | (~coll(X0,X1,X3) | ~coll(X0,X1,X2)))),
  inference(ennf_transformation,[],[f3])).
fof(f3,axiom,(
  ! [X0,X1,X2,X3] : ((coll(X0,X1,X3) & coll(X0,X1,X2)) => coll(X2,X3,X0))),
  file('./6_GDD_FULL/more/E023-17.p',ruleD3)).
fof(f2415,plain,(
  spl33_208 | spl33_210),
  inference(avatar_split_clause,[],[f2392,f2413,f2410])).
fof(f2392,plain,(
  ( ! [X47,X45,X48,X46,X44,X49] : (~para(X44,X45,X46,X47) | ~perp(X48,X49,X48,X49) | perp(X44,X45,X46,X47)) )),
  inference(resolution,[],[f805,f434])).
fof(f2255,plain,(
  spl33_204 | spl33_206),
  inference(avatar_split_clause,[],[f2226,f2253,f2250])).
fof(f2226,plain,(
  ( ! [X2,X0,X3,X1] : (~para(X0,X1,X0,X1) | para(X2,X3,X2,X3)) )),
  inference(resolution,[],[f776,f398])).
fof(f776,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (eqangle(X0,X1,X2,X3,X0,X1,X4,X5) | ~para(X2,X3,X4,X5)) )),
  inference(resolution,[],[f378,f399])).
fof(f378,plain,(
  ( ! [X6,X4,X2,X0,X7,X5,X3,X1] : (~eqangle(X0,X1,X2,X3,X4,X5,X6,X7) | eqangle(X2,X3,X0,X1,X6,X7,X4,X5)) )),
  inference(cnf_transformation,[],[f194])).
fof(f194,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6,X7] : (eqangle(X2,X3,X0,X1,X6,X7,X4,X5) | ~eqangle(X0,X1,X2,X3,X4,X5,X6,X7))),
  inference(ennf_transformation,[],[f101])).
fof(f101,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6,X7] : (eqangle(X0,X1,X2,X3,X4,X5,X6,X7) => eqangle(X2,X3,X0,X1,X6,X7,X4,X5))),
  inference(rectify,[],[f19])).
fof(f19,axiom,(
  ! [X0,X1,X2,X3,X8,X9,X10,X11] : (eqangle(X0,X1,X2,X3,X8,X9,X10,X11) => eqangle(X2,X3,X0,X1,X10,X11,X8,X9))),
  file('./6_GDD_FULL/more/E023-17.p',ruleD19)).
fof(f1237,plain,(
  ~spl33_97 | spl33_98),
  inference(avatar_split_clause,[],[f1222,f1235,f1229])).
fof(f1229,plain,(
  spl33_97 <=> ~midp(sK28,sK24,sK25)),
  introduced(avatar_definition,[new_symbols(naming,[spl33_97])])).
fof(f1222,plain,(
  cong(sK20,sK24,sK20,sK25) | ~midp(sK28,sK24,sK25)),
  inference(resolution,[],[f1010,f415])).
fof(f1010,plain,(
  perp(sK20,sK28,sK24,sK25)),
  inference(resolution,[],[f529,f366])).
fof(f1106,plain,(
  spl33_76 | spl33_78),
  inference(avatar_split_clause,[],[f1099,f1104,f1101])).
fof(f1099,plain,(
  ( ! [X6,X10,X8,X7,X11,X9] : (~cong(X6,X7,X7,X8) | cong(X9,X10,X10,X11) | ~midp(X6,X7,X8) | ~midp(X9,X10,X11)) )),
  inference(resolution,[],[f435,f430])).
fof(f430,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (eqratio(X4,X0,X0,X1,X5,X2,X2,X3) | ~midp(X5,X2,X3) | ~midp(X4,X0,X1)) )),
  inference(cnf_transformation,[],[f274])).
fof(f274,plain,(
  ! [X0,X1,X2,X3,X4,X5] : (eqratio(X4,X0,X0,X1,X5,X2,X2,X3) | ~midp(X5,X2,X3) | ~midp(X4,X0,X1))),
  inference(flattening,[],[f273])).
fof(f273,plain,(
  ! [X0,X1,X2,X3,X4,X5] : (eqratio(X4,X0,X0,X1,X5,X2,X2,X3) | (~midp(X5,X2,X3) | ~midp(X4,X0,X1)))),
  inference(ennf_transformation,[],[f145])).
fof(f145,plain,(
  ! [X0,X1,X2,X3,X4,X5] : ((midp(X5,X2,X3) & midp(X4,X0,X1)) => eqratio(X4,X0,X0,X1,X5,X2,X2,X3))),
  inference(rectify,[],[f71])).
fof(f71,axiom,(
  ! [X0,X1,X2,X3,X6,X16] : ((midp(X16,X2,X3) & midp(X6,X0,X1)) => eqratio(X6,X0,X0,X1,X16,X2,X2,X3))),
  file('./6_GDD_FULL/more/E023-17.p',ruleD70)).
fof(f435,plain,(
  ( ! [X6,X4,X2,X0,X7,X5,X3,X1] : (~eqratio(X0,X1,X2,X3,X4,X5,X6,X7) | ~cong(X4,X5,X6,X7) | cong(X0,X1,X2,X3)) )),
  inference(cnf_transformation,[],[f284])).
fof(f284,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6,X7] : (cong(X0,X1,X2,X3) | ~cong(X4,X5,X6,X7) | ~eqratio(X0,X1,X2,X3,X4,X5,X6,X7))),
  inference(flattening,[],[f283])).
fof(f283,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6,X7] : (cong(X0,X1,X2,X3) | (~cong(X4,X5,X6,X7) | ~eqratio(X0,X1,X2,X3,X4,X5,X6,X7)))),
  inference(ennf_transformation,[],[f148])).
fof(f148,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6,X7] : ((cong(X4,X5,X6,X7) & eqratio(X0,X1,X2,X3,X4,X5,X6,X7)) => cong(X0,X1,X2,X3))),
  inference(rectify,[],[f76])).
fof(f76,axiom,(
  ! [X0,X1,X2,X3,X8,X9,X10,X11] : ((cong(X8,X9,X10,X11) & eqratio(X0,X1,X2,X3,X8,X9,X10,X11)) => cong(X0,X1,X2,X3))),
  file('./6_GDD_FULL/more/E023-17.p',ruleD75)).
fof(f1023,plain,(
  ~spl33_65 | spl33_66),
  inference(avatar_split_clause,[],[f1008,f1021,f1015])).
fof(f1008,plain,(
  cong(sK20,sK25,sK20,sK24) | ~midp(sK28,sK25,sK24)),
  inference(resolution,[],[f529,f415])).
% SZS output end Proof for E023-17
% ------------------------------
% Version: Vampire 4.4.0 (commit 7916d27 on 2019-08-23 08:50:16 +0100)
% Termination reason: Refutation

% Memory used [KB]: 18549
% Time elapsed: 0.487 s
% ------------------------------
% ------------------------------
