% Refutation found. Thanks to Tanya!
% SZS status Theorem for B016x
% SZS output start Proof for B016x
23. ! [X0,X1,X2,X3] : (cong(X0,X1,X2,X3) => cong(X0,X1,X3,X2)) [input]
25. ! [X0,X1,X2,X3,X4,X5] : ((cong(X2,X3,X4,X5) & cong(X0,X1,X2,X3)) => cong(X0,X1,X4,X5)) [input]
68. ! [X0,X1,X2] : (midp(X0,X1,X2) => cong(X0,X1,X0,X2)) [input]
94. ! [X0,X1,X2,X3,X4,X5,X12] : (ratio(X12,X4,X0,X3) = 1/1 & midp(X5,X1,X0) & midp(X4,X2,X0) & midp(X3,X1,X2)) => cong(X2,X5,X12,X1) [input]
95. ~! [X1,X12,X5,X2] : (! [X0,X1,X2,X3,X4,X5,X12] : (ratio(X12,X4,X0,X3) = 1/1 & midp(X5,X1,X0) & midp(X4,X2,X0) & midp(X3,X1,X2)) => cong(X2,X5,X12,X1)) [negated conjecture 94]
165. ~! [X0,X1,X2,X3] : (! [X4,X5,X6,X7,X8,X9,X10] : (ratio(X10,X8,X4,X7) = 1/1 & midp(X9,X5,X4) & midp(X8,X6,X4) & midp(X7,X5,X6)) => cong(X3,X2,X1,X0)) [rectify 95]
196. ! [X0,X1,X2,X3] : (cong(X0,X1,X3,X2) | ~cong(X0,X1,X2,X3)) [ennf transformation 23]
198. ! [X0,X1,X2,X3,X4,X5] : (cong(X0,X1,X4,X5) | (~cong(X2,X3,X4,X5) | ~cong(X0,X1,X2,X3))) [ennf transformation 25]
199. ! [X0,X1,X2,X3,X4,X5] : (cong(X0,X1,X4,X5) | ~cong(X2,X3,X4,X5) | ~cong(X0,X1,X2,X3)) [flattening 198]
266. ! [X0,X1,X2] : (cong(X0,X1,X0,X2) | ~midp(X0,X1,X2)) [ennf transformation 68]
315. ? [X0,X1,X2,X3] : (~cong(X3,X2,X1,X0) & ! [X4,X5,X6,X7,X8,X9,X10] : (ratio(X10,X8,X4,X7) = 1/1 & midp(X9,X5,X4) & midp(X8,X6,X4) & midp(X7,X5,X6))) [ennf transformation 165]
352. ? [X0,X1,X2,X3] : (~cong(X3,X2,X1,X0) & ! [X4,X5,X6,X7,X8,X9,X10] : (ratio(X10,X8,X4,X7) = 1/1 & midp(X9,X5,X4) & midp(X8,X6,X4) & midp(X7,X5,X6))) => (~cong(sK23,sK22,sK21,sK20) & ! [X4,X5,X6,X7,X8,X9,X10] : (ratio(X10,X8,X4,X7) = 1/1 & midp(X9,X5,X4) & midp(X8,X6,X4) & midp(X7,X5,X6))) [choice axiom]
353. ~cong(sK23,sK22,sK21,sK20) & ! [X4,X5,X6,X7,X8,X9,X10] : (ratio(X10,X8,X4,X7) = 1/1 & midp(X9,X5,X4) & midp(X8,X6,X4) & midp(X7,X5,X6)) [skolemisation 315,352]
376. ~cong(X0,X1,X2,X3) | cong(X0,X1,X3,X2) [cnf transformation 196]
378. ~cong(X2,X3,X4,X5) | cong(X0,X1,X4,X5) | ~cong(X0,X1,X2,X3) [cnf transformation 199]
422. cong(X0,X1,X0,X2) | ~midp(X0,X1,X2) [cnf transformation 266]
469. midp(X9,X5,X4) [cnf transformation 353]
471. ~cong(sK23,sK22,sK21,sK20) [cnf transformation 353]
476. cong(X0,X1,X2,X0) | ~midp(X0,X1,X2) [resolution 376,422]
477. cong(X0,X1,X2,X0) [subsumption resolution 476,469]
539. ~cong(X10,X11,X13,X14) | cong(X10,X11,X12,X13) [resolution 378,477]
550. cong(X7,X8,X9,X10) [resolution 539,477]
553. $false [resolution 550,471]
% SZS output end Proof for B016x
% ------------------------------
% Version: Vampire 4.4.0 (commit 7916d27 on 2019-08-23 08:50:16 +0100)
% Termination reason: Refutation

% Memory used [KB]: 5117
% Time elapsed: 0.027 s
% ------------------------------
% ------------------------------
