% Refutation found. Thanks to Tanya!
% SZS status Theorem for L022-1
% SZS output start Proof for L022-1
fof(f12068,plain,(
  $false),
  inference(avatar_sat_refutation,[],[f1725,f2310,f2545,f2779,f3504,f3506,f3564,f3834,f4416,f7289,f8658,f12058,f12067])).
fof(f12067,plain,(
  ~spl26_88 | ~spl26_142 | spl26_295),
  inference(avatar_contradiction_clause,[],[f12066])).
fof(f12066,plain,(
  $false | (~spl26_88 | ~spl26_142 | ~spl26_295)),
  inference(subsumption_resolution,[],[f7288,f8561])).
fof(f8561,plain,(
  ( ! [X0] : (midp(X0,sK20,sK23)) ) | (~spl26_88 | ~spl26_142)),
  inference(resolution,[],[f8503,f370])).
fof(f370,plain,(
  ( ! [X2,X0,X1] : (~midp(X2,X1,X0) | midp(X2,X0,X1)) )),
  inference(cnf_transformation,[],[f183])).
fof(f183,plain,(
  ! [X0,X1,X2] : (midp(X2,X0,X1) | ~midp(X2,X1,X0))),
  inference(ennf_transformation,[],[f97])).
fof(f97,plain,(
  ! [X0,X1,X2] : (midp(X2,X1,X0) => midp(X2,X0,X1))),
  inference(rectify,[],[f11])).
fof(f11,axiom,(
  ! [X0,X1,X6] : (midp(X6,X1,X0) => midp(X6,X0,X1))),
  file('./Other/gao/L/L022-1.p',ruleD11)).
fof(f8503,plain,(
  ( ! [X40] : (midp(X40,sK23,sK20)) ) | (~spl26_88 | ~spl26_142)),
  inference(resolution,[],[f8433,f3503])).
fof(f3503,plain,(
  ( ! [X1] : (~midp(X1,sK20,sK20) | midp(X1,sK23,sK20)) ) | ~spl26_142),
  inference(avatar_component_clause,[],[f3502])).
fof(f3502,plain,(
  spl26_142 <=> ! [X1] : (midp(X1,sK23,sK20) | ~midp(X1,sK20,sK20))),
  introduced(avatar_definition,[new_symbols(naming,[spl26_142])])).
fof(f8433,plain,(
  ( ! [X8,X7] : (midp(X7,X8,X8)) ) | ~spl26_88),
  inference(subsumption_resolution,[],[f8393,f2550])).
fof(f2550,plain,(
  ( ! [X0,X1] : (coll(X0,X1,X1)) ) | ~spl26_88),
  inference(resolution,[],[f1721,f426])).
fof(f426,plain,(
  ( ! [X2,X0,X1] : (~para(X0,X1,X0,X2) | coll(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f268])).
fof(f268,plain,(
  ! [X0,X1,X2] : (coll(X0,X1,X2) | ~para(X0,X1,X0,X2))),
  inference(ennf_transformation,[],[f67])).
fof(f67,axiom,(
  ! [X0,X1,X2] : (para(X0,X1,X0,X2) => coll(X0,X1,X2))),
  file('./Other/gao/L/L022-1.p',ruleD66)).
fof(f1721,plain,(
  ( ! [X2,X3] : (para(X2,X3,X2,X3)) ) | ~spl26_88),
  inference(avatar_component_clause,[],[f1720])).
fof(f1720,plain,(
  spl26_88 <=> ! [X3,X2] : para(X2,X3,X2,X3)),
  introduced(avatar_definition,[new_symbols(naming,[spl26_88])])).
fof(f8393,plain,(
  ( ! [X8,X7] : (~coll(X7,X8,X8) | midp(X7,X8,X8)) ) | ~spl26_88),
  inference(resolution,[],[f8390,f427])).
fof(f427,plain,(
  ( ! [X2,X0,X1] : (~cong(X0,X1,X0,X2) | ~coll(X0,X1,X2) | midp(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f270])).
fof(f270,plain,(
  ! [X0,X1,X2] : (midp(X0,X1,X2) | ~coll(X0,X1,X2) | ~cong(X0,X1,X0,X2))),
  inference(flattening,[],[f269])).
fof(f269,plain,(
  ! [X0,X1,X2] : (midp(X0,X1,X2) | (~coll(X0,X1,X2) | ~cong(X0,X1,X0,X2)))),
  inference(ennf_transformation,[],[f68])).
fof(f68,axiom,(
  ! [X0,X1,X2] : ((coll(X0,X1,X2) & cong(X0,X1,X0,X2)) => midp(X0,X1,X2))),
  file('./Other/gao/L/L022-1.p',ruleD67)).
fof(f8390,plain,(
  ( ! [X83,X84] : (cong(X83,X84,X83,X84)) ) | ~spl26_88),
  inference(subsumption_resolution,[],[f8389,f7619])).
fof(f7619,plain,(
  ( ! [X12,X10,X11,X9] : (cyclic(X9,X10,X11,X12)) ) | ~spl26_88),
  inference(subsumption_resolution,[],[f7597,f6834])).
fof(f6834,plain,(
  ( ! [X6,X8,X7] : (cyclic(X6,X6,X7,X8)) ) | ~spl26_88),
  inference(resolution,[],[f2540,f1721])).
fof(f2540,plain,(
  ( ! [X68,X66,X67] : (~para(X66,X67,X66,X67) | cyclic(X67,X67,X66,X68)) )),
  inference(subsumption_resolution,[],[f1882,f1883])).
fof(f1883,plain,(
  ( ! [X70,X71,X69] : (~para(X69,X70,X69,X70) | coll(X69,X71,X70) | cyclic(X70,X70,X69,X71)) )),
  inference(resolution,[],[f789,f401])).
fof(f401,plain,(
  ( ! [X2,X0,X3,X1] : (~eqangle(X2,X0,X2,X1,X3,X0,X3,X1) | coll(X2,X3,X0) | cyclic(X0,X1,X2,X3)) )),
  inference(cnf_transformation,[],[f223])).
fof(f223,plain,(
  ! [X0,X1,X2,X3] : (cyclic(X0,X1,X2,X3) | coll(X2,X3,X0) | ~eqangle(X2,X0,X2,X1,X3,X0,X3,X1))),
  inference(flattening,[],[f222])).
fof(f222,plain,(
  ! [X0,X1,X2,X3] : (cyclic(X0,X1,X2,X3) | (coll(X2,X3,X0) | ~eqangle(X2,X0,X2,X1,X3,X0,X3,X1)))),
  inference(ennf_transformation,[],[f121])).
fof(f121,plain,(
  ! [X0,X1,X2,X3] : ((~coll(X2,X3,X0) & eqangle(X2,X0,X2,X1,X3,X0,X3,X1)) => cyclic(X0,X1,X2,X3))),
  inference(rectify,[],[f42])).
fof(f42,axiom,(
  ! [X0,X1,X8,X9] : ((~coll(X8,X9,X0) & eqangle(X8,X0,X8,X1,X9,X0,X9,X1)) => cyclic(X0,X1,X8,X9))),
  file('./Other/gao/L/L022-1.p',ruleD42a)).
fof(f789,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (eqangle(X0,X1,X2,X3,X4,X5,X4,X5) | ~para(X0,X1,X2,X3)) )),
  inference(resolution,[],[f380,f399])).
fof(f399,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (eqangle(X0,X1,X4,X5,X2,X3,X4,X5) | ~para(X0,X1,X2,X3)) )),
  inference(cnf_transformation,[],[f220])).
fof(f220,plain,(
  ! [X0,X1,X2,X3,X4,X5] : (eqangle(X0,X1,X4,X5,X2,X3,X4,X5) | ~para(X0,X1,X2,X3))),
  inference(ennf_transformation,[],[f119])).
fof(f119,plain,(
  ! [X0,X1,X2,X3,X4,X5] : (para(X0,X1,X2,X3) => eqangle(X0,X1,X4,X5,X2,X3,X4,X5))),
  inference(rectify,[],[f40])).
fof(f40,axiom,(
  ! [X0,X1,X2,X3,X8,X9] : (para(X0,X1,X2,X3) => eqangle(X0,X1,X8,X9,X2,X3,X8,X9))),
  file('./Other/gao/L/L022-1.p',ruleD40)).
fof(f380,plain,(
  ( ! [X6,X4,X2,X0,X7,X5,X3,X1] : (~eqangle(X0,X1,X2,X3,X4,X5,X6,X7) | eqangle(X0,X1,X4,X5,X2,X3,X6,X7)) )),
  inference(cnf_transformation,[],[f196])).
fof(f196,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6,X7] : (eqangle(X0,X1,X4,X5,X2,X3,X6,X7) | ~eqangle(X0,X1,X2,X3,X4,X5,X6,X7))),
  inference(ennf_transformation,[],[f103])).
fof(f103,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6,X7] : (eqangle(X0,X1,X2,X3,X4,X5,X6,X7) => eqangle(X0,X1,X4,X5,X2,X3,X6,X7))),
  inference(rectify,[],[f21])).
fof(f21,axiom,(
  ! [X0,X1,X2,X3,X8,X9,X10,X11] : (eqangle(X0,X1,X2,X3,X8,X9,X10,X11) => eqangle(X0,X1,X8,X9,X2,X3,X10,X11))),
  file('./Other/gao/L/L022-1.p',ruleD21)).
fof(f1882,plain,(
  ( ! [X68,X66,X67] : (~para(X66,X67,X66,X67) | ~coll(X66,X68,X67) | cyclic(X67,X67,X66,X68)) )),
  inference(resolution,[],[f789,f402])).
fof(f402,plain,(
  ( ! [X2,X0,X3,X1] : (~eqangle(X2,X0,X2,X1,X3,X0,X3,X1) | ~coll(X2,X3,X1) | cyclic(X0,X1,X2,X3)) )),
  inference(cnf_transformation,[],[f225])).
fof(f225,plain,(
  ! [X0,X1,X2,X3] : (cyclic(X0,X1,X2,X3) | ~coll(X2,X3,X1) | ~eqangle(X2,X0,X2,X1,X3,X0,X3,X1))),
  inference(flattening,[],[f224])).
fof(f224,plain,(
  ! [X0,X1,X2,X3] : (cyclic(X0,X1,X2,X3) | (~coll(X2,X3,X1) | ~eqangle(X2,X0,X2,X1,X3,X0,X3,X1)))),
  inference(ennf_transformation,[],[f122])).
fof(f122,plain,(
  ! [X0,X1,X2,X3] : ((coll(X2,X3,X1) & eqangle(X2,X0,X2,X1,X3,X0,X3,X1)) => cyclic(X0,X1,X2,X3))),
  inference(rectify,[],[f43])).
fof(f43,axiom,(
  ! [X0,X1,X8,X9] : ((coll(X8,X9,X1) & eqangle(X8,X0,X8,X1,X9,X0,X9,X1)) => cyclic(X0,X1,X8,X9))),
  file('./Other/gao/L/L022-1.p',ruleD42b)).
fof(f7597,plain,(
  ( ! [X12,X10,X11,X9] : (cyclic(X9,X10,X11,X12) | ~cyclic(X9,X9,X10,X11)) ) | ~spl26_88),
  inference(resolution,[],[f6834,f376])).
fof(f376,plain,(
  ( ! [X4,X2,X0,X3,X1] : (~cyclic(X0,X1,X2,X4) | cyclic(X1,X2,X3,X4) | ~cyclic(X0,X1,X2,X3)) )),
  inference(cnf_transformation,[],[f192])).
fof(f192,plain,(
  ! [X0,X1,X2,X3,X4] : (cyclic(X1,X2,X3,X4) | ~cyclic(X0,X1,X2,X4) | ~cyclic(X0,X1,X2,X3))),
  inference(flattening,[],[f191])).
fof(f191,plain,(
  ! [X0,X1,X2,X3,X4] : (cyclic(X1,X2,X3,X4) | (~cyclic(X0,X1,X2,X4) | ~cyclic(X0,X1,X2,X3)))),
  inference(ennf_transformation,[],[f17])).
fof(f17,axiom,(
  ! [X0,X1,X2,X3,X4] : ((cyclic(X0,X1,X2,X4) & cyclic(X0,X1,X2,X3)) => cyclic(X1,X2,X3,X4))),
  file('./Other/gao/L/L022-1.p',ruleD17)).
fof(f8389,plain,(
  ( ! [X85,X83,X84] : (cong(X83,X84,X83,X84) | ~cyclic(X83,X84,X85,X84)) ) | ~spl26_88),
  inference(subsumption_resolution,[],[f8388,f7619])).
fof(f8388,plain,(
  ( ! [X85,X83,X84] : (cong(X83,X84,X83,X84) | ~cyclic(X83,X84,X85,X83) | ~cyclic(X83,X84,X85,X84)) ) | ~spl26_88),
  inference(subsumption_resolution,[],[f8377,f7595])).
fof(f7595,plain,(
  ( ! [X4,X5,X3] : (cyclic(X3,X4,X3,X5)) ) | ~spl26_88),
  inference(resolution,[],[f6834,f374])).
fof(f374,plain,(
  ( ! [X2,X0,X3,X1] : (~cyclic(X0,X1,X2,X3) | cyclic(X0,X2,X1,X3)) )),
  inference(cnf_transformation,[],[f189])).
fof(f189,plain,(
  ! [X0,X1,X2,X3] : (cyclic(X0,X2,X1,X3) | ~cyclic(X0,X1,X2,X3))),
  inference(ennf_transformation,[],[f15])).
fof(f15,axiom,(
  ! [X0,X1,X2,X3] : (cyclic(X0,X1,X2,X3) => cyclic(X0,X2,X1,X3))),
  file('./Other/gao/L/L022-1.p',ruleD15)).
fof(f8377,plain,(
  ( ! [X85,X83,X84] : (~cyclic(X83,X84,X83,X85) | cong(X83,X84,X83,X84) | ~cyclic(X83,X84,X85,X83) | ~cyclic(X83,X84,X85,X84)) )),
  inference(duplicate_literal_removal,[],[f8367])).
fof(f8367,plain,(
  ( ! [X85,X83,X84] : (~cyclic(X83,X84,X83,X85) | cong(X83,X84,X83,X84) | ~cyclic(X83,X84,X85,X83) | ~cyclic(X83,X84,X85,X84) | ~cyclic(X83,X84,X85,X83)) )),
  inference(resolution,[],[f1643,f403])).
fof(f403,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (~eqangle(X2,X0,X2,X1,X5,X3,X5,X4) | cong(X0,X1,X3,X4) | ~cyclic(X0,X1,X2,X5) | ~cyclic(X0,X1,X2,X4) | ~cyclic(X0,X1,X2,X3)) )),
  inference(cnf_transformation,[],[f227])).
fof(f227,plain,(
  ! [X0,X1,X2,X3,X4,X5] : (cong(X0,X1,X3,X4) | ~eqangle(X2,X0,X2,X1,X5,X3,X5,X4) | ~cyclic(X0,X1,X2,X5) | ~cyclic(X0,X1,X2,X4) | ~cyclic(X0,X1,X2,X3))),
  inference(flattening,[],[f226])).
fof(f226,plain,(
  ! [X0,X1,X2,X3,X4,X5] : (cong(X0,X1,X3,X4) | (~eqangle(X2,X0,X2,X1,X5,X3,X5,X4) | ~cyclic(X0,X1,X2,X5) | ~cyclic(X0,X1,X2,X4) | ~cyclic(X0,X1,X2,X3)))),
  inference(ennf_transformation,[],[f123])).
fof(f123,plain,(
  ! [X0,X1,X2,X3,X4,X5] : ((eqangle(X2,X0,X2,X1,X5,X3,X5,X4) & cyclic(X0,X1,X2,X5) & cyclic(X0,X1,X2,X4) & cyclic(X0,X1,X2,X3)) => cong(X0,X1,X3,X4))),
  inference(rectify,[],[f44])).
fof(f44,axiom,(
  ! [X0,X1,X2,X8,X9,X14] : ((eqangle(X2,X0,X2,X1,X14,X8,X14,X9) & cyclic(X0,X1,X2,X14) & cyclic(X0,X1,X2,X9) & cyclic(X0,X1,X2,X8)) => cong(X0,X1,X8,X9))),
  file('./Other/gao/L/L022-1.p',ruleD43)).
fof(f1643,plain,(
  ( ! [X14,X12,X13,X11] : (eqangle(X14,X11,X14,X12,X11,X13,X13,X12) | ~cyclic(X11,X12,X13,X14)) )),
  inference(resolution,[],[f769,f379])).
fof(f379,plain,(
  ( ! [X6,X4,X2,X0,X7,X5,X3,X1] : (~eqangle(X0,X1,X2,X3,X4,X5,X6,X7) | eqangle(X4,X5,X6,X7,X0,X1,X2,X3)) )),
  inference(cnf_transformation,[],[f195])).
fof(f195,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6,X7] : (eqangle(X4,X5,X6,X7,X0,X1,X2,X3) | ~eqangle(X0,X1,X2,X3,X4,X5,X6,X7))),
  inference(ennf_transformation,[],[f102])).
fof(f102,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6,X7] : (eqangle(X0,X1,X2,X3,X4,X5,X6,X7) => eqangle(X4,X5,X6,X7,X0,X1,X2,X3))),
  inference(rectify,[],[f20])).
fof(f20,axiom,(
  ! [X0,X1,X2,X3,X8,X9,X10,X11] : (eqangle(X0,X1,X2,X3,X8,X9,X10,X11) => eqangle(X8,X9,X10,X11,X0,X1,X2,X3))),
  file('./Other/gao/L/L022-1.p',ruleD20)).
fof(f769,plain,(
  ( ! [X6,X8,X7,X9] : (eqangle(X6,X7,X7,X8,X9,X6,X9,X8) | ~cyclic(X6,X8,X7,X9)) )),
  inference(resolution,[],[f377,f400])).
fof(f400,plain,(
  ( ! [X2,X0,X3,X1] : (eqangle(X2,X0,X2,X1,X3,X0,X3,X1) | ~cyclic(X0,X1,X2,X3)) )),
  inference(cnf_transformation,[],[f221])).
fof(f221,plain,(
  ! [X0,X1,X2,X3] : (eqangle(X2,X0,X2,X1,X3,X0,X3,X1) | ~cyclic(X0,X1,X2,X3))),
  inference(ennf_transformation,[],[f120])).
fof(f120,plain,(
  ! [X0,X1,X2,X3] : (cyclic(X0,X1,X2,X3) => eqangle(X2,X0,X2,X1,X3,X0,X3,X1))),
  inference(rectify,[],[f41])).
fof(f41,axiom,(
  ! [X0,X1,X8,X9] : (cyclic(X0,X1,X8,X9) => eqangle(X8,X0,X8,X1,X9,X0,X9,X1))),
  file('./Other/gao/L/L022-1.p',ruleD41)).
fof(f377,plain,(
  ( ! [X6,X4,X2,X0,X7,X5,X3,X1] : (~eqangle(X0,X1,X2,X3,X4,X5,X6,X7) | eqangle(X1,X0,X2,X3,X4,X5,X6,X7)) )),
  inference(cnf_transformation,[],[f193])).
fof(f193,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6,X7] : (eqangle(X1,X0,X2,X3,X4,X5,X6,X7) | ~eqangle(X0,X1,X2,X3,X4,X5,X6,X7))),
  inference(ennf_transformation,[],[f100])).
fof(f100,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6,X7] : (eqangle(X0,X1,X2,X3,X4,X5,X6,X7) => eqangle(X1,X0,X2,X3,X4,X5,X6,X7))),
  inference(rectify,[],[f18])).
fof(f18,axiom,(
  ! [X0,X1,X2,X3,X8,X9,X10,X11] : (eqangle(X0,X1,X2,X3,X8,X9,X10,X11) => eqangle(X1,X0,X2,X3,X8,X9,X10,X11))),
  file('./Other/gao/L/L022-1.p',ruleD18)).
fof(f7288,plain,(
  ~midp(sK21,sK20,sK23) | ~spl26_295),
  inference(avatar_component_clause,[],[f7287])).
fof(f7287,plain,(
  spl26_295 <=> ~midp(sK21,sK20,sK23)),
  introduced(avatar_definition,[new_symbols(naming,[spl26_295])])).
fof(f12058,plain,(
  ~spl26_54 | ~spl26_88 | spl26_293),
  inference(avatar_contradiction_clause,[],[f12057])).
fof(f12057,plain,(
  $false | (~spl26_54 | ~spl26_88 | ~spl26_293)),
  inference(resolution,[],[f11930,f7282])).
fof(f7282,plain,(
  ~midp(sK20,sK24,sK21) | ~spl26_293),
  inference(avatar_component_clause,[],[f7281])).
fof(f7281,plain,(
  spl26_293 <=> ~midp(sK20,sK24,sK21)),
  introduced(avatar_definition,[new_symbols(naming,[spl26_293])])).
fof(f11930,plain,(
  ( ! [X4] : (midp(X4,sK24,sK21)) ) | (~spl26_54 | ~spl26_88)),
  inference(subsumption_resolution,[],[f11929,f8433])).
fof(f11929,plain,(
  ( ! [X4] : (midp(X4,sK24,sK21) | ~midp(X4,sK24,sK24)) ) | (~spl26_54 | ~spl26_88)),
  inference(subsumption_resolution,[],[f11917,f9034])).
fof(f9034,plain,(
  para(sK24,sK24,sK24,sK21) | (~spl26_54 | ~spl26_88)),
  inference(subsumption_resolution,[],[f9016,f1721])).
fof(f9016,plain,(
  para(sK24,sK24,sK24,sK21) | ~para(sK24,sK21,sK24,sK21) | ~spl26_54),
  inference(resolution,[],[f1420,f941])).
fof(f941,plain,(
  ( ! [X21,X19,X20] : (~cong(X21,X19,X21,X20) | para(X21,X19,X19,X20) | ~para(X19,X20,X21,X20)) )),
  inference(resolution,[],[f433,f406])).
fof(f406,plain,(
  ( ! [X2,X0,X1] : (eqangle(X2,X0,X0,X1,X0,X1,X2,X1) | ~cong(X2,X0,X2,X1)) )),
  inference(cnf_transformation,[],[f232])).
fof(f232,plain,(
  ! [X0,X1,X2] : (eqangle(X2,X0,X0,X1,X0,X1,X2,X1) | ~cong(X2,X0,X2,X1))),
  inference(ennf_transformation,[],[f126])).
fof(f126,plain,(
  ! [X0,X1,X2] : (cong(X2,X0,X2,X1) => eqangle(X2,X0,X0,X1,X0,X1,X2,X1))),
  inference(rectify,[],[f47])).
fof(f47,axiom,(
  ! [X0,X1,X7] : (cong(X7,X0,X7,X1) => eqangle(X7,X0,X0,X1,X0,X1,X7,X1))),
  file('./Other/gao/L/L022-1.p',ruleD46)).
fof(f433,plain,(
  ( ! [X6,X4,X2,X0,X7,X5,X3,X1] : (~eqangle(X0,X1,X2,X3,X4,X5,X6,X7) | ~para(X4,X5,X6,X7) | para(X0,X1,X2,X3)) )),
  inference(cnf_transformation,[],[f280])).
fof(f280,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6,X7] : (para(X0,X1,X2,X3) | ~para(X4,X5,X6,X7) | ~eqangle(X0,X1,X2,X3,X4,X5,X6,X7))),
  inference(flattening,[],[f279])).
fof(f279,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6,X7] : (para(X0,X1,X2,X3) | (~para(X4,X5,X6,X7) | ~eqangle(X0,X1,X2,X3,X4,X5,X6,X7)))),
  inference(ennf_transformation,[],[f146])).
fof(f146,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6,X7] : ((para(X4,X5,X6,X7) & eqangle(X0,X1,X2,X3,X4,X5,X6,X7)) => para(X0,X1,X2,X3))),
  inference(rectify,[],[f74])).
fof(f74,axiom,(
  ! [X0,X1,X2,X3,X8,X9,X10,X11] : ((para(X8,X9,X10,X11) & eqangle(X0,X1,X2,X3,X8,X9,X10,X11)) => para(X0,X1,X2,X3))),
  file('./Other/gao/L/L022-1.p',ruleD73)).
fof(f1420,plain,(
  cong(sK24,sK24,sK24,sK21) | ~spl26_54),
  inference(avatar_component_clause,[],[f1419])).
fof(f1419,plain,(
  spl26_54 <=> cong(sK24,sK24,sK24,sK21)),
  introduced(avatar_definition,[new_symbols(naming,[spl26_54])])).
fof(f11917,plain,(
  ( ! [X4] : (midp(X4,sK24,sK21) | ~para(sK24,sK24,sK24,sK21) | ~midp(X4,sK24,sK24)) ) | ~spl26_54),
  inference(resolution,[],[f9030,f424])).
fof(f424,plain,(
  ( ! [X4,X2,X0,X3,X1] : (~para(X0,X3,X1,X2) | midp(X4,X2,X3) | ~para(X0,X2,X1,X3) | ~midp(X4,X0,X1)) )),
  inference(cnf_transformation,[],[f265])).
fof(f265,plain,(
  ! [X0,X1,X2,X3,X4] : (midp(X4,X2,X3) | ~para(X0,X3,X1,X2) | ~para(X0,X2,X1,X3) | ~midp(X4,X0,X1))),
  inference(flattening,[],[f264])).
fof(f264,plain,(
  ! [X0,X1,X2,X3,X4] : (midp(X4,X2,X3) | (~para(X0,X3,X1,X2) | ~para(X0,X2,X1,X3) | ~midp(X4,X0,X1)))),
  inference(ennf_transformation,[],[f143])).
fof(f143,plain,(
  ! [X0,X1,X2,X3,X4] : ((para(X0,X3,X1,X2) & para(X0,X2,X1,X3) & midp(X4,X0,X1)) => midp(X4,X2,X3))),
  inference(rectify,[],[f65])).
fof(f65,axiom,(
  ! [X0,X1,X2,X3,X6] : ((para(X0,X3,X1,X2) & para(X0,X2,X1,X3) & midp(X6,X0,X1)) => midp(X6,X2,X3))),
  file('./Other/gao/L/L022-1.p',ruleD64)).
fof(f9030,plain,(
  para(sK24,sK21,sK24,sK24) | ~spl26_54),
  inference(resolution,[],[f1420,f1840])).
fof(f1840,plain,(
  ( ! [X0,X1] : (~cong(X0,X0,X0,X1) | para(X0,X1,X0,X0)) )),
  inference(resolution,[],[f784,f398])).
fof(f398,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (~eqangle(X0,X1,X4,X5,X2,X3,X4,X5) | para(X0,X1,X2,X3)) )),
  inference(cnf_transformation,[],[f219])).
fof(f219,plain,(
  ! [X0,X1,X2,X3,X4,X5] : (para(X0,X1,X2,X3) | ~eqangle(X0,X1,X4,X5,X2,X3,X4,X5))),
  inference(ennf_transformation,[],[f118])).
fof(f118,plain,(
  ! [X0,X1,X2,X3,X4,X5] : (eqangle(X0,X1,X4,X5,X2,X3,X4,X5) => para(X0,X1,X2,X3))),
  inference(rectify,[],[f39])).
fof(f39,axiom,(
  ! [X0,X1,X2,X3,X8,X9] : (eqangle(X0,X1,X8,X9,X2,X3,X8,X9) => para(X0,X1,X2,X3))),
  file('./Other/gao/L/L022-1.p',ruleD39)).
fof(f784,plain,(
  ( ! [X12,X10,X11] : (eqangle(X10,X11,X12,X11,X12,X10,X10,X11) | ~cong(X12,X10,X12,X11)) )),
  inference(resolution,[],[f379,f406])).
fof(f8658,plain,(
  spl26_55 | ~spl26_88 | ~spl26_120),
  inference(avatar_contradiction_clause,[],[f8657])).
fof(f8657,plain,(
  $false | (~spl26_55 | ~spl26_88 | ~spl26_120)),
  inference(resolution,[],[f8499,f2395])).
fof(f2395,plain,(
  ~midp(sK24,sK21,sK24) | ~spl26_55),
  inference(resolution,[],[f1423,f515])).
fof(f515,plain,(
  ( ! [X2,X0,X1] : (cong(X0,X1,X0,X2) | ~midp(X0,X2,X1)) )),
  inference(resolution,[],[f383,f428])).
fof(f428,plain,(
  ( ! [X2,X0,X1] : (cong(X0,X1,X0,X2) | ~midp(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f271])).
fof(f271,plain,(
  ! [X0,X1,X2] : (cong(X0,X1,X0,X2) | ~midp(X0,X1,X2))),
  inference(ennf_transformation,[],[f69])).
fof(f69,axiom,(
  ! [X0,X1,X2] : (midp(X0,X1,X2) => cong(X0,X1,X0,X2))),
  file('./Other/gao/L/L022-1.p',ruleD68)).
fof(f383,plain,(
  ( ! [X2,X0,X3,X1] : (~cong(X0,X1,X2,X3) | cong(X2,X3,X0,X1)) )),
  inference(cnf_transformation,[],[f200])).
fof(f200,plain,(
  ! [X0,X1,X2,X3] : (cong(X2,X3,X0,X1) | ~cong(X0,X1,X2,X3))),
  inference(ennf_transformation,[],[f24])).
fof(f24,axiom,(
  ! [X0,X1,X2,X3] : (cong(X0,X1,X2,X3) => cong(X2,X3,X0,X1))),
  file('./Other/gao/L/L022-1.p',ruleD24)).
fof(f1423,plain,(
  ~cong(sK24,sK24,sK24,sK21) | ~spl26_55),
  inference(avatar_component_clause,[],[f1422])).
fof(f1422,plain,(
  spl26_55 <=> ~cong(sK24,sK24,sK24,sK21)),
  introduced(avatar_definition,[new_symbols(naming,[spl26_55])])).
fof(f8499,plain,(
  ( ! [X31,X32] : (midp(X31,X32,X31)) ) | (~spl26_88 | ~spl26_120)),
  inference(resolution,[],[f8433,f6282])).
fof(f6282,plain,(
  ( ! [X10,X11,X9] : (~midp(X9,X10,X11) | midp(X10,X9,X11)) ) | (~spl26_88 | ~spl26_120)),
  inference(subsumption_resolution,[],[f6245,f2601])).
fof(f2601,plain,(
  ( ! [X2,X0,X1] : (coll(X0,X1,X2)) ) | ~spl26_88),
  inference(subsumption_resolution,[],[f2597,f2582])).
fof(f2582,plain,(
  ( ! [X6,X5] : (coll(X5,X5,X6)) ) | ~spl26_88),
  inference(resolution,[],[f2564,f360])).
fof(f360,plain,(
  ( ! [X2,X0,X1] : (~coll(X0,X1,X2) | coll(X0,X2,X1)) )),
  inference(cnf_transformation,[],[f169])).
fof(f169,plain,(
  ! [X0,X1,X2] : (coll(X0,X2,X1) | ~coll(X0,X1,X2))),
  inference(ennf_transformation,[],[f1])).
fof(f1,axiom,(
  ! [X0,X1,X2] : (coll(X0,X1,X2) => coll(X0,X2,X1))),
  file('./Other/gao/L/L022-1.p',ruleD1)).
fof(f2564,plain,(
  ( ! [X4,X3] : (coll(X3,X4,X3)) ) | ~spl26_88),
  inference(resolution,[],[f2550,f361])).
fof(f361,plain,(
  ( ! [X2,X0,X1] : (~coll(X0,X1,X2) | coll(X1,X0,X2)) )),
  inference(cnf_transformation,[],[f170])).
fof(f170,plain,(
  ! [X0,X1,X2] : (coll(X1,X0,X2) | ~coll(X0,X1,X2))),
  inference(ennf_transformation,[],[f2])).
fof(f2,axiom,(
  ! [X0,X1,X2] : (coll(X0,X1,X2) => coll(X1,X0,X2))),
  file('./Other/gao/L/L022-1.p',ruleD2)).
fof(f2597,plain,(
  ( ! [X2,X0,X1] : (coll(X0,X1,X2) | ~coll(X2,X2,X0)) ) | ~spl26_88),
  inference(resolution,[],[f2582,f362])).
fof(f362,plain,(
  ( ! [X2,X0,X3,X1] : (~coll(X0,X1,X3) | coll(X2,X3,X0) | ~coll(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f172])).
fof(f172,plain,(
  ! [X0,X1,X2,X3] : (coll(X2,X3,X0) | ~coll(X0,X1,X3) | ~coll(X0,X1,X2))),
  inference(flattening,[],[f171])).
fof(f171,plain,(
  ! [X0,X1,X2,X3] : (coll(X2,X3,X0) | (~coll(X0,X1,X3) | ~coll(X0,X1,X2)))),
  inference(ennf_transformation,[],[f3])).
fof(f3,axiom,(
  ! [X0,X1,X2,X3] : ((coll(X0,X1,X3) & coll(X0,X1,X2)) => coll(X2,X3,X0))),
  file('./Other/gao/L/L022-1.p',ruleD3)).
fof(f6245,plain,(
  ( ! [X10,X11,X9] : (~midp(X9,X10,X11) | ~coll(X10,X9,X11) | midp(X10,X9,X11)) ) | ~spl26_120),
  inference(resolution,[],[f2305,f427])).
fof(f2305,plain,(
  ( ! [X14,X15,X13] : (cong(X14,X13,X14,X15) | ~midp(X13,X14,X15)) ) | ~spl26_120),
  inference(avatar_component_clause,[],[f2304])).
fof(f2304,plain,(
  spl26_120 <=> ! [X13,X15,X14] : (~midp(X13,X14,X15) | cong(X14,X13,X14,X15))),
  introduced(avatar_definition,[new_symbols(naming,[spl26_120])])).
fof(f7289,plain,(
  ~spl26_293 | ~spl26_295),
  inference(avatar_split_clause,[],[f7269,f7287,f7281])).
fof(f7269,plain,(
  ~midp(sK21,sK20,sK23) | ~midp(sK20,sK24,sK21)),
  inference(resolution,[],[f1589,f477])).
fof(f477,plain,(
  ~cong(sK20,sK24,sK21,sK23)),
  inference(cnf_transformation,[],[f359])).
fof(f359,plain,(
  ~cong(sK20,sK24,sK21,sK23) & coll(sK24,sK21,sK22) & cong(sK24,sK22,sK22,sK23) & coll(sK23,sK20,sK22) & perp(sK20,sK21,sK25,sK22)),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK20,sK21,sK22,sK23,sK24,sK25])],[f321,f358])).
fof(f358,plain,(
  ? [X0,X1,X2,X3,X4,X5] : (~cong(X0,X4,X1,X3) & coll(X4,X1,X2) & cong(X4,X2,X2,X3) & coll(X3,X0,X2) & perp(X0,X1,X5,X2)) => (~cong(sK20,sK24,sK21,sK23) & coll(sK24,sK21,sK22) & cong(sK24,sK22,sK22,sK23) & coll(sK23,sK20,sK22) & perp(sK20,sK21,sK25,sK22))),
  introduced(choice_axiom,[])).
fof(f321,plain,(
  ? [X0,X1,X2,X3,X4,X5] : (~cong(X0,X4,X1,X3) & coll(X4,X1,X2) & cong(X4,X2,X2,X3) & coll(X3,X0,X2) & perp(X0,X1,X5,X2))),
  inference(flattening,[],[f320])).
fof(f320,plain,(
  ? [X0,X1,X2,X3,X4,X5] : (~cong(X0,X4,X1,X3) & (coll(X4,X1,X2) & cong(X4,X2,X2,X3) & coll(X3,X0,X2) & perp(X0,X1,X5,X2)))),
  inference(ennf_transformation,[],[f168])).
fof(f168,plain,(
  ~! [X0,X1,X2,X3,X4,X5] : ((coll(X4,X1,X2) & cong(X4,X2,X2,X3) & coll(X3,X0,X2) & perp(X0,X1,X5,X2)) => cong(X0,X4,X1,X3))),
  inference(pure_predicate_removal,[],[f167])).
fof(f167,plain,(
  ~! [X0,X1,X2,X3,X4,X5] : ((coll(X4,X1,X2) & cong(X4,X2,X2,X3) & coll(X3,X0,X2) & perp(X0,X1,X5,X2) & midpoint(X5,X0,X1)) => cong(X0,X4,X1,X3))),
  inference(rectify,[],[f96])).
fof(f96,negated_conjecture,(
  ~! [X0,X1,X2,X3,X4,X18] : ((coll(X4,X1,X2) & cong(X4,X2,X2,X3) & coll(X3,X0,X2) & perp(X0,X1,X18,X2) & midpoint(X18,X0,X1)) => cong(X0,X4,X1,X3))),
  inference(negated_conjecture,[],[f95])).
fof(f95,conjecture,(
  ! [X0,X1,X2,X3,X4,X18] : ((coll(X4,X1,X2) & cong(X4,X2,X2,X3) & coll(X3,X0,X2) & perp(X0,X1,X18,X2) & midpoint(X18,X0,X1)) => cong(X0,X4,X1,X3))),
  file('./Other/gao/L/L022-1.p',exemploOthergaoLL0221)).
fof(f1589,plain,(
  ( ! [X14,X12,X15,X13] : (cong(X12,X13,X14,X15) | ~midp(X14,X12,X15) | ~midp(X12,X13,X14)) )),
  inference(resolution,[],[f734,f511])).
fof(f511,plain,(
  ( ! [X2,X0,X1] : (cong(X0,X1,X2,X0) | ~midp(X0,X1,X2)) )),
  inference(resolution,[],[f382,f428])).
fof(f382,plain,(
  ( ! [X2,X0,X3,X1] : (~cong(X0,X1,X2,X3) | cong(X0,X1,X3,X2)) )),
  inference(cnf_transformation,[],[f199])).
fof(f199,plain,(
  ! [X0,X1,X2,X3] : (cong(X0,X1,X3,X2) | ~cong(X0,X1,X2,X3))),
  inference(ennf_transformation,[],[f23])).
fof(f23,axiom,(
  ! [X0,X1,X2,X3] : (cong(X0,X1,X2,X3) => cong(X0,X1,X3,X2))),
  file('./Other/gao/L/L022-1.p',ruleD23)).
fof(f734,plain,(
  ( ! [X4,X2,X0,X3,X1] : (~cong(X0,X1,X2,X4) | cong(X0,X1,X2,X3) | ~midp(X2,X4,X3)) )),
  inference(resolution,[],[f384,f428])).
fof(f384,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (~cong(X2,X3,X4,X5) | cong(X0,X1,X4,X5) | ~cong(X0,X1,X2,X3)) )),
  inference(cnf_transformation,[],[f202])).
fof(f202,plain,(
  ! [X0,X1,X2,X3,X4,X5] : (cong(X0,X1,X4,X5) | ~cong(X2,X3,X4,X5) | ~cong(X0,X1,X2,X3))),
  inference(flattening,[],[f201])).
fof(f201,plain,(
  ! [X0,X1,X2,X3,X4,X5] : (cong(X0,X1,X4,X5) | (~cong(X2,X3,X4,X5) | ~cong(X0,X1,X2,X3)))),
  inference(ennf_transformation,[],[f25])).
fof(f25,axiom,(
  ! [X0,X1,X2,X3,X4,X5] : ((cong(X2,X3,X4,X5) & cong(X0,X1,X2,X3)) => cong(X0,X1,X4,X5))),
  file('./Other/gao/L/L022-1.p',ruleD25)).
fof(f4416,plain,(
  spl26_128 | ~spl26_78 | ~spl26_88),
  inference(avatar_split_clause,[],[f4410,f1720,f1497,f2726])).
fof(f2726,plain,(
  spl26_128 <=> para(sK23,sK22,sK22,sK22)),
  introduced(avatar_definition,[new_symbols(naming,[spl26_128])])).
fof(f1497,plain,(
  spl26_78 <=> cong(sK22,sK22,sK22,sK23)),
  introduced(avatar_definition,[new_symbols(naming,[spl26_78])])).
fof(f4410,plain,(
  para(sK23,sK22,sK22,sK22) | (~spl26_78 | ~spl26_88)),
  inference(resolution,[],[f4051,f364])).
fof(f364,plain,(
  ( ! [X2,X0,X3,X1] : (~para(X0,X1,X2,X3) | para(X2,X3,X0,X1)) )),
  inference(cnf_transformation,[],[f174])).
fof(f174,plain,(
  ! [X0,X1,X2,X3] : (para(X2,X3,X0,X1) | ~para(X0,X1,X2,X3))),
  inference(ennf_transformation,[],[f5])).
fof(f5,axiom,(
  ! [X0,X1,X2,X3] : (para(X0,X1,X2,X3) => para(X2,X3,X0,X1))),
  file('./Other/gao/L/L022-1.p',ruleD5)).
fof(f4051,plain,(
  para(sK22,sK22,sK23,sK22) | (~spl26_78 | ~spl26_88)),
  inference(resolution,[],[f3859,f363])).
fof(f363,plain,(
  ( ! [X2,X0,X3,X1] : (~para(X0,X1,X2,X3) | para(X0,X1,X3,X2)) )),
  inference(cnf_transformation,[],[f173])).
fof(f173,plain,(
  ! [X0,X1,X2,X3] : (para(X0,X1,X3,X2) | ~para(X0,X1,X2,X3))),
  inference(ennf_transformation,[],[f4])).
fof(f4,axiom,(
  ! [X0,X1,X2,X3] : (para(X0,X1,X2,X3) => para(X0,X1,X3,X2))),
  file('./Other/gao/L/L022-1.p',ruleD4)).
fof(f3859,plain,(
  para(sK22,sK22,sK22,sK23) | (~spl26_78 | ~spl26_88)),
  inference(subsumption_resolution,[],[f3849,f1721])).
fof(f3849,plain,(
  para(sK22,sK22,sK22,sK23) | ~para(sK22,sK23,sK22,sK23) | ~spl26_78),
  inference(resolution,[],[f1498,f941])).
fof(f1498,plain,(
  cong(sK22,sK22,sK22,sK23) | ~spl26_78),
  inference(avatar_component_clause,[],[f1497])).
fof(f3834,plain,(
  spl26_79 | ~spl26_88),
  inference(avatar_contradiction_clause,[],[f3833])).
fof(f3833,plain,(
  $false | (~spl26_79 | ~spl26_88)),
  inference(subsumption_resolution,[],[f3829,f3530])).
fof(f3530,plain,(
  ( ! [X7] : (midp(X7,X7,sK23)) ) | ~spl26_88),
  inference(resolution,[],[f3294,f370])).
fof(f3294,plain,(
  ( ! [X3] : (midp(X3,sK23,X3)) ) | ~spl26_88),
  inference(resolution,[],[f2781,f3253])).
fof(f3253,plain,(
  ( ! [X26] : (midp(sK23,X26,X26)) ) | ~spl26_88),
  inference(resolution,[],[f2559,f2721])).
fof(f2721,plain,(
  midp(sK23,sK22,sK22)),
  inference(subsumption_resolution,[],[f2705,f584])).
fof(f584,plain,(
  coll(sK23,sK22,sK22)),
  inference(resolution,[],[f581,f361])).
fof(f581,plain,(
  coll(sK22,sK23,sK22)),
  inference(resolution,[],[f577,f360])).
fof(f577,plain,(
  coll(sK22,sK22,sK23)),
  inference(resolution,[],[f524,f474])).
fof(f474,plain,(
  coll(sK23,sK20,sK22)),
  inference(cnf_transformation,[],[f359])).
fof(f524,plain,(
  ( ! [X6] : (~coll(sK23,sK20,X6) | coll(X6,sK22,sK23)) )),
  inference(resolution,[],[f362,f474])).
fof(f2705,plain,(
  ~coll(sK23,sK22,sK22) | midp(sK23,sK22,sK22)),
  inference(resolution,[],[f2392,f427])).
fof(f2392,plain,(
  cong(sK23,sK22,sK23,sK22)),
  inference(resolution,[],[f1344,f382])).
fof(f1344,plain,(
  cong(sK23,sK22,sK22,sK23)),
  inference(resolution,[],[f741,f514])).
fof(f514,plain,(
  cong(sK23,sK22,sK24,sK22)),
  inference(resolution,[],[f383,f510])).
fof(f510,plain,(
  cong(sK24,sK22,sK23,sK22)),
  inference(resolution,[],[f382,f475])).
fof(f475,plain,(
  cong(sK24,sK22,sK22,sK23)),
  inference(cnf_transformation,[],[f359])).
fof(f741,plain,(
  ( ! [X17,X18] : (~cong(X17,X18,sK24,sK22) | cong(X17,X18,sK22,sK23)) )),
  inference(resolution,[],[f384,f475])).
fof(f2559,plain,(
  ( ! [X10,X8,X9] : (~midp(X8,X10,X10) | midp(X8,X9,X9)) ) | ~spl26_88),
  inference(subsumption_resolution,[],[f2553,f1721])).
fof(f2553,plain,(
  ( ! [X10,X8,X9] : (midp(X8,X9,X9) | ~para(X10,X9,X10,X9) | ~midp(X8,X10,X10)) ) | ~spl26_88),
  inference(resolution,[],[f1721,f424])).
fof(f2781,plain,(
  ( ! [X12,X13,X11] : (~midp(X13,X12,X13) | midp(X11,X12,X11)) ) | ~spl26_88),
  inference(subsumption_resolution,[],[f2554,f2564])).
fof(f2554,plain,(
  ( ! [X12,X13,X11] : (~coll(X11,X12,X11) | midp(X11,X12,X11) | ~midp(X13,X12,X13)) ) | ~spl26_88),
  inference(resolution,[],[f1721,f405])).
fof(f405,plain,(
  ( ! [X4,X2,X0,X3,X1] : (~para(X3,X4,X1,X2) | ~coll(X4,X0,X2) | midp(X4,X0,X2) | ~midp(X3,X0,X1)) )),
  inference(cnf_transformation,[],[f231])).
fof(f231,plain,(
  ! [X0,X1,X2,X3,X4] : (midp(X4,X0,X2) | ~coll(X4,X0,X2) | ~para(X3,X4,X1,X2) | ~midp(X3,X0,X1))),
  inference(flattening,[],[f230])).
fof(f230,plain,(
  ! [X0,X1,X2,X3,X4] : (midp(X4,X0,X2) | (~coll(X4,X0,X2) | ~para(X3,X4,X1,X2) | ~midp(X3,X0,X1)))),
  inference(ennf_transformation,[],[f125])).
fof(f125,plain,(
  ! [X0,X1,X2,X3,X4] : ((coll(X4,X0,X2) & para(X3,X4,X1,X2) & midp(X3,X0,X1)) => midp(X4,X0,X2))),
  inference(rectify,[],[f46])).
fof(f46,axiom,(
  ! [X0,X1,X2,X4,X5] : ((coll(X5,X0,X2) & para(X4,X5,X1,X2) & midp(X4,X0,X1)) => midp(X5,X0,X2))),
  file('./Other/gao/L/L022-1.p',ruleD45)).
fof(f3829,plain,(
  ~midp(sK22,sK22,sK23) | ~spl26_79),
  inference(resolution,[],[f1501,f428])).
fof(f1501,plain,(
  ~cong(sK22,sK22,sK22,sK23) | ~spl26_79),
  inference(avatar_component_clause,[],[f1500])).
fof(f1500,plain,(
  spl26_79 <=> ~cong(sK22,sK22,sK22,sK23)),
  introduced(avatar_definition,[new_symbols(naming,[spl26_79])])).
fof(f3564,plain,(
  spl26_75 | ~spl26_88),
  inference(avatar_contradiction_clause,[],[f3563])).
fof(f3563,plain,(
  $false | (~spl26_75 | ~spl26_88)),
  inference(subsumption_resolution,[],[f3559,f3530])).
fof(f3559,plain,(
  ~midp(sK20,sK20,sK23) | ~spl26_75),
  inference(resolution,[],[f1488,f428])).
fof(f1488,plain,(
  ~cong(sK20,sK20,sK20,sK23) | ~spl26_75),
  inference(avatar_component_clause,[],[f1487])).
fof(f1487,plain,(
  spl26_75 <=> ~cong(sK20,sK20,sK20,sK23)),
  introduced(avatar_definition,[new_symbols(naming,[spl26_75])])).
fof(f3506,plain,(
  spl26_140 | ~spl26_74 | ~spl26_88),
  inference(avatar_split_clause,[],[f3493,f1720,f1484,f3496])).
fof(f3496,plain,(
  spl26_140 <=> para(sK20,sK23,sK20,sK20)),
  introduced(avatar_definition,[new_symbols(naming,[spl26_140])])).
fof(f1484,plain,(
  spl26_74 <=> cong(sK20,sK20,sK20,sK23)),
  introduced(avatar_definition,[new_symbols(naming,[spl26_74])])).
fof(f3493,plain,(
  para(sK20,sK23,sK20,sK20) | (~spl26_74 | ~spl26_88)),
  inference(resolution,[],[f3393,f364])).
fof(f3393,plain,(
  para(sK20,sK20,sK20,sK23) | (~spl26_74 | ~spl26_88)),
  inference(subsumption_resolution,[],[f3384,f1721])).
fof(f3384,plain,(
  para(sK20,sK20,sK20,sK23) | ~para(sK20,sK23,sK20,sK23) | ~spl26_74),
  inference(resolution,[],[f1485,f941])).
fof(f1485,plain,(
  cong(sK20,sK20,sK20,sK23) | ~spl26_74),
  inference(avatar_component_clause,[],[f1484])).
fof(f3504,plain,(
  ~spl26_141 | spl26_142 | ~spl26_74 | ~spl26_88),
  inference(avatar_split_clause,[],[f3490,f1720,f1484,f3502,f3499])).
fof(f3499,plain,(
  spl26_141 <=> ~para(sK20,sK23,sK20,sK20)),
  introduced(avatar_definition,[new_symbols(naming,[spl26_141])])).
fof(f3490,plain,(
  ( ! [X1] : (midp(X1,sK23,sK20) | ~para(sK20,sK23,sK20,sK20) | ~midp(X1,sK20,sK20)) ) | (~spl26_74 | ~spl26_88)),
  inference(resolution,[],[f3393,f424])).
fof(f2779,plain,(
  ~spl26_42 | ~spl26_88 | ~spl26_128),
  inference(avatar_contradiction_clause,[],[f2778])).
fof(f2778,plain,(
  $false | (~spl26_42 | ~spl26_88 | ~spl26_128)),
  inference(resolution,[],[f2774,f2721])).
fof(f2774,plain,(
  ( ! [X1] : (~midp(sK23,X1,sK22)) ) | (~spl26_42 | ~spl26_88 | ~spl26_128)),
  inference(subsumption_resolution,[],[f2773,f2332])).
fof(f2332,plain,(
  ( ! [X0,X1] : (~midp(X0,X1,X0)) ) | ~spl26_42),
  inference(subsumption_resolution,[],[f2325,f370])).
fof(f2325,plain,(
  ( ! [X0,X1] : (~midp(X0,X0,X1) | ~midp(X0,X1,X0)) ) | ~spl26_42),
  inference(resolution,[],[f963,f515])).
fof(f963,plain,(
  ( ! [X6,X8,X7] : (~cong(X6,X7,X7,X8) | ~midp(X6,X7,X8)) ) | ~spl26_42),
  inference(avatar_component_clause,[],[f962])).
fof(f962,plain,(
  spl26_42 <=> ! [X7,X8,X6] : (~cong(X6,X7,X7,X8) | ~midp(X6,X7,X8))),
  introduced(avatar_definition,[new_symbols(naming,[spl26_42])])).
fof(f2773,plain,(
  ( ! [X1] : (midp(sK22,X1,sK22) | ~midp(sK23,X1,sK22)) ) | (~spl26_88 | ~spl26_128)),
  inference(subsumption_resolution,[],[f2768,f2564])).
fof(f2768,plain,(
  ( ! [X1] : (~coll(sK22,X1,sK22) | midp(sK22,X1,sK22) | ~midp(sK23,X1,sK22)) ) | ~spl26_128),
  inference(resolution,[],[f2727,f405])).
fof(f2727,plain,(
  para(sK23,sK22,sK22,sK22) | ~spl26_128),
  inference(avatar_component_clause,[],[f2726])).
fof(f2545,plain,(
  ~spl26_90),
  inference(avatar_contradiction_clause,[],[f2544])).
fof(f2544,plain,(
  $false | ~spl26_90),
  inference(resolution,[],[f1724,f1329])).
fof(f1329,plain,(
  para(sK25,sK22,sK25,sK22)),
  inference(resolution,[],[f711,f501])).
fof(f501,plain,(
  perp(sK25,sK22,sK20,sK21)),
  inference(resolution,[],[f367,f473])).
fof(f473,plain,(
  perp(sK20,sK21,sK25,sK22)),
  inference(cnf_transformation,[],[f359])).
fof(f367,plain,(
  ( ! [X2,X0,X3,X1] : (~perp(X0,X1,X2,X3) | perp(X2,X3,X0,X1)) )),
  inference(cnf_transformation,[],[f178])).
fof(f178,plain,(
  ! [X0,X1,X2,X3] : (perp(X2,X3,X0,X1) | ~perp(X0,X1,X2,X3))),
  inference(ennf_transformation,[],[f8])).
fof(f8,axiom,(
  ! [X0,X1,X2,X3] : (perp(X0,X1,X2,X3) => perp(X2,X3,X0,X1))),
  file('./Other/gao/L/L022-1.p',ruleD8)).
fof(f711,plain,(
  ( ! [X6,X7] : (~perp(X6,X7,sK20,sK21) | para(X6,X7,sK25,sK22)) )),
  inference(resolution,[],[f368,f473])).
fof(f368,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (~perp(X2,X3,X4,X5) | para(X0,X1,X4,X5) | ~perp(X0,X1,X2,X3)) )),
  inference(cnf_transformation,[],[f180])).
fof(f180,plain,(
  ! [X0,X1,X2,X3,X4,X5] : (para(X0,X1,X4,X5) | ~perp(X2,X3,X4,X5) | ~perp(X0,X1,X2,X3))),
  inference(flattening,[],[f179])).
fof(f179,plain,(
  ! [X0,X1,X2,X3,X4,X5] : (para(X0,X1,X4,X5) | (~perp(X2,X3,X4,X5) | ~perp(X0,X1,X2,X3)))),
  inference(ennf_transformation,[],[f9])).
fof(f9,axiom,(
  ! [X0,X1,X2,X3,X4,X5] : ((perp(X2,X3,X4,X5) & perp(X0,X1,X2,X3)) => para(X0,X1,X4,X5))),
  file('./Other/gao/L/L022-1.p',ruleD9)).
fof(f1724,plain,(
  ( ! [X0,X1] : (~para(X0,X1,X0,X1)) ) | ~spl26_90),
  inference(avatar_component_clause,[],[f1723])).
fof(f1723,plain,(
  spl26_90 <=> ! [X1,X0] : ~para(X0,X1,X0,X1)),
  introduced(avatar_definition,[new_symbols(naming,[spl26_90])])).
fof(f2310,plain,(
  spl26_120 | spl26_42),
  inference(avatar_split_clause,[],[f2297,f962,f2304])).
fof(f2297,plain,(
  ( ! [X14,X12,X10,X15,X13,X11] : (~midp(X10,X11,X12) | ~midp(X13,X14,X15) | ~cong(X10,X11,X11,X12) | cong(X14,X13,X14,X15)) )),
  inference(resolution,[],[f798,f435])).
fof(f435,plain,(
  ( ! [X6,X4,X2,X0,X7,X5,X3,X1] : (~eqratio(X0,X1,X2,X3,X4,X5,X6,X7) | ~cong(X4,X5,X6,X7) | cong(X0,X1,X2,X3)) )),
  inference(cnf_transformation,[],[f284])).
fof(f284,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6,X7] : (cong(X0,X1,X2,X3) | ~cong(X4,X5,X6,X7) | ~eqratio(X0,X1,X2,X3,X4,X5,X6,X7))),
  inference(flattening,[],[f283])).
fof(f283,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6,X7] : (cong(X0,X1,X2,X3) | (~cong(X4,X5,X6,X7) | ~eqratio(X0,X1,X2,X3,X4,X5,X6,X7)))),
  inference(ennf_transformation,[],[f148])).
fof(f148,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6,X7] : ((cong(X4,X5,X6,X7) & eqratio(X0,X1,X2,X3,X4,X5,X6,X7)) => cong(X0,X1,X2,X3))),
  inference(rectify,[],[f76])).
fof(f76,axiom,(
  ! [X0,X1,X2,X3,X8,X9,X10,X11] : ((cong(X8,X9,X10,X11) & eqratio(X0,X1,X2,X3,X8,X9,X10,X11)) => cong(X0,X1,X2,X3))),
  file('./Other/gao/L/L022-1.p',ruleD75)).
fof(f798,plain,(
  ( ! [X6,X10,X8,X7,X11,X9] : (eqratio(X6,X7,X6,X8,X9,X10,X10,X11) | ~midp(X9,X10,X11) | ~midp(X7,X6,X8)) )),
  inference(resolution,[],[f385,f430])).
fof(f430,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (eqratio(X4,X0,X0,X1,X5,X2,X2,X3) | ~midp(X5,X2,X3) | ~midp(X4,X0,X1)) )),
  inference(cnf_transformation,[],[f274])).
fof(f274,plain,(
  ! [X0,X1,X2,X3,X4,X5] : (eqratio(X4,X0,X0,X1,X5,X2,X2,X3) | ~midp(X5,X2,X3) | ~midp(X4,X0,X1))),
  inference(flattening,[],[f273])).
fof(f273,plain,(
  ! [X0,X1,X2,X3,X4,X5] : (eqratio(X4,X0,X0,X1,X5,X2,X2,X3) | (~midp(X5,X2,X3) | ~midp(X4,X0,X1)))),
  inference(ennf_transformation,[],[f145])).
fof(f145,plain,(
  ! [X0,X1,X2,X3,X4,X5] : ((midp(X5,X2,X3) & midp(X4,X0,X1)) => eqratio(X4,X0,X0,X1,X5,X2,X2,X3))),
  inference(rectify,[],[f71])).
fof(f71,axiom,(
  ! [X0,X1,X2,X3,X6,X16] : ((midp(X16,X2,X3) & midp(X6,X0,X1)) => eqratio(X6,X0,X0,X1,X16,X2,X2,X3))),
  file('./Other/gao/L/L022-1.p',ruleD70)).
fof(f385,plain,(
  ( ! [X6,X4,X2,X0,X7,X5,X3,X1] : (~eqratio(X0,X1,X2,X3,X4,X5,X6,X7) | eqratio(X1,X0,X2,X3,X4,X5,X6,X7)) )),
  inference(cnf_transformation,[],[f203])).
fof(f203,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6,X7] : (eqratio(X1,X0,X2,X3,X4,X5,X6,X7) | ~eqratio(X0,X1,X2,X3,X4,X5,X6,X7))),
  inference(ennf_transformation,[],[f105])).
fof(f105,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6,X7] : (eqratio(X0,X1,X2,X3,X4,X5,X6,X7) => eqratio(X1,X0,X2,X3,X4,X5,X6,X7))),
  inference(rectify,[],[f26])).
fof(f26,axiom,(
  ! [X0,X1,X2,X3,X8,X9,X10,X11] : (eqratio(X0,X1,X2,X3,X8,X9,X10,X11) => eqratio(X1,X0,X2,X3,X8,X9,X10,X11))),
  file('./Other/gao/L/L022-1.p',ruleD26)).
fof(f1725,plain,(
  spl26_88 | spl26_90),
  inference(avatar_split_clause,[],[f1696,f1723,f1720])).
fof(f1696,plain,(
  ( ! [X2,X0,X3,X1] : (~para(X0,X1,X0,X1) | para(X2,X3,X2,X3)) )),
  inference(resolution,[],[f775,f398])).
fof(f775,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (eqangle(X0,X1,X2,X3,X0,X1,X4,X5) | ~para(X2,X3,X4,X5)) )),
  inference(resolution,[],[f378,f399])).
fof(f378,plain,(
  ( ! [X6,X4,X2,X0,X7,X5,X3,X1] : (~eqangle(X0,X1,X2,X3,X4,X5,X6,X7) | eqangle(X2,X3,X0,X1,X6,X7,X4,X5)) )),
  inference(cnf_transformation,[],[f194])).
fof(f194,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6,X7] : (eqangle(X2,X3,X0,X1,X6,X7,X4,X5) | ~eqangle(X0,X1,X2,X3,X4,X5,X6,X7))),
  inference(ennf_transformation,[],[f101])).
fof(f101,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6,X7] : (eqangle(X0,X1,X2,X3,X4,X5,X6,X7) => eqangle(X2,X3,X0,X1,X6,X7,X4,X5))),
  inference(rectify,[],[f19])).
fof(f19,axiom,(
  ! [X0,X1,X2,X3,X8,X9,X10,X11] : (eqangle(X0,X1,X2,X3,X8,X9,X10,X11) => eqangle(X2,X3,X0,X1,X10,X11,X8,X9))),
  file('./Other/gao/L/L022-1.p',ruleD19)).
% SZS output end Proof for L022-1
% ------------------------------
% Version: Vampire 4.4.0 (commit 7916d27 on 2019-08-23 08:50:16 +0100)
% Termination reason: Refutation

% Memory used [KB]: 11129
% Time elapsed: 0.239 s
% ------------------------------
% ------------------------------
